<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2023-03-07 Chebychev-Ungleichung</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <div id="oer-aufgabe-19-info">
        <details>
            <summary>
                Inhalt und Umfang dieser Aufgabe
            </summary>
            <p>
                In dieser Aufgabe lernen Sie, eine Wahrscheinlichkeit mit der Chebychev-Ungleichung nach unten abzusch&#228;tzen. Sie besteht aus den folgenden vier Teilaufgaben.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    bestimmen Sie die Verteilung der Zufallsvariablen \(X\).
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    berechnen Sie den Erwartungswert und die Varianz von \(X\).
                </li>
                <li>
                    In
                    <b>
                        (c)
                    </b>
                    bringen Sie die gesuchte Wahrscheinlichkeit auf die Form der Chebychev-Ungleichung.
                </li>
                <li>
                    In
                    <b>
                        (d)
                    </b>
                    sch&#228;tzen Sie die gesuchte Wahrscheinlichkeit mit der Chebychev-Ungleichung nach unten ab.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-19">
        <p>
            Saatk&#246;rner einer bestimmten Sorte keimen unabh&#228;ngig voneinander mit Wahrscheinlichkeit \(p={@dp(p)@}\). Insgesamt werden \(n={@n@}\) Saatk&#246;rner ges&#228;t. Die zuf&#228;llige Anzahl der davon gekeimten Saatk&#246;rner wird durch eine Zufallsvariable \(X\) beschrieben.
        </p>
        <hr/>
    </div>
    <div id="oer-aufgabe-19a">
        <p>
            <b>
                (a)
            </b>
            Welche Verteilung besitzt die Zufallsvariable \(X\)?
        </p>
        <p>
            [[input:ansA]] [[validation:ansA]] [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-19b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            Die Zufallsvariable \(X\) hat eine \(\text{Bin}({@n@},{@dp(p)@})\)-Verteilung. Berechnen Sie den Erwartungswert und die Varianz von \(X\).
        </p>
        <p>
            Der Erwartungswert von \(X\) ist \(\text{E}(X)=\) [[input:ansB1]]. [[validation:ansB1]] [[feedback:prtB1]]
        </p>
        <p>
            Die Varianz von \(X\) ist \(\text{Var}(X)=\) [[input:ansB2]]. [[validation:ansB2]] [[feedback:prtB2]]
        </p>
        <p>
            [[feedback:prtB]]
        </p>
    </div>
    <div id="oer-aufgabe-19c" style="display:none;">
        <p>
            <b>
                (c)
            </b>
            Schreiben Sie die Wahrscheinlichkeit \(P({@a@} &lt; X &lt; {@b@})\) Schritt f&#252;r Schritt so um, dass die Chebychev-Ungleichung anwendbar wird: Subtrahieren Sie zuerst auf allen Seiten der Ungleichungskette, die das Ereignis beschreibt, den Erwartungswert von \(X\), den Sie in der vorherigen Teilaufgabe berechnet haben. Fassen Sie anschlie&#223;end die Ungleichungskette zu einer Betragsungleichung zusammen. Gehen Sie zuletzt zur Gegenwahrscheinlichkeit &#252;ber, damit das Ereignis durch ein \(\geq\)-Zeichen statt ein \(&lt;\)-Zeichen beschrieben wird.
        </p>
        <p>
            <i>
                Beim Klick auf "Pr&#252;fen" bekommen Sie erst dann ein Feedback, wenn Sie alle Eingabefelder ausgef&#252;llt haben.
            </i>
        </p>
        <p>
            <br/>
            \(P({@a@} &lt; X &lt; {@b@})\)
            <br/>
            \(=P(\) [[input:ansC1]] \(&lt; X-\text{E}(X) &lt;\) [[input:ansC2]] \()\)
            <br/>
            \(=P(|X-\text{E}(X)| &lt;\) [[input:ansC3]] \()\)
            <br/>
            \(=1-P(|X-\text{E}(X)|\geq\) [[input:ansC4]] \()\)
[[validation:ansC1]] [[validation:ansC2]] [[validation:ansC3]] [[validation:ansC4]] [[feedback:prtC]]
        </p>
    </div>
    <div id="oer-aufgabe-19d" style="display:none;">
        <p>
            <b>
                (d)
            </b>
            Sch&#228;tzen Sie die Wahrscheinlichkeit, dass mehr als {@a@} und weniger als {@b@} Saatk&#246;rner keimen, nach unten ab, indem Sie mit der Chebychev-Ungleichung eine reelle Zahl berechnen, sodass \(P({@a@} &lt; X &lt; {@b@})\geq\) [[input:ansD]]. [[validation:ansD]] [[feedback:prtD]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
<span class="multilang" lang="en">
    <div id="oer-aufgabe-19-info">
        <details>
            <summary>
                Content and scope of this task
            </summary>
            <p>
                In this task you will learn to estimate a likelihood with the Chebychev equation down. It consists of the following four subtasks.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    determine the distribution of random variables \(X\).
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    calculate the expected value and the variance of \(X\).
                </li>
                <li>
                    In
                    <b>
                        (c)
                    </b>
                    bring the desired probability to the form of chebychev imbalance.
                </li>
                <li>
                    In
                    <b>
                        (d)
                    </b>
                    estimate the desired probability with the Chebychev equation down.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-19">
        <p>
            Seeds of a particular variety germinate independently of one another with probability \(p={@dp(p)@}\). Total \(n={@n@}\) seeds sown. The random number of seed grains germinated therefrom is variable by chance \(X\) described.
        </p>
        <hr/>
    </div>
    <div id="oer-aufgabe-19a">
        <p>
            <b>
                (a)
            </b>
            What distribution is the random variable \(X\)?
        </p>
        <p>
            [[input:ansA]] [[validation:ansA]] [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-19b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            The random variable \(X\) has \(\text{Bin}({@n@},{@dp(p)@})\)- Distribution. Calculate the expected value and the variance of \(X\).
        </p>
        <p>
            The expected value of \(X\) is \(\text{E}(X)=\) [[input:ansB1]]. [[validation:ansB1]] [[feedback:prtB1]]
        </p>
        <p>
            The variance of \(X\) is \(\text{Var}(X)=\) [[input:ansB2]]. [[validation:ansB2]] [[feedback:prtB2]]
        </p>
        <p>
            [[feedback:prtB]]
        </p>
    </div>
    <div id="oer-aufgabe-19c" style="display:none;">
        <p>
            <b>
                (c)
            </b>
            Write the probability \(P({@a@} &lt; X &lt; {@b@})\) Step by step so that the Chebychev equation will be applicable: Subtract first on all sides of the non-equalification chain that describes the event the expected value of \(X\), which you have calculated in the previous part task. Then merge the non-equalification chain to a balance. Lastly, go over to the counter-probability so that the event is through a \(\geq\)- Signs instead of a \(&lt;\)-signs are described.
        </p>
        <p>
            <i>
                When clicking on "Check" get You only have a feedback when you fill out all input fields.
            </i>
        </p>
        <p>
            <br/>
            \(P({@a@} &lt; X &lt; {@b@})\)
            <br/>
            \(=P(\) [[input:ansC1]] \(&lt; X-\text{E}(X) &lt;\) [[input:ansC2]] \()\)
            <br/>
            \(=P(|X-\text{E}(X)| &lt;\) [[input:ansC3]] \()\)
            <br/>
            \(=1-P(|X-\text{E}(X)|\geq\) [[input:ansC4]] \()\)[[validation:ansC1]] [[validation:ansC2]] [[validation:ansC3]] [[validation:ansC4]] [[feedback:prtC]]
        </p>
    </div>
    <div id="oer-aufgabe-19d" style="display:none;">
        <p>
            <b>
                (d)
            </b>
            Guess the probability that more than {@a@} and less than {@b@} seed grains germinate, down by calculating a real number with the chebychev balance, so that \(P({@a@} &lt; X &lt; {@b@})\geq\) [[input:ansD]]. [[validation:ansD]] [[feedback:prtD]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="moodle_auto_format">
            <text/>
        </generalfeedback>
        <defaultgrade>4.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
Chebychev-Ungleichung

wurde entwickelt von
  
Christian M&#252;ller <c.mueller(at)hhu.de>

nach einer Idee von

Axel B&#252;cher <axel.buecher(at)hhu.de>
Peter Kern <kern(at)hhu.de>

an der Heinrich-Heine-Universit&#228;t D&#252;sseldorf.

Dieses Werk ist lizenziert unter einer Creative Commons Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie https://creativecommons.org/licenses/by-sa/4.0/.

SPDX-License-Identifier: CC-BY-SA-4.0

Technische Informationen:

Diese Aufgabe verwendet das Maxima Paket 'distrib'.

Diese Aufgabe bindet das Skript 'stackselbstlern.js' von Michael Kallweit f&#252;r die Aufgabennavigation ein.
*/

/* &#252;berpr&#252;fe, ob ans als Wahrscheinlichkeit interpretiert werden kann */
probabilityp(ans) := is(numberp(float(ans)) and ans >= 0 and ans <= 1)
/* runde x auf 3 Nachkommastellen */
dp(x) := decimalplaces(x, 3)
/* stelle x in wissenschaftlicher Notation dar */
SCI(x) := scientific_notation(x, 2)
/* Rundungsgenauigkeit bei numerischen Antworttests */
epsilon: 0.001

/* Anzahl der Saatk&#246;rner */
n: rand([1000, 2000, 3000, 4000])
/* Keimwahrscheinlichkeit */
p: rand([1 / 5, 2 / 5, 3 / 5, 4 / 5])
/* Mindest-/H&#246;chstanzahl der gekeimten Saatk&#246;rner */
c: rand([50, 100, 150, 200])
a: -c + n * p
b: c + n * p

tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
/* Abk&#252;rzungen f&#252;r Verteilungen */
Bin: tr("Binomialverteilung","binomial distribution")
Ber: tr("Bernoulli-Verteilung","bernoulli distribution")
U:   tr("Laplace-Verteilung","laplace distribution")
Geo: tr("Geometrische Verteilung","geometric distribution")

/* L&#246;sungen */
TansA: random_permutation([[1, true, Bin], [2, false, Ber], [3, false, Geo], [4, false, U]])
TansB1: n * p
TansB2: n * p * (1 - p)
TansC1: -c
TansC2: c
TansC3: c
TansC4: c
TansD: 1 - n * p * (1 - p) / c^2]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(n={@n@}\), \(p={@p@}\), \(c={@c@}\)</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ansA</name>
            <type>radio</type>
            <tans>TansA</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ansB1</name>
            <type>numerical</type>
            <tans>TansB1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansB2</name>
            <type>numerical</type>
            <tans>TansB2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansC1</name>
            <type>numerical</type>
            <tans>TansC1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansC2</name>
            <type>numerical</type>
            <tans>TansC2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansC3</name>
            <type>numerical</type>
            <tans>TansC3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansC4</name>
            <type>numerical</type>
            <tans>TansC4</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansD</name>
            <type>numerical</type>
            <tans>TansD</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prtA</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Mit der Binomialverteilung wird die Anzahl der Erfolge in \(n\) unabh&#228;ngigen Bernoulli-Experimenten mit Erfolgswahrscheinlichkeit \(p\) modelliert.
    </p>
    <p>
        Ob ein einzelnes Saatkorn keimt oder nicht, entspricht einem Bernoulli-Experiment mit Erfolgswahrscheinlichkeit \({@dp(p)@}\). Da die \({@n@}\) Saatk&#246;rner nach Voraussetzung unabh&#228;ngig voneinander keimen, besitzt die Anzahl der gekeimten Saatk&#246;rner somit eine Binomialverteilung mit \(n={@n@}\) und \(p={@dp(p)@}\).
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-19a');
show('oer-aufgabe-19b');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        With the binomial distribution, the number of successes in \(n\) independent Bernoulli experiments with probability of success \(p\) modelled.
    </p>
    <p>
        Whether a single seed grain germinates or not, corresponds to a Bernoulli experiment with probability of success \({@dp(p)@}\). There \({@n@}\) seed grains germinate independently of one another, so the number of germinated seed grains has a binomial distribution with \(n={@n@}\) and \(p={@dp(p)@}\).
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-19a');
show('oer-aufgabe-19b');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prtA-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Bernoulli-verteilte Zufallsvariablen treten bei Zufallsexperimenten auf, die nur zwei m&#246;gliche Ergebnisse haben oder bei denen man sich nur daf&#252;r interessiert, ob ein bestimmtes Ereignis eintritt oder nicht.
    </p>
    <p>
        Ob ein einzelnes Saatkorn keimt oder nicht, wird also durch eine Bernoulli-verteilte Zufallsvariable beschrieben. Gesucht ist hier jedoch die Verteilung der Anzahl aller gekeimten Saatk&#246;rner.
    </p>
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Bernoulli-distributed random variables occur in random experiments that have only two possible results or where one is only interested in whether or not a particular event occurs.
    </p>
    <p>
        Whether a single seed grain germinates or not, is thus described by a Bernoulli-distributed random variable. Here, however, the distribution of the number of germinated seed grains is sought.
    </p>
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtA-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Wird ein Zufallsexperiment so lange unabh&#228;ngig wiederholt, bis ein bestimmtes Ereignis erstmals eingetreten ist, hat die Anzahl der Fehlversuche vor dem ersten Eintreten des Ereignisses eine geometrische Verteilung.
    </p>
    <p>
        Werden die Saatk&#246;rner einzeln nacheinander ausges&#228;t, wird die Anzahl der nicht gekeimten Saatk&#246;rner, bis erstmals ein Saatkorn keimt, also durch eine geometrisch verteilte Zufallsvariable beschrieben. Gesucht ist hier jedoch die Verteilung der Anzahl aller gekeimten Saatk&#246;rner unter den {@n@} ausges&#228;ten Saatk&#246;rnern.
    </p>
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        If a random experiment is repeated independently until a certain event has occurred for the first time, the number of faulty tests before the first occurrence of the event has a geometric distribution.
    </p>
    <p>
        If the seed grains are sown individually one after the other, the number of unglued seed grains is described until a seed grain is germinated for the first time, i.e. by a geometrically distributed random variable. However, the distribution of the number of germinated seed grains among the {@n@} seed seeds.
    </p>
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prtA-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die
        <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='auch bekannt unter "diskrete Gleichverteilung"'>
            Laplace-Verteilung
        </span>
        wird dann verwendet, wenn eine Zufallsvariable einen endlichen Wertebereich hat und jeder m&#246;gliche Wert mit derselben Wahrscheinlichkeit angenommen wird.
    </p>
    <p>
        Zwar wird hier angenommen, dass jedes einzelne Saatkorn mit derselben Wahrscheinlichkeit keimt, gesucht ist jedoch die Verteilung der Anzahl aller gekeimten Saatk&#246;rner, die keine Gleichverteilung mehr ist.
    </p>
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The
        <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='auch bekannt unter "diskrete Gleichverteilung"'>
            Laplace distribution
        </span>
        is used when a random variable has a finite value range and each possible value is assumed with the same probability.
    </p>
    <p>
        Although it is assumed here that each individual seed grain germinates with the same probability, the distribution of the number of all germinated seed grains which is no longer a uniform distribution is sought.
    </p>
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtA-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtB</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>fullratsimp([ansB1,ansB2])</sans>
                <tans>[TansB1,TansB2]</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-19b');
show('oer-aufgabe-19c');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-19b');
show('oer-aufgabe-19c');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtB-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtB1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>fullratsimp(ansB1)</sans>
                <tans>TansB1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie haben den Erwartungswert richtig berechnet. Da \(X\) eine \(\text{Bin}({@n@},{@dp(p)@})\)-Verteilung besitzt, ist der Erwartungswert \(\text{E}(X)=np={@n@}\cdot {@dp(p)@}={@n*p@}\).
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You have calculated the expected value correctly. There. \(X\) a \(\text{Bin}({@n@},{@dp(p)@})\)-distribution is the expected value \(\text{E}(X)=np={@n@}\cdot {@dp(p)@}={@n*p@}\).
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtB1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Sie haben den Erwartungswert falsch berechnet. Erinnern Sie sich an die Formel, mit der Sie bei einer binomialverteilten Zufallsvariablen den Erwartungswert berechnen k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        They calculated the expected value incorrectly. Remember the formula with which you can calculate the expected value for a binomial-distributed random variable.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtB2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>fullratsimp(ansB2)</sans>
                <tans>TansB2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie haben die Varianz richtig berechnet. Da \(X\) eine \(\text{Bin}({@n@},{@dp(p)@})\)-Verteilung besitzt, ist die Varianz \(\text{Var}(X)=np(1-p)={@n@}\cdot {@dp(p)@}\cdot(1-{@dp(p)@})={@n*p*(1-p)@}\).
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You have calculated the variance correctly. There. \(X\) a \(\text{Bin}({@n@},{@dp(p)@})\)- distribution is the variance \(\text{Var}(X)=np(1-p)={@n@}\cdot {@dp(p)@}\cdot(1-{@dp(p)@})={@n*p*(1-p)@}\).
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtB2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Sie haben die Varianz falsch berechnet. Erinnern Sie sich an die Formel, mit der Sie bei einer binomialverteilten Zufallsvariablen die Varianz berechnen k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        You calculated the variance incorrectly. Remember the formula with which you can calculate the variance for a binomial-distributed random variables.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtC</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[/* Umwandlung mit fullratsimp, damit mit Dezimalzahlen und mit Br&#252;chen gerechnet werden kann */
AnsC1: fullratsimp(ansC1)
AnsC2: fullratsimp(ansC2)
AnsC3: fullratsimp(ansC3)
AnsC4: fullratsimp(ansC4)
TAnsC1: fullratsimp(TansC1)
TAnsC2: fullratsimp(TansC2)
TAnsC3: fullratsimp(TansC3)
TAnsC4: fullratsimp(TansC4)

checkAll: is([AnsC1, AnsC2, AnsC3, AnsC4] = [TAnsC1, TAnsC2, TAnsC3, TAnsC4])

checkFirst: is([AnsC1, AnsC2] = [TAnsC1, TAnsC2])
checkFirstEmptyEvent: is(AnsC1 >= AnsC2)
checkFirstWrongNumber: is([AnsC1, AnsC2] = [TAnsC2, TAnsC1])

checkSecond: is(AnsC3 = TAnsC3 and AnsC3 = AnsC2)
checkSecondEmptyEvent: is(AnsC3 < 0)
checkSecondWrongNumber: is(AnsC3 = AnsC2 and AnsC3 = -AnsC1)

checkThird: is(AnsC4 = TAnsC4 and AnsC4 = AnsC3)
checkThirdEmptyEvent: is(AnsC4 < 0)
checkThirdWrongNumber: is(AnsC4 = AnsC3)]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkAll</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtC-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben alle Werte richtig eingegeben und die Wahrscheinlichkeit korrekt umgeformt.
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-19c');
show('oer-aufgabe-19d');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have entered all values correctly and the likelihood has been correctly transformed.
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-19c');
show('oer-aufgabe-19d');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prtC-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkFirst</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prtC-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im ersten Umformungsschritt haben Sie die Ungleichungskette richtig umgeformt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the first forming step, you have properly reshaped the unequalification chain.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtC-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>10</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkAll</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtC-11-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtC-11-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkFirstEmptyEvent</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prtC-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im ersten Umformungsschritt ist die linke Seite der Ungleichungskette nicht kleiner als die rechte Seite. Das resultiert in einem leeren Ereignis mit der Wahrscheinlichkeit 0.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the first forming step, the left side of the imbalance chain is not smaller than the right side. This results in an empty event with probability 0.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prtC-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkFirstWrongNumber</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prtC-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im ersten Umformungsschritt haben Sie die Ungleichungskette falsch umgeformt, aber blo&#223; beide Eingaben miteinander verwechselt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the first forming step, you have incorrectly reshaped the non-uniform chain, but only both inputs are confused with each other.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prtC-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im ersten Umformungsschritt haben Sie die Ungleichungskette falsch umgeformt. Subtrahieren Sie auf allen Seiten der Ungleichungskette \({@a@} &lt; X &lt; {@b@}\) den Erwartungswert \(\text{E}(X)={@n*p@}\).
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the first forming step, you have incorrectly reshaped the unequalification chain. Subtract on all sides of the non-uniform chain \({@a@} &lt; X &lt; {@b@}\) the expected value \(\text{E}(X)={@n*p@}\).
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkSecond</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>7</truenextnode>
                <trueanswernote>prtC-5-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im zweiten Umformungsschritt haben Sie die Ungleichungskette richtig zu einer Betragsungleichung zusammengefasst.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the second reshaping step, you have correctly summarized the unequalification chain to an amount imbalance.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>5</falsenextnode>
                <falseanswernote>prtC-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkSecondEmptyEvent</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>7</truenextnode>
                <trueanswernote>prtC-6-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im zweiten Umformungsschritt haben Sie eine negative Zahl eingegeben. Das resultiert in einem leeren Ereignis mit der Wahrscheinlichkeit 0.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the second forming step, you entered a negative number. This results in an empty event with probability 0.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>6</falsenextnode>
                <falseanswernote>prtC-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>6</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkSecondWrongNumber</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>7</truenextnode>
                <trueanswernote>prtC-7-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im zweiten Umformungsschritt haben Sie zwar beide Ungleichungen richtig zu einer Betragsungleichung zusammengefasst, allerdings mit einem falschen Zahlenwert aus dem ersten Schritt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the second forming step, you have correctly combined both inequalities to an amount imbalance, but with an incorrect numerical value from the first step.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>7</falsenextnode>
                <falseanswernote>prtC-7-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im zweiten Umformungsschritt haben Sie die Ungleichungskette falsch zusammengefasst.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the second forming step, you have incorrectly combined the unequalification chain.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>7</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkThird</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2500000</truescore>
                <truepenalty/>
                <truenextnode>10</truenextnode>
                <trueanswernote>prtC-8-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im dritten Umformungsschritt haben Sie die Wahrscheinlichkeit richtig auf die Gegenwahrscheinlichkeit umgeschrieben.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the third forming step, you have correctly rewritten the probability to the counter probability.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>8</falsenextnode>
                <falseanswernote>prtC-8-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>8</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkThirdEmptyEvent</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>10</truenextnode>
                <trueanswernote>prtC-9-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im dritten Umformungsschritt haben Sie eine negative Zahl eingegeben. Das resultiert in einem sicheren Ereignis mit der Wahrscheinlichkeit 1.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the third forming step, you entered a negative number. This results in a safe event with probability 1.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>9</falsenextnode>
                <falseanswernote>prtC-9-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>9</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>checkThirdWrongNumber</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>10</truenextnode>
                <trueanswernote>prtC-10-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im dritten Umformungsschritt haben Sie die Wahrscheinlichkeit zwar richtig auf die Gegenwahrscheinlichkeit umgeschrieben, allerdings mit einem falschen Zahlenwert aus dem zweiten Schritt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the third forming step, you have correctly rewritten the probability to the counter probability, but with an incorrect numerical value from the second step.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>10</falsenextnode>
                <falseanswernote>prtC-10-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im dritten Umformungsschritt sind Sie falsch zur Gegenwahrscheinlichkeit &#252;bergegangen. Nur das Ungleichheitszeichen, mit dem das Ereignis beschrieben wird, &#228;ndert sich, aber nicht die linke oder rechte Seite der Ungleichung.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the third forming step, you have gone wrong to counterprobability. Only the inequality mark with which the event is described changes, but not the left or right side of the inequality.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtD</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>prob: cdf_binomial(b + 1, n, p) - cdf_binomial(a, n, p)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans><![CDATA[is(numberp(float(ansD)) and ansD<=1)]]></sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prtD-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtD-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihr Ergebnis ist keine Zahl oder ist gr&#246;&#223;er als 1 und kann daher keine untere Schranke f&#252;r eine Wahrscheinlichkeit sein.
    </p>
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Their result is no number or is greater than 1 and therefore cannot be a lower limit for a probability.
    </p>
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans><![CDATA[is(dp(TansD)>dp(ansD))]]></sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prtD-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihre Antwort ist zwar eine untere Schranke f&#252;r die gesuchte Wahrscheinlichkeit, jedoch erh&#228;lt man mithilfe der Chebychev-Ungleichung ein noch besseres Ergebnis.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Their answer is a lower limit for the desired probability, but with the help of the Chebychev equation you get an even better result.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtD-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans><![CDATA[is(dp(prob)>dp(ansD) and dp(ansD)>dp(TansD))]]></sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prtD-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihre Antwort ist zwar eine untere Schranke f&#252;r die gesuchte Wahrscheinlichkeit, jedoch ist sie gr&#246;&#223;er als diejenige, die man mithilfe der Chebychev-Ungleichung erh&#228;lt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Their answer is a lower limit for the likelihood sought, but it is greater than that obtained by the Chebychev imbalance.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prtD-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>fullratsimp(ansD)</sans>
                <tans>TansD</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtD-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben mit der Chebychev-Ungleichung die untere Schranke richtig berechnet. Eine m&#246;gliche Rechnung sieht vollst&#228;ndig aufgeschrieben folgenderma&#223;en aus:
\begin{align*}
P({@a@} &lt; X &lt; {@b@})
&amp;=P(-{@c@} &lt; X-{@n*p@} &lt; {@c@})\\
&amp;=P(|X-{@n*p@}| &lt; {@c@})\\
&amp;=P(|X-\text{E}(X)| &lt; {@c@})\\
&amp;=1-P(|X-\text{E}(X)|\geq {@c@})\\
&amp;\geq1-\frac{\text{Var}(X)}{{@c@}^2}\\
&amp;=1-\frac{{@n*p*(1-p)@}}{{@c@}^2}\\
&amp;\approx {@dp(TansD)@}.
\end{align*}
Eine auf diese Weise mit der Chebychev-Ungleichung berechnete Schranke ist allerdings typischerweise eine sehr grobe Absch&#228;tzung, die stark vom tats&#228;chlichen Wert der gesuchten Wahrscheinlichkeit abweichen kann. In manchen F&#228;llen kann die Chebychev-Ungleichung sogar eine untere bzw. obere Schranke ergeben, die negativ bzw. gr&#246;&#223;er als 1 ist, und damit keine brauchbare Information beinhaltet.
    </p>
    <p>
        Zum Vergleich ist im Beispiel dieser Aufgabe die mit der Binomialverteilung exakt berechnete Wahrscheinlichkeit (auf drei Nachkommastellen gerundet) \[P({@a@} &lt; X &lt; {@b@})\approx {@dp(prob)@}.\]
    </p>
    <hr/>
    <p>
        Sie haben die Aufgabe erfolgreich gel&#246;st! Jetzt sollten Sie mit der Chebychev-Ungleichung sicher rechnen k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have correctly calculated the lower limit with the Chebychev equation. A possible invoice looks fully written as follows:
\begin{align*}
P({@a@} &lt; X &lt; {@b@})
&amp;=P(-{@c@} &lt; X-{@n*p@} &lt; {@c@})\\
&amp;=P(|X-{@n*p@}| &lt; {@c@})\\
&amp;=P(|X-\text{E}(X)| &lt; {@c@})\\
&amp;=1-P(|X-\text{E}(X)|\geq {@c@})\\
&amp;\geq1-\frac{\text{Var}(X)}{{@c@}^2}\\
&amp;=1-\frac{{@n*p*(1-p)@}}{{@c@}^2}\\
&amp;\approx {@dp(TansD)@}.
\end{align*}However, a barrier calculated in this way with the chebychev imbalance is typically a very rough estimate which can differ greatly from the actual value of the desired probability. In some cases, the chebychev imbalance can even result in a lower or upper barrier, which is negative or greater than 1, and thus does not contain any useful information.
    </p>
    <p>
        For comparison, in the example of this task, the probability exactly calculated with the binomial distribution (rounded to three decimal places) \[P({@a@} &lt; X &lt; {@b@})\approx {@dp(prob)@}.\]
    </p>
    <hr/>
    <p>
        You have successfully solved the task! Now you should be able to count on the Chebychev balance.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtD-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben in (c) die Wahrscheinlichkeit \(P({@a@} &lt; X &lt; {@b@})\) schrittweise umgeschrieben und gesehen, dass \(P({@a@} &lt; X &lt; {@b@})=1-P(|X-\text{E}(X)|\geq {@c@})\) gilt. Wenden Sie auf den Term \(P(|X-\text{E}(X)|\geq {@c@})\) die Chebychev-Ungleichung an. Achten Sie auf das negative Vorzeichen, wodurch Sie eine Absch&#228;tzung nach unten bekommen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have in (c) the likelihood \(P({@a@} &lt; X &lt; {@b@})\) gradually rewritten and seen that \(P({@a@} &lt; X &lt; {@b@})=1-P(|X-\text{E}(X)|\geq {@c@})\) Article 2 Turn to the Term \(P(|X-\text{E}(X)|\geq {@c@})\) the Chebychev balance. Pay attention to the negative sign, making you get an estimate down.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1029091889</deployedseed>
        <deployedseed>644772454</deployedseed>
        <deployedseed>1291568217</deployedseed>
        <deployedseed>1204500129</deployedseed>
        <deployedseed>1625240249</deployedseed>
        <deployedseed>957844951</deployedseed>
        <deployedseed>1175735061</deployedseed>
        <deployedseed>56155399</deployedseed>
        <deployedseed>1855564291</deployedseed>
        <deployedseed>427554921</deployedseed>
        <deployedseed>229315034</deployedseed>
        <deployedseed>715678512</deployedseed>
        <deployedseed>962708878</deployedseed>
        <deployedseed>1463059177</deployedseed>
        <deployedseed>860117173</deployedseed>
        <deployedseed>1529413544</deployedseed>
        <deployedseed>1242484322</deployedseed>
        <deployedseed>1176127975</deployedseed>
        <deployedseed>1814572979</deployedseed>
        <deployedseed>1810937147</deployedseed>
        <deployedseed>375935930</deployedseed>
        <deployedseed>390509940</deployedseed>
        <deployedseed>588549409</deployedseed>
        <deployedseed>406982316</deployedseed>
        <deployedseed>1199943589</deployedseed>
        <deployedseed>702558086</deployedseed>
        <deployedseed>1670075134</deployedseed>
        <deployedseed>1670135982</deployedseed>
        <deployedseed>364546351</deployedseed>
        <deployedseed>1946872573</deployedseed>
        <deployedseed>1657423101</deployedseed>
        <deployedseed>1542430868</deployedseed>
        <deployedseed>906105267</deployedseed>
        <deployedseed>99798838</deployedseed>
        <deployedseed>1546082309</deployedseed>
        <deployedseed>1847665975</deployedseed>
        <deployedseed>1859243605</deployedseed>
        <deployedseed>1221794485</deployedseed>
        <deployedseed>177001041</deployedseed>
        <deployedseed>1631028820</deployedseed>
        <deployedseed>2099784464</deployedseed>
    </question>
    

</quiz>