<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2023-03-07 Totale Wahrscheinlichkeit</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <div id="oer-aufgabe-31-info">
        <details>
            <summary>
                Inhalt und Umfang dieser Aufgabe
            </summary>
            <p>
                In dieser Aufgabe lernen Sie, Wahrscheinlichkeiten mit der Formel von der totalen Wahrscheinlichkeit zu berechnen. Sie besteht aus den folgenden zwei Teilaufgaben.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    geben Sie die aus dem Sachkontext bekannten Wahrscheinlichkeiten in Form einer Tabelle an.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    berechnen Sie mithilfe der Informationen aus (a) die gesuchte Wahrscheinlichkeit.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-31">
        <p>
            Bei Zellen einer bestimmten Art wird innerhalb einer Zeiteinheit das folgende Verhalten beobachtet: Unabh&#228;ngig voneinander teilen sich die Zellen mit Wahrscheinlichkeit {@dp(p2)@}, sie bleiben mit Wahrscheinlichkeit {@dp(p1)@} unver&#228;ndert oder sie sterben mit Wahrscheinlichkeit {@dp(p0)@} ab. Ausgehend von einer einzelnen Zelle bezeichnen wir f&#252;r \(j\in\{0,1,2\}\) mit \(A_j\) das Ereignis, dass nach einer Zeiteinheit \(j\) Zellen vorhanden sind, und f&#252;r \(i\in\{0,1,2,3,4\}\) mit \(B_i\) das Ereignis, dass nach zwei Zeiteinheiten \(i\) Zellen vorhanden sind.
        </p>
        <hr/>
    </div>
    <div id="oer-aufgabe-31a">
        <p>
            <b>
                (a)
            </b>
            Berechnen Sie f&#252;r alle m&#246;glichen Paare von \(i\) und \(j\) die bedingte Wahrscheinlichkeit \(P(B_i|A_j)\) und tragen Sie diese in die folgende Tabelle ein.
        </p>
        <p>
            <i>
                Beim Klick auf "Pr&#252;fen" bekommen Sie erst dann ein Feedback, wenn Sie alle Eingabefelder ausgef&#252;llt haben.
            </i>
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(B_i|A_j)\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=0\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=1\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=2\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=3\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=4\)
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=0\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA11]] [[validation:ansA11]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA12]] [[validation:ansA12]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA13]] [[validation:ansA13]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA14]] [[validation:ansA14]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA15]] [[validation:ansA15]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=1\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA21]] [[validation:ansA21]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA22]] [[validation:ansA22]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA23]] [[validation:ansA23]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA24]] [[validation:ansA24]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA25]] [[validation:ansA25]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=2\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA31]] [[validation:ansA31]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA32]] [[validation:ansA32]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA33]] [[validation:ansA33]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA34]] [[validation:ansA34]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA35]] [[validation:ansA35]]
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-31b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            Berechnen Sie die Wahrscheinlichkeit \(P(B_{@N[1]-1@})\), dass nach zwei Zeiteinheiten {@N[2]@} vorhanden {@N[3]@}. Zur Erinnerung sind die Ergebnisse aus Teil (a) in der folgenden Tabelle zusammengefasst.
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(B_i|A_j)\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=0\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=1\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=2\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=3\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=4\)
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=0\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,1]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,2]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,3]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,4]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,5]@}\)
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=1\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,1]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,2]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,3]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,4]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,5]@}\)
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=2\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,1]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,2]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,3]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,4]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,5]@}\)
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            Die Wahrscheinlichkeit f&#252;r {@N[2]@} nach zwei Zeiteinheiten ist \(P(B_{@N[1]-1@})=\) [[input:ansB]].
            <i>
                Geben Sie den exakten Wert als Formel ein oder runden Sie das Ergebnis auf mindestens drei Nachkommastellen.
            </i>
            [[validation:ansB]] [[feedback:prtB]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
<span class="multilang" lang="en">
    <div id="oer-aufgabe-31-info">
        <details>
            <summary>
                Content and scope of this task
            </summary>
            <p>
                In this task you will learn to calculate probabilities with the formula of total probability. It consists of the following two subtasks.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    specify the probabilities known from the context in the form of a table.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    use the information from (a) to calculate the desired probability.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-31">
        <p>
            In cells of a specific type, the following behavior is observed within a time unit: {@dp(p2)@}, they remain with probability {@dp(p1)@} unchanged or they die with probability {@dp(p0)@} off. Starting from a single cell, we refer to \(j\in\{0,1,2\}\) with \(A_j\) the event that after a time unit \(j\) cells are present, and for \(i\in\{0,1,2,3,4\}\) with \(B_i\) the event that after two time units \(i\) Cells are present.
        </p>
        <hr/>
    </div>
    <div id="oer-aufgabe-31a">
        <p>
            <b>
                (a)
            </b>
            Calculate for all possible pairs of \(i\) and \(j\) the conditional probability \(P(B_i|A_j)\) and enter it in the table below.
        </p>
        <p>
            <i>
                When clicking on "Check" get You only have a feedback when you fill out all input fields.
            </i>
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(B_i|A_j)\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=0\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=1\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=2\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=3\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=4\)
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=0\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA11]] [[validation:ansA11]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA12]] [[validation:ansA12]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA13]] [[validation:ansA13]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA14]] [[validation:ansA14]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA15]] [[validation:ansA15]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=1\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA21]] [[validation:ansA21]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA22]] [[validation:ansA22]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA23]] [[validation:ansA23]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA24]] [[validation:ansA24]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA25]] [[validation:ansA25]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=2\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA31]] [[validation:ansA31]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA32]] [[validation:ansA32]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA33]] [[validation:ansA33]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA34]] [[validation:ansA34]]
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA35]] [[validation:ansA35]]
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-31b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            Calculate the probability \(P(B_{@N[1]-1@})\)that after two time units {@N[2]@} available {@N[3]@}. The results from Part (a) are summarized in the following table.
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(B_i|A_j)\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=0\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=1\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=2\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=3\)
                        </td>
                        <td style="border: 1px solid black;">
                            \(i=4\)
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=0\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,1]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,2]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,3]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,4]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[1,5]@}\)
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=1\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,1]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,2]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,3]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,4]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[2,5]@}\)
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(j=2\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,1]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,2]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,3]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,4]@}\)
                        </td>
                        <td style="border: 1px solid black;">
                            \({@TansA[3,5]@}\)
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            The probability of {@N[2]@} after two time units \(P(B_{@N[1]-1@})=\) [[input:ansB]].
            <i>
                Enter the exact value as a formula or round the result to at least three decimal places.
            </i>
            [[validation:ansB]] [[feedback:prtB]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>2.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
Totale Wahrscheinlichkeit

wurde entwickelt von
  
Christian M&#252;ller <c.mueller(at)hhu.de>

nach einer Idee von

Axel B&#252;cher <axel.buecher(at)hhu.de>
Peter Kern <kern(at)hhu.de>

an der Heinrich-Heine-Universit&#228;t D&#252;sseldorf.

Dieses Werk ist lizenziert unter einer Creative Commons Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie https://creativecommons.org/licenses/by-sa/4.0/.

SPDX-License-Identifier: CC-BY-SA-4.0

Technische Informationen:

Diese Aufgabe bindet das Skript 'stackselbstlern.js' von Michael Kallweit f&#252;r die Aufgabennavigation ein.
*/

/* &#252;berpr&#252;fe, ob ans als Wahrscheinlichkeit interpretiert werden kann */
probabilityp(ans) := is(numberp(float(ans)) and ans >= 0 and ans <= 1)
/* runde x auf 4 Nachkommastellen */
dp(x) := decimalplaces(x, 4)
/* stelle x in wissenschaftlicher Notation dar */
SCI(x) := scientific_notation(x, 2)
/* Rundungsgenauigkeit bei numerischen Antworttests */
epsilon: 0.001

/* Teilungswahrscheinlichkeiten */
p0: dp(rand_with_step(0.05, 0.25, 0.05))
p1: dp(rand_with_step(0.05, 0.25, 0.05))
p2: dp(1 - p0 - p1)
/* Anzahl der Tochterzellen in der zweiten Generation */
tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
N: rand([[1, tr("0 Zellen","0 cells"), tr("sind","are")], [2, tr("1 Zelle","1 cell"), tr("ist","is")], [3, tr("2 Zellen","2 cells"), tr("sind","are")], [4, tr("3 Zellen","3 cells"), tr("sind","are")], [5, tr("4 Zellen","4 cells"), tr("sind","are")]])

/* L&#246;sungen */
TansA: matrix([1, 0, 0, 0, 0], [dp(p0), dp(p1), dp(p2), 0, 0], [dp(p0^2), dp(2 * p0 * p1), dp(p1^2 + 2 * p0 * p2), dp(2 * p1 * p2), dp(p2^2)])
TansB: col(TansA, N[1]).matrix([p0], [p1], [p2])]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(p_0={@p0@}\), \(p_1={@p1@}\), \(p_2={@p2@}\), \(N={@N[1]@}\)</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ansA11</name>
            <type>numerical</type>
            <tans>TansA[1,1]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA12</name>
            <type>numerical</type>
            <tans>TansA[1,2]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA13</name>
            <type>numerical</type>
            <tans>TansA[1,3]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA14</name>
            <type>numerical</type>
            <tans>TansA[1,4]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA15</name>
            <type>numerical</type>
            <tans>TansA[1,5]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA21</name>
            <type>numerical</type>
            <tans>TansA[2,1]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA22</name>
            <type>numerical</type>
            <tans>TansA[2,2]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA23</name>
            <type>numerical</type>
            <tans>TansA[2,3]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA24</name>
            <type>numerical</type>
            <tans>TansA[2,4]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA25</name>
            <type>numerical</type>
            <tans>TansA[2,5]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA31</name>
            <type>numerical</type>
            <tans>TansA[3,1]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA32</name>
            <type>numerical</type>
            <tans>TansA[3,2]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA33</name>
            <type>numerical</type>
            <tans>TansA[3,3]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA34</name>
            <type>numerical</type>
            <tans>TansA[3,4]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA35</name>
            <type>numerical</type>
            <tans>TansA[3,5]</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansB</name>
            <type>numerical</type>
            <tans>TansB</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prtA</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ansA: matrix([ansA11, ansA12, ansA13, ansA14, ansA15], [ansA21, ansA22, ansA23, ansA24, ansA25], [ansA31, ansA32, ansA33, ansA34, ansA35])
f(x) := is(equal(x, 1))

/* Umwandlung mit fullratsimp, damit mit Dezimalzahlen und mit Br&#252;chen gerechnet werden kann */
AnsA: fullratsimp(ansA)
TAnsA: fullratsimp(TansA)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>every(probabilityp, AnsA)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prtA-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtA-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Nicht alle Tabelleneintr&#228;ge liegen im Intervall \([0,1]\). Diese Werte k&#246;nnen daher nicht als Wahrscheinlichkeiten interpretiert werden.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Not all table entries are in interval \([0,1]\). These values cannot therefore be interpreted as probabilities.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA</sans>
                <tans>TAnsA</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben alle bedingten Wahrscheinlichkeiten richtig berechnet.
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-31a');
show('oer-aufgabe-31b');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have calculated all conditional probabilities correctly.
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-31a');
show('oer-aufgabe-31b');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtA-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>10</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[2,3]</sans>
                <tans>TAnsA[2,3]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>11</truenextnode>
                <trueanswernote>prtA-11-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>11</falsenextnode>
                <falseanswernote>prtA-11-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_2|A_1)\) ist falsch. Wenn 1 Zelle vorhanden ist, k&#246;nnen daraus nur dann 2 Zellen entstehen, wenn sich diese Zelle teilt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_2|A_1)\) is wrong. If 1 cell is present, 2 cells can only arise from this when this cell splits.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>11</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[2,4]</sans>
                <tans>TAnsA[2,4]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>12</truenextnode>
                <trueanswernote>prtA-12-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>12</falsenextnode>
                <falseanswernote>prtA-12-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_3|A_1)\) ist falsch. Wenn 1 Zelle vorhanden ist, k&#246;nnen daraus nicht 3 Zellen entstehen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_3|A_1)\) is wrong. If 1 cell is present, 3 cells cannot arise from it.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>12</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[2,5]</sans>
                <tans>TAnsA[2,5]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>13</truenextnode>
                <trueanswernote>prtA-13-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>13</falsenextnode>
                <falseanswernote>prtA-13-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_4|A_1)\) ist falsch. Wenn 1 Zelle vorhanden ist, k&#246;nnen daraus nicht 4 Zellen entstehen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_4|A_1)\) is wrong. If 1 cell is present, there can be no 4 cells.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>13</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[3,1]</sans>
                <tans>TAnsA[3,1]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>14</truenextnode>
                <trueanswernote>prtA-14-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>14</falsenextnode>
                <falseanswernote>prtA-14-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_0|A_2)\) ist falsch. Wenn 2 Zellen vorhanden sind, k&#246;nnen daraus nur dann 0 Zellen entstehen, wenn beide Zellen absterben.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_0|A_2)\) is wrong. If 2 cells are present, 0 cells can only arise from this when both cells die.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>14</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[3,2]</sans>
                <tans>TAnsA[3,2]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>15</truenextnode>
                <trueanswernote>prtA-15-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>15</falsenextnode>
                <falseanswernote>prtA-15-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_1|A_2)\) ist falsch. Wenn 2 Zellen vorhanden sind, kann daraus nur dann 1 Zelle entstehen, wenn eine der beiden Zellen abstirbt und die andere unver&#228;ndert bleibt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_1|A_2)\) is wrong. If 2 cells are present, this can only result in 1 cell if one of the two cells dies and the other remains unchanged.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>15</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[3,3]</sans>
                <tans>TAnsA[3,3]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>16</truenextnode>
                <trueanswernote>prtA-16-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>16</falsenextnode>
                <falseanswernote>prtA-16-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_2|A_2)\) ist falsch. Wenn 2 Zellen vorhanden sind, k&#246;nnen daraus nur dann 2 Zellen entstehen, wenn entweder beide Zellen unver&#228;ndert bleiben oder wenn eine der beiden Zellen abstirbt und die andere sich teilt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_2|A_2)\) is wrong. If 2 cells are present, then only 2 cells can arise from this if either both cells remain unchanged or if one of the two cells dies and the other splits.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>16</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[3,4]</sans>
                <tans>TAnsA[3,4]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>17</truenextnode>
                <trueanswernote>prtA-17-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>17</falsenextnode>
                <falseanswernote>prtA-17-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_3|A_2)\) ist falsch. Wenn 2 Zellen vorhanden sind, k&#246;nnen daraus nur dann 3 Zellen entstehen, wenn sich eine der beiden Zellen teilt und die andere unver&#228;ndert bleibt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_3|A_2)\) is wrong. If 2 cells are present, 3 cells can only arise from this if one of the two cells splits and the other remains unchanged.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>17</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[3,5]</sans>
                <tans>TAnsA[3,5]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-18-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtA-18-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_4|A_2)\) ist falsch. Wenn 2 Zellen vorhanden sind, k&#246;nnen daraus nur dann 4 Zellen entstehen, wenn sich beide Zellen teilen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_4|A_2)\) is wrong. If 2 cells are present, then only 4 cells can be formed if both cells share.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>every(f, makelist(sum(dp(AnsA[i, j]), j, 1, 5), i, 3))</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prtA-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prtA-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Nicht alle Zeilensummen ergeben 1. Das muss aufgrund von \(\sum\limits_{i=0}^4P(B_i|A_j)=1\) f&#252;r jedes \(j\in\{0,1,2\}\) aber erf&#252;llt sein.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Not all row sums result 1. This must be due to \(\sum\limits_{i=0}^4P(B_i|A_j)=1\) for each \(j\in\{0,1,2\}\) but be fulfilled.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[1,1]</sans>
                <tans>TAnsA[1,1]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prtA-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prtA-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_0|A_0)\) ist falsch. Wenn 0 Zellen vorhanden sind, wird es sicher bei 0 Zellen bleiben.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_0|A_0)\) is wrong. If 0 cells are present, it will stay safe at 0 cells.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[1,2]</sans>
                <tans>TAnsA[1,2]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>5</truenextnode>
                <trueanswernote>prtA-5-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>5</falsenextnode>
                <falseanswernote>prtA-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_1|A_0)\) ist falsch. Wenn 0 Zellen vorhanden sind, kann daraus nicht 1 Zelle entstehen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_1|A_0)\) is wrong. If 0 cells are present, there cannot be 1 cell.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[1,3]</sans>
                <tans>TAnsA[1,3]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>6</truenextnode>
                <trueanswernote>prtA-6-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>6</falsenextnode>
                <falseanswernote>prtA-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_2|A_0)\) ist falsch. Wenn 0 Zellen vorhanden sind, k&#246;nnen daraus nicht 2 Zellen entstehen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_2|A_0)\) is wrong. If 0 cells are present, there cannot be 2 cells.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>6</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[1,4]</sans>
                <tans>TAnsA[1,4]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>7</truenextnode>
                <trueanswernote>prtA-7-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>7</falsenextnode>
                <falseanswernote>prtA-7-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_3|A_0)\) ist falsch. Wenn 0 Zellen vorhanden sind, k&#246;nnen daraus nicht 3 Zellen entstehen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_3|A_0)\) is wrong. If 0 cells are present, 3 cells cannot arise from this.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>7</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[1,5]</sans>
                <tans>TAnsA[1,5]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>8</truenextnode>
                <trueanswernote>prtA-8-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>8</falsenextnode>
                <falseanswernote>prtA-8-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_4|A_0)\) ist falsch. Wenn 0 Zellen vorhanden sind, k&#246;nnen daraus nicht 4 Zellen entstehen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_4|A_0)\) is wrong. If 0 cells are present, there can be no 4 cells.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>8</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[2,1]</sans>
                <tans>TAnsA[2,1]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>9</truenextnode>
                <trueanswernote>prtA-9-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>9</falsenextnode>
                <falseanswernote>prtA-9-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_0|A_1)\) ist falsch. Wenn 1 Zelle vorhanden ist, k&#246;nnen daraus nur dann 0 Zellen entstehen, wenn diese Zelle abstirbt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_0|A_1)\) is wrong. If 1 cell is present, 0 cells can only arise from this if this cell is lost.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>9</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>AnsA[2,2]</sans>
                <tans>TAnsA[2,2]</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>10</truenextnode>
                <trueanswernote>prtA-10-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>10</falsenextnode>
                <falseanswernote>prtA-10-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(B_1|A_1)\) ist falsch. Wenn 1 Zelle vorhanden ist, kann daraus nur dann 1 Zelle entstehen, wenn diese Zelle unver&#228;ndert bleibt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(B_1|A_1)\) is wrong. If 1 cell is present, it can only result in 1 cell if this cell remains unchanged.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtB</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>S: matrix([ansA11, ansA12, ansA13, ansA14, ansA15], [ansA21, ansA22, ansA23, ansA24, ansA25], [ansA31, ansA32, ansA33, ansA34, ansA35])
Sprob(j) := dp(col(S, j).matrix([p0], [p1], [p2]))
list: makelist(Sprob(j), j, 5)</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>probabilityp(ansB)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prtB-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtB-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihr Ergebnis liegt nicht im Intervall \([0,1]\) und kann daher nicht als Wahrscheinlichkeit interpretiert werden.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your result is not in interval \([0,1]\) and therefore cannot be interpreted as probability.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ansB</sans>
                <tans>TansB</tans>
                <testoptions>epsilon</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben die gesuchte Wahrscheinlichkeit richtig berechnet. Eine m&#246;gliche Rechnung mit der Formel von der totalen Wahrscheinlichkeit sieht vollst&#228;ndig aufgeschrieben folgenderma&#223;en aus:
\begin{align*}
P(B_{@N[1]-1@})
&amp; =\sum_{j=0}^2P(B_{@N[1]-1@}|A_j)P(A_j) \\
&amp; =P(B_{@N[1]-1@}|A_0)P(A_0)+P(B_{@N[1]-1@}|A_1)P(A_1)+P(B_{@N[1]-1@}|A_2)P(A_2) \\
&amp; ={@TansA[1,N[1]]@}\cdot{@dp(p0)@}+{@TansA[2,N[1]]@}\cdot{@dp(p1)@}+{@TansA[3,N[1]]@}\cdot{@dp(p2)@} \\
&amp; \approx {@dp(TansB)@}.
\end{align*}
    </p>
    <hr/>
    <p>
        Sie haben die Aufgabe erfolgreich gel&#246;st! Jetzt sollten Sie mit der Formel von der totalen Wahrscheinlichkeit sicher rechnen k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have calculated the desired probability correctly. A possible invoice with the formula of total probability looks fully written as follows:
\begin{align*}
P(B_{@N[1]-1@})
&amp; =\sum_{j=0}^2P(B_{@N[1]-1@}|A_j)P(A_j) \\
&amp; =P(B_{@N[1]-1@}|A_0)P(A_0)+P(B_{@N[1]-1@}|A_1)P(A_1)+P(B_{@N[1]-1@}|A_2)P(A_2) \\
&amp; ={@TansA[1,N[1]]@}\cdot{@dp(p0)@}+{@TansA[2,N[1]]@}\cdot{@dp(p1)@}+{@TansA[3,N[1]]@}\cdot{@dp(p2)@} \\
&amp; \approx {@dp(TansB)@}.
\end{align*}
    </p>
    <hr/>
    <p>
        You have successfully solved the task! Now you should be sure to count on the formula of total probability.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtB-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ansB</sans>
                <tans>Sprob(N[1])</tans>
                <testoptions>epsilon</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben die Formel von der totalen Wahrscheinlichkeit richtig angewendet. Allerdings sind in (a) noch falsche Ergebnisse f&#252;r \(i={@N[1]-1@}\) und Sie haben somit einen Folgefehler gemacht. Korrigieren Sie Ihre L&#246;sung in (a), wiederholen Sie anschlie&#223;end die Rechnung mit den richtigen Zahlenwerten und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have applied the formula of total probability correctly. However, in (a) there are still false results for \(i={@N[1]-1@}\) and you have made a follow-up error. Correct your solution in (a), then repeat the invoice with the correct numerical values and then click "Check".
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prtB-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>member(ansB, list)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.5000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben die Formel von der totalen Wahrscheinlichkeit richtig angewendet. Allerdings haben Sie nicht mit den Zahlenwerten f&#252;r \(i={@N[1]-1@}\) aus (a) gerechnet. Wiederholen Sie die Rechnung mit den Zahlenwerten f&#252;r \(i={@N[1]-1@}\) und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have applied the formula of total probability correctly. However, you do not have the numerical values for \(i={@N[1]-1@}\) from (a). Repeat the invoice with the numerical values for \(i={@N[1]-1@}\) and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtB-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Verwenden Sie die Formel von der totalen Wahrscheinlichkeit und Ihre Ergebnisse aus (a), um die gesuchte Wahrscheinlichkeit zu berechnen.
    </p>
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Use the formula of total probability and your results from (a) to calculate the desired probability.
    </p>
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1512224325</deployedseed>
        <deployedseed>1285031254</deployedseed>
        <deployedseed>29268737</deployedseed>
        <deployedseed>1673968635</deployedseed>
        <deployedseed>202056078</deployedseed>
        <deployedseed>1960682702</deployedseed>
        <deployedseed>163003143</deployedseed>
        <deployedseed>1422693915</deployedseed>
        <deployedseed>1018842449</deployedseed>
        <deployedseed>2126702145</deployedseed>
        <deployedseed>1632419053</deployedseed>
        <deployedseed>1177984436</deployedseed>
        <deployedseed>218280321</deployedseed>
        <deployedseed>1667992433</deployedseed>
        <deployedseed>1592134522</deployedseed>
        <deployedseed>693670988</deployedseed>
        <deployedseed>1476685950</deployedseed>
        <deployedseed>356447129</deployedseed>
        <deployedseed>355890059</deployedseed>
        <deployedseed>2067225818</deployedseed>
        <deployedseed>1467618281</deployedseed>
        <deployedseed>704125380</deployedseed>
        <deployedseed>452244612</deployedseed>
        <deployedseed>226208943</deployedseed>
        <deployedseed>528641357</deployedseed>
        <deployedseed>2015521835</deployedseed>
        <deployedseed>1877035097</deployedseed>
        <deployedseed>1549921989</deployedseed>
        <deployedseed>496915429</deployedseed>
        <deployedseed>1693425811</deployedseed>
        <deployedseed>1949732295</deployedseed>
        <deployedseed>1890712380</deployedseed>
        <deployedseed>898958390</deployedseed>
        <deployedseed>1998677089</deployedseed>
        <deployedseed>1734194520</deployedseed>
        <deployedseed>1950918134</deployedseed>
        <deployedseed>422344341</deployedseed>
        <deployedseed>863599198</deployedseed>
        <deployedseed>1622308351</deployedseed>
        <deployedseed>1238444903</deployedseed>
        <deployedseed>1656506934</deployedseed>
        <deployedseed>1507645354</deployedseed>
        <deployedseed>713060140</deployedseed>
        <deployedseed>1633624998</deployedseed>
        <deployedseed>1800437600</deployedseed>
        <deployedseed>1771898308</deployedseed>
        <deployedseed>1387735983</deployedseed>
        <deployedseed>1816955636</deployedseed>
        <deployedseed>1695621408</deployedseed>
        <deployedseed>135314183</deployedseed>
        <deployedseed>327807827</deployedseed>
        <deployedseed>1282676900</deployedseed>
        <deployedseed>110738903</deployedseed>
        <deployedseed>215164460</deployedseed>
        <deployedseed>842620524</deployedseed>
        <deployedseed>2026610644</deployedseed>
        <deployedseed>1628291633</deployedseed>
        <deployedseed>447773691</deployedseed>
        <deployedseed>1061557664</deployedseed>
        <deployedseed>1878321790</deployedseed>
        <deployedseed>822202328</deployedseed>
        <deployedseed>1645873872</deployedseed>
        <deployedseed>403172903</deployedseed>
        <deployedseed>1363056966</deployedseed>
        <deployedseed>2085238224</deployedseed>
        <deployedseed>1270565505</deployedseed>
        <deployedseed>990441390</deployedseed>
        <deployedseed>1566569956</deployedseed>
        <deployedseed>1758508061</deployedseed>
    </question>
    

</quiz>