<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2023-03-07 Satz von Bayes</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <div id="oer-aufgabe-30-info">
        <details>
            <summary>
                Inhalt und Umfang dieser Aufgabe
            </summary>
            <p>
                In dieser Aufgabe lernen Sie, Wahrscheinlichkeiten mit dem Satz von Bayes zu berechnen. Sie besteht aus den folgenden drei Teilaufgaben.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    geben Sie die aus dem Sachkontext bekannten Wahrscheinlichkeiten in Form einer Tabelle an.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    und
                    <b>
                        (c)
                    </b>
                    berechnen Sie mithilfe der Informationen aus (a) die gesuchten bedingten Wahrscheinlichkeiten.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-30">
        <p>
            Bei einer eingehenden E-Mail handelt es sich mit einer Wahrscheinlichkeit von {@dp(probS)@} um eine Spammail. Wenn eine E-Mail eine Spammail ist, enth&#228;lt sie mit einer Wahrscheinlichkeit von {@dp(probWS)@} das Wort "Gewinn". Wenn eine E-Mail keine Spammail ist, enth&#228;lt sie das Wort "Gewinn" nur mit einer Wahrscheinlichkeit von {@dp(probWSc)@}. Sei \(S\) das Ereignis, dass eine Spammail eingeht, und sei \(W\) das Ereignis, dass eine E-Mail das Wort "Gewinn" enth&#228;lt.
        </p>
        <hr/>
    </div>
    <div id="oer-aufgabe-30a">
        <p>
            <b>
                (a)
            </b>
            Geben Sie die folgenden Wahrscheinlichkeiten an.
        </p>
        <p>
            <i>
                Beim Klick auf "Pr&#252;fen" bekommen Sie erst dann ein Feedback, wenn Sie alle Eingabefelder ausgef&#252;llt haben.
            </i>
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA1]] [[validation:ansA1]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA2]] [[validation:ansA2]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA3]] [[validation:ansA3]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA4]] [[validation:ansA4]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA5]] [[validation:ansA5]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA6]] [[validation:ansA6]]
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-30b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            Sie erhalten eine E-Mail, die das Wort "Gewinn" enth&#228;lt. Berechnen Sie die Wahrscheinlichkeit, dass es sich {@listB[2]@} handelt. Zur Erinnerung sind die Ergebnisse aus Teil (a) in der folgenden Tabelle zusammengefasst.
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probSc@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWcS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWSc@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWcSc@}
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            Die Wahrscheinlichkeit betr&#228;gt [[input:ansB]].
            <i>
                Geben Sie den exakten Wert als Formel ein oder runden Sie das Ergebnis auf mindestens drei Nachkommastellen.
            </i>
            [[validation:ansB]] [[feedback:prtB]]
        </p>
    </div>
    <div id="oer-aufgabe-30c" style="display:none;">
        <p>
            <b>
                (c)
            </b>
            Sie erhalten eine E-Mail, die das Wort "Gewinn" nicht enth&#228;lt. Berechnen Sie die Wahrscheinlichkeit, dass es sich {@listC[2]@} handelt. Zur Erinnerung sind die Ergebnisse aus Teil (a) in der folgenden Tabelle zusammengefasst.
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probSc@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWcS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWSc@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWcSc@}
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            Die Wahrscheinlichkeit betr&#228;gt [[input:ansC]].
            <i>
                Geben Sie den exakten Wert als Formel ein oder runden Sie das Ergebnis auf mindestens drei Nachkommastellen.
            </i>
            [[validation:ansC]] [[feedback:prtC]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
<span class="multilang" lang="en">
    <div id="oer-aufgabe-30-info">
        <details>
            <summary>
                Content and scope of this task
            </summary>
            <p>
                In this task you will learn to calculate probabilities with the set of Bayes. It consists of the following three subtasks.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    specify the probabilities known from the context in the form of a table.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    and
                    <b>
                        (c)
                    </b>
                    use the information from (a) to calculate the desired conditional probabilities.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-30">
        <p>
            An incoming email is likely to be {@dp(probS)@} a spammail. If an email is a spammail, it contains a probability of {@dp(probWS)@} the word "winning". If an email is not a spammail, it contains the word "winning" only with a probability of {@dp(probWSc)@}. Be careful \(S\) the event that a spammail enters, and \(W\) the event that an email contains the word "winning".
        </p>
        <hr/>
    </div>
    <div id="oer-aufgabe-30a">
        <p>
            <b>
                (a)
            </b>
            Specify the following probabilities.
        </p>
        <p>
            <i>
                When clicking on "Check" get You only have a feedback when you fill out all input fields.
            </i>
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA1]] [[validation:ansA1]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA2]] [[validation:ansA2]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA3]] [[validation:ansA3]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA4]] [[validation:ansA4]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA5]] [[validation:ansA5]]
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            [[input:ansA6]] [[validation:ansA6]]
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-30b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            You will receive an email containing the word "winning". Calculate the probability that it will {@listB[2]@} in the case of the The results from Part (a) are summarized in the following table.
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probSc@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWcS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWSc@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWcSc@}
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            The probability is [[input:ansB]].
            <i>
                Enter the exact value as a formula or round the result to at least three decimal places.
            </i>
            [[validation:ansB]] [[feedback:prtB]]
        </p>
    </div>
    <div id="oer-aufgabe-30c" style="display:none;">
        <p>
            <b>
                (c)
            </b>
            You will receive an email that does not contain the word "win". Calculate the probability that it will {@listC[2]@} in the case of the The results from Part (a) are summarized in the following table.
        </p>
        <p>
            <table style="margin-left: auto; margin-right: auto;">
                <tbody>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probSc@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWcS@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWSc@}
                        </td>
                    </tr>
                    <tr>
                        <td style="border: 1px solid black;">
                            \(P(W^c|S^c)\)
                        </td>
                        <td style="border: 1px solid black;">
                            {@probWcSc@}
                        </td>
                    </tr>
                </tbody>
            </table>
        </p>
        <p>
            The probability is [[input:ansC]].
            <i>
                Enter the exact value as a formula or round the result to at least three decimal places.
            </i>
            [[validation:ansC]] [[feedback:prtC]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>3.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
Satz von Bayes

wurde entwickelt von
  
Christian M&#252;ller <c.mueller(at)hhu.de>

nach einer Idee von

Axel B&#252;cher <axel.buecher(at)hhu.de>
Peter Kern <kern(at)hhu.de>

an der Heinrich-Heine-Universit&#228;t D&#252;sseldorf.

Dieses Werk ist lizenziert unter einer Creative Commons Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie https://creativecommons.org/licenses/by-sa/4.0/.

SPDX-License-Identifier: CC-BY-SA-4.0

Technische Informationen:

Diese Aufgabe bindet das Skript 'stackselbstlern.js' von Michael Kallweit f&#252;r die Aufgabennavigation ein.
*/

/* &#252;berpr&#252;fe, ob ans als Wahrscheinlichkeit interpretiert werden kann */
probabilityp(ans) := is(numberp(float(ans)) and ans >= 0 and ans <= 1)
/* runde x auf 3 Nachkommastellen */
dp(x) := decimalplaces(x, 3)
/* stelle x in wissenschaftlicher Notation dar */
SCI(x) := scientific_notation(x, 2)
/* Rundungsgenauigkeit bei numerischen Antworttests */
epsilon: 0.001

/* gebe Wahrscheinlichkeiten vor */
probS: dp(rand_with_step(0.80, 0.90, 0.01))
probWS: dp(rand_with_step(0.85, 0.95, 0.01))
probWSc: dp(rand_with_step(0.05, 0.15, 0.01))

/* berechne die &#252;brigen Wahrscheinlichkeiten */
probSc: dp(1 - probS)
probWcS: dp(1 - probWS)
probWcSc: dp(1 - probWSc)
probSW: probWS * probS / (probWS * probS + probWSc * probSc)
probScW: probWSc * probSc / (probWS * probS + probWSc * probSc)
probSWc: probWcS * probS / (probWcS * probS + probWcSc * probSc)
probScWc: probWcSc * probSc / (probWcS * probS + probWcSc * probSc)

/* Formulierung der Aufgabenstellungen */
listB1: [1, "um eine Spammail", probSW, "P(S|W)", "P(W|S)P(S)", probWS, probS]
listB2: [2, "nicht um eine Spammail", probScW, "P(S^c|W)", "P(W|S^c)P(S^c)", probWSc, probSc]
listC1: [1, "um eine Spammail", probSWc, "P(S|W^c)", "P(W^c|S)P(S)", probWcS, probS]
listC2: [2, "nicht um eine Spammail", probScWc, "P(S^c|W^c)", "P(W^c|S^c)P(S^c)", probWcSc, probSc]

/* Auswahl der Aufgabenstellungen */
listB: rand([listB1, listB2])
listC: rand([listC1, listC2])

/* L&#246;sungen */
TansA1: probS
TansA2: probSc
TansA3: probWS
TansA4: probWcS
TansA5: probWSc
TansA6: probWcSc
TansA: [TansA1, TansA2, TansA3, TansA4, TansA5, TansA6]
TansB: listB[3]
TansC: listC[3]]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(P(S)={@probS@}\), \(P(W|S)={@probWS@}\), \(P(W|S^c)={@probWSc@}\), B={@listB[1]@}, C={@listC[1]@}</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ansA1</name>
            <type>numerical</type>
            <tans>TansA1</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA2</name>
            <type>numerical</type>
            <tans>TansA2</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA3</name>
            <type>numerical</type>
            <tans>TansA3</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA4</name>
            <type>numerical</type>
            <tans>TansA4</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA5</name>
            <type>numerical</type>
            <tans>TansA5</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA6</name>
            <type>numerical</type>
            <tans>TansA6</tans>
            <boxsize>5</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansB</name>
            <type>numerical</type>
            <tans>TansB</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansC</name>
            <type>numerical</type>
            <tans>TansC</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prtA</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>ansA: [ansA1, ansA2, ansA3, ansA4, ansA5, ansA6]</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>every(probabilityp,ansA)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prtA-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtA-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Nicht alle Tabelleneintr&#228;ge liegen im Intervall \([0,1]\). Diese Werte k&#246;nnen daher nicht als Wahrscheinlichkeiten interpretiert werden.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Not all table entries are in interval \([0,1]\). These values cannot therefore be interpreted as probabilities.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA</sans>
                <tans>TansA</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben alle Wahrscheinlichkeiten richtig angegeben.
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-30a');
show('oer-aufgabe-30b');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have all the probabilities correctly stated.
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-30a');
show('oer-aufgabe-30b');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtA-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA1</sans>
                <tans>TansA1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prtA-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prtA-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(S)\) ist falsch. Die Wahrscheinlichkeit, dass es sich bei einer eingehenden E-Mail um eine Spammail handelt, ist in der Aufgabenstellung gegeben.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(S)\) is wrong. The probability that an incoming email is a spammail is given in the task.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA2</sans>
                <tans>TansA2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prtA-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prtA-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(S^c)\) ist falsch. Berechnen Sie die Gegenwahrscheinlichkeit von \(P(S)\).
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(S^c)\) is wrong. Calculate the counter probability of \(P(S)\).
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA3</sans>
                <tans>TansA3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>5</truenextnode>
                <trueanswernote>prtA-5-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>5</falsenextnode>
                <falseanswernote>prtA-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(W|S)\) ist falsch. Die bedingte Wahrscheinlichkeit, dass eine E-Mail das Wort "Gewinn" enth&#228;lt, wenn sie eine Spammail ist, ist in der Aufgabenstellung gegeben.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(W|S)\) is wrong. The conditional probability that an email contains the word "winning" when it is a spammail is given in the task.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA4</sans>
                <tans>TansA4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>6</truenextnode>
                <trueanswernote>prtA-6-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>6</falsenextnode>
                <falseanswernote>prtA-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(W^c|S)\) ist falsch. Berechnen Sie die Gegenwahrscheinlichkeit von \(P(W|S)\).
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(W^c|S)\) is wrong. Calculate the counter probability of \(P(W|S)\).
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>6</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA5</sans>
                <tans>TansA5</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>7</truenextnode>
                <trueanswernote>prtA-7-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>7</falsenextnode>
                <falseanswernote>prtA-7-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(W|S^c)\) ist falsch. Die bedingte Wahrscheinlichkeit, dass eine E-Mail das Wort "Gewinn" enth&#228;lt, wenn sie keine Spammail ist, ist in der Aufgabenstellung gegeben.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(W|S^c)\) is wrong. The conditional probability that an email contains the word "winning" if it is not a spammail is given in the task.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>7</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA6</sans>
                <tans>TansA6</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-8-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtA-8-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Eintrag \(P(W^c|S^c)\) ist falsch. Berechnen Sie die Gegenwahrscheinlichkeit von \(P(W|S^c)\).
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The entry \(P(W^c|S^c)\) is wrong. Calculate the counter probability of \(P(W|S^c)\).
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtB</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[setB: set(dp(probSW), dp(probScW), dp(probSWc), dp(probScWc))

textB: if dp(ansB) = dp(probSW) then "P(S|W)"
elseif dp(ansB) = dp(probScW) then "P(S^c|W)"
elseif dp(ansB) = dp(probSWc) then "P(S|W^c)"
elseif dp(ansB) = dp(probScWc) then "P(S^c|W^c)"]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>probabilityp(ansB)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prtB-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtB-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihr Ergebnis liegt nicht im Intervall \([0,1]\) und kann daher nicht als Wahrscheinlichkeit interpretiert werden.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your result is not in interval \([0,1]\) and therefore cannot be interpreted as probability.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ansB</sans>
                <tans>TansB</tans>
                <testoptions>epsilon</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben die gesuchte Wahrscheinlichkeit \({@listB[4]@}\) richtig berechnet. Eine m&#246;gliche Rechnung mit dem Satz von Bayes sieht vollst&#228;ndig aufgeschrieben folgenderma&#223;en aus:
\begin{align*}
{@listB[4]@}
&amp; =\frac{{@listB[5]@}}{P(W)} \\
&amp; =\frac{{@listB[5]@}}{P(W|S)P(S)+P(W|S^c)P(S^c)} \\
&amp; =\frac{{@dp(listB[6])@}\cdot{@dp(listB[7])@}}{{@dp(probWS)@}\cdot{@dp(probS)@}+{@dp(probWSc)@}\cdot{@dp(probSc)@}} \\
&amp; \approx{@dp(listB[3])@}.
\end{align*}
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-30b');
show('oer-aufgabe-30c');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have the desired probability \({@listB[4]@}\) correctly calculated. A possible bill with the set of Bayes looks fully written as follows:
\begin{align*}
{@listB[4]@}
&amp; =\frac{{@listB[5]@}}{P(W)} \\
&amp; =\frac{{@listB[5]@}}{P(W|S)P(S)+P(W|S^c)P(S^c)} \\
&amp; =\frac{{@dp(listB[6])@}\cdot{@dp(listB[7])@}}{{@dp(probWS)@}\cdot{@dp(probS)@}+{@dp(probWSc)@}\cdot{@dp(probSc)@}} \\
&amp; \approx{@dp(listB[3])@}.
\end{align*}
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-30b');
show('oer-aufgabe-30c');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtB-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>member(dp(ansB),setB)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben vermutlich die bedingte Wahrscheinlichkeit \({@textB@}\) berechnet, gesucht ist aber \({@listB[4]@}\).
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You probably have the conditional probability \({@textB@}\) calculated, but searched \({@listB[4]@}\).
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtB-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Gesucht ist die bedingte Wahrscheinlichkeit \({@listB[4]@}\). Verwenden Sie zu deren Berechnung den Satz von Bayes.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The demand is the conditional probability \({@listB[4]@}\). Use the set of Bayes for their calculation.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtC</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[setC: set(dp(probSW), dp(probScW), dp(probSWc), dp(probScWc))

textC: if dp(ansC) = dp(probSW) then "P(S|W)"
elseif dp(ansC) = dp(probScW) then "P(S^c|W)"
elseif dp(ansC) = dp(probSWc) then "P(S|W^c)"
elseif dp(ansC) = dp(probScWc) then "P(S^c|W^c)"]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>probabilityp(ansC)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prtC-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtC-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihr Ergebnis liegt nicht im Intervall \([0,1]\) und kann daher nicht als Wahrscheinlichkeit interpretiert werden.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your result is not in interval \([0,1]\) and therefore cannot be interpreted as probability.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ansC</sans>
                <tans>TansC</tans>
                <testoptions>epsilon</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtC-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben die gesuchte Wahrscheinlichkeit \({@listC[4]@}\) richtig berechnet. Eine m&#246;gliche Rechnung mit dem Satz von Bayes sieht vollst&#228;ndig aufgeschrieben folgenderma&#223;en aus:
\begin{align*}
{@listC[4]@}
&amp; =\frac{{@listC[5]@}}{P(W^c)} \\
&amp; =\frac{{@listC[5]@}}{P(W^c|S)P(S)+P(W^c|S^c)P(S^c)} \\
&amp; =\frac{{@dp(listC[6])@}\cdot{@dp(listC[7])@}}{{@dp(probWcS)@}\cdot{@dp(probS)@}+{@dp(probWcSc)@}\cdot{@dp(probSc)@}} \\
&amp; \approx{@dp(listC[3])@}.
\end{align*}
    </p>
    <hr/>
    <p>
        Sie haben die Aufgabe erfolgreich gel&#246;st! Jetzt sollten Sie mit dem Satz von Bayes sicher rechnen k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have the desired probability \({@listC[4]@}\) correctly calculated. A possible bill with the set of Bayes looks fully written as follows:
\begin{align*}
{@listC[4]@}
&amp; =\frac{{@listC[5]@}}{P(W^c)} \\
&amp; =\frac{{@listC[5]@}}{P(W^c|S)P(S)+P(W^c|S^c)P(S^c)} \\
&amp; =\frac{{@dp(listC[6])@}\cdot{@dp(listC[7])@}}{{@dp(probWcS)@}\cdot{@dp(probS)@}+{@dp(probWcSc)@}\cdot{@dp(probSc)@}} \\
&amp; \approx{@dp(listC[3])@}.
\end{align*}
    </p>
    <hr/>
    <p>
        You have successfully solved the task! Now you should be able to count on the set of Bayes.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtC-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>member(dp(ansC),setC)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtC-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben vermutlich die bedingte Wahrscheinlichkeit \({@textC@}\) berechnet, gesucht ist aber \({@listC[4]@}\).
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You probably have the conditional probability \({@textC@}\) calculated, but searched \({@listC[4]@}\).
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtC-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Gesucht ist die bedingte Wahrscheinlichkeit \({@listC[4]@}\). Verwenden Sie zu deren Berechnung den Satz von Bayes.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The demand is the conditional probability \({@listC[4]@}\). Use the set of Bayes for their calculation.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>46</deployedseed>
        <deployedseed>719862284</deployedseed>
        <deployedseed>90133813</deployedseed>
        <deployedseed>1274194721</deployedseed>
        <deployedseed>855476367</deployedseed>
        <deployedseed>1172057739</deployedseed>
        <deployedseed>1831690210</deployedseed>
        <deployedseed>780847783</deployedseed>
        <deployedseed>2088320236</deployedseed>
        <deployedseed>789261352</deployedseed>
        <deployedseed>1576833348</deployedseed>
    </question>
    

</quiz>