<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2023-03-07 Ziegenproblem</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <div id="oer-aufgabe-8-info">
        <details>
            <summary>
                Inhalt und Umfang dieser Aufgabe
            </summary>
            <p>
                In dieser Aufgabe lernen Sie, das Ziegenproblem stochastisch zu l&#246;sen. Sie besteht aus den folgenden drei Teilaufgaben.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    stellen Sie die Situation beim Ziegenproblem mit einem Baumdiagram dar.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    berechnen Sie die Gewinnwahrscheinlichkeit in Abh&#228;ngigkeit von der gew&#228;hlten Spielstrategie.
                </li>
                <li>
                    In
                    <b>
                        (c)
                    </b>
                    vergleichen Sie verschiedene Spielstrategien und entscheiden, welche davon die Gewinnwahrscheinlichkeit maximiert.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-8">
        <p>
            Beim
            <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='auch bekannt unter "Drei-T&#252;ren-Problem"'>
                Ziegenproblem
            </span>
            geht es um die folgende Situation: In einer Spielshow muss sich die Kandidatin f&#252;r eine von drei T&#252;ren entscheiden. Hinter einer T&#252;r befindet sich der Gewinn, beispielsweise ein Luxusauto, hinter den beiden anderen T&#252;ren befindet sich jeweils eine Ziege. Der Gewinn wird von den Produzenten der Show zuf&#228;llig hinter einer der drei T&#252;ren platziert, und nur der Showmaster wei&#223;, hinter welcher T&#252;r er sich befindet. Nach der Wahl der Kandidatin &#246;ffnet er zun&#228;chst eine der beiden nicht gew&#228;hlten T&#252;ren, hinter der sich auf jeden Fall eine Ziege befindet. Anschlie&#223;end stellt er die Kandidatin vor die Wahl, bei der zuerst gew&#228;hlten T&#252;r zu bleiben oder zur verbleibenden T&#252;r zu wechseln. Das dahinter befindliche Objekt geht schlie&#223;lich in den Besitz der Kandidatin &#252;ber. Wie sollte sich die Kandidatin entscheiden, d.h. mit welcher Spielstrategie besteht die gr&#246;&#223;te Chance auf den Gewinn?
        </p>
        <hr/>
    </div>
    <div id="oer-aufgabe-8a">
        <p>
            <b>
                (a)
            </b>
            Zur L&#246;sung des Problems stellen wir alle m&#246;glichen Spielverl&#228;ufe in einem Baumdiagramm dar.
            <ul>
                <li>
                    Im ersten Schritt w&#228;hlt die Kandidatin eine der drei T&#252;ren aus. Da der Gewinn zuvor zuf&#228;llig platziert wurde, w&#228;hlt sie ihre T&#252;r zuf&#228;llig mit Wahrscheinlichkeit \(1/3\) aus. Wir bezeichnen die T&#252;r, hinter der sich der Gewinn befindet, mit \(G\) und die T&#252;ren, hinter denen sich die Ziegen befinden, mit \(Z_1\) bzw. \(Z_2\).
                </li>
                <li>
                    Im zweiten Schritt &#246;ffnet der Showmaster eine der beiden nicht gew&#228;hlten T&#252;ren, hinter der sich auf jeden Fall eine Ziege befindet. Falls die Kandidatin die T&#252;r \(G\) ausgew&#228;hlt hat, wird der Showmaster mit Wahrscheinlichkeit \(p\in[0,1]\) die T&#252;r \(Z_1\) oder mit Wahrscheinlichkeit \(1-p\) die T&#252;r \(Z_2\) &#246;ffnen. Falls die Kandidatin eine der beiden T&#252;ren mit einer Ziege dahinter ausgew&#228;hlt hat, wird der Showmaster mit Wahrscheinlichkeit \(1\) die jeweils andere T&#252;r mit einer Ziege dahinter &#246;ffnen.
                </li>
                <li>
                    Im dritten Schritt bleibt die Kandidatin mit einer von ihr vorab gew&#228;hlten Wahrscheinlichkeit \(q\in[0,1]\) bei der zuerst gew&#228;hlten T&#252;r oder wechselt mit Wahrscheinlichkeit \(1-q\) zur anderen noch verschlossenen T&#252;r. Der Wert von \(q\) definiert somit ihre Spielstrategie.
                </li>
            </ul>
            Vervollst&#228;ndigen Sie das Baumdiagramm, indem Sie die blauen, rautenf&#246;rmigen Felder und die roten, quadratischen Felder auf die richtigen Knotenpunkte verschieben.
        </p>
        <p>
            [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-8-JSXGraph">
        [[jsxgraph
input-ref-ansA11="ans2ARef"
input-ref-ansA12="ans2BRef"
input-ref-ansA13="ans2CRef"
input-ref-ansA14="ans2DRef"
input-ref-ansA21="ans3ARef"
input-ref-ansA22="ans3BRef"
input-ref-ansA23="ans3CRef"
input-ref-ansA24="ans3DRef"
input-ref-ansA25="ans3ERef"
input-ref-ansA26="ans3FRef"
input-ref-ansA27="ans3GRef"
input-ref-ansA28="ans3HRef"
]]
var board = JXG.JSXGraph.initBoard(divid, {
boundingbox: [-1.1, 1.1, 1.1, -1.1],
axis: false,
drag: { enabled: true },
grid: false,
showCopyright: false,
showFullscreen: false,
showInfobox: false,
showNavigation: false,
showScreenshot: false
});




//Schritt 0: Spielbeginn
var point0 = board.create('point', [0, 1], { color: 'grey', fixed: true, highlight: false, name: "", opacity: .5, showInfobox: false, size: 5 });




//Schritt 1: w&#228;hle eine der drei T&#252;ren G, Z1, Z2
//erzeuge Punkte f&#252;r die Knoten im Entscheidungsbaum und beschrifte diese mit Ereignissen
var pointOption1 = {
color: 'grey',
face: 'circle',
fixed: true,
highlight: false,
name: "",
opacity: .5,
showInfobox: false,
size: 15
};
var point1A = board.create('point', [-.5, .5], pointOption1);
var point1B = board.create('point', [0, .5], pointOption1);
var point1C = board.create('point', [.5, .5], pointOption1);
var textPoint1A = board.create('text', [-.02, 0, "G"], { anchor: point1A, fixed: true, highlight: false });
var textPoint1B = board.create('text', [-.025, 0, "Z1"], { anchor: point1B, fixed: true, highlight: false });
var textPoint1C = board.create('text', [-.025, 0, "Z2"], { anchor: point1C, fixed: true, highlight: false });
//erzeuge Segmente f&#252;r die Kanten im Entscheidungsbaum und beschrifte diese mit Wahrscheinlichkeiten
var segment01A = board.create('segment', [point0, point1A], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment01B = board.create('segment', [point0, point1B], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment01C = board.create('segment', [point0, point1C], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var textSegment01A = board.create('text', [0, 0, "1/3"], { anchor: segment01A, fixed: true, highlight: false });
var textSegment01B = board.create('text', [0, 0, "1/3"], { anchor: segment01B, fixed: true, highlight: false });
var textSegment01C = board.create('text', [0, 0, "1/3"], { anchor: segment01C, fixed: true, highlight: false });




//Schritt 2: Showmaster &#246;ffnet eine T&#252;r mit einer Ziege dahinter
//erzeuge Punkte f&#252;r die Knoten im Entscheidungsbaum
var pointOption2 = {
color: 'blue',
face: 'diamond',
fixed: true,
highlight: false,
name: "",
opacity: .1,
showInfobox: false,
size: 20
};
var point2A = board.create('point', [-.75, 0], pointOption2);
var point2B = board.create('point', [-.25, 0], pointOption2);
var point2C = board.create('point', [.25, 0], pointOption2);
var point2D = board.create('point', [.75, 0], pointOption2);
//erzeuge Segmente f&#252;r die Kanten im Entscheidungsbaum und beschrifte diese mit bedingten Wahrscheinlichkeiten
var segment1A2A = board.create('segment', [point1A, point2A], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment1A2B = board.create('segment', [point1A, point2B], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment1B2C = board.create('segment', [point1B, point2C], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment1C2D = board.create('segment', [point1C, point2D], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var textSegment1A2A = board.create('text', [0, 0, "p"], { anchor: segment1A2A, fixed: true, highlight: false });
var textSegment1A2B = board.create('text', [0, 0, "1-p"], { anchor: segment1A2B, fixed: true, highlight: false });
var textSegment1B2C = board.create('text', [0, 0, "1"], { anchor: segment1B2C, fixed: true, highlight: false });
var textSegment1C2D = board.create('text', [0, 0, "1"], { anchor: segment1C2D, fixed: true, highlight: false });




//erzeuge bewegbare Punkte und beschrifte diese mit Ereignissen
var pointDragOption2 = {
attractors: [point2A, point2B, point2C, point2D],
attractorDistance: .1,
snatchDistance: .2,
color: 'blue',
face: 'diamond',
highlight: false,
name: "",
opacity: .5,
showInfobox: false,
size: 20
};
var pointDrag2A = board.create('point', [-3 / 4, -.7], pointDragOption2);
var pointDrag2B = board.create('point', [-1 / 4, -.7], pointDragOption2);
var pointDrag2C = board.create('point', [1 / 4, -.7], pointDragOption2);
var pointDrag2D = board.create('point', [3 / 4, -.7], pointDragOption2);
var textPointDrag2A = board.create('text', [-.025, 0, "Z1"], { anchor: pointDrag2A, fixed: true, highlight: false });
var textPointDrag2B = board.create('text', [-.025, 0, "Z1"], { anchor: pointDrag2B, fixed: true, highlight: false });
var textPointDrag2C = board.create('text', [-.025, 0, "Z2"], { anchor: pointDrag2C, fixed: true, highlight: false });
var textPointDrag2D = board.create('text', [-.025, 0, "Z2"], { anchor: pointDrag2D, fixed: true, highlight: false });




//Schritt 3: bleibe bei der zuerst gew&#228;hlten T&#252;r oder wechsle zur anderen noch verschlossenen T&#252;r
//erzeuge Punkte f&#252;r die Knoten im Entscheidungsbaum
var pointOption3 = {
color: 'red',
face: 'square',
fixed: true,
highlight: false,
name: "",
opacity: .1,
showInfobox: false,
size: 15
};
var point3A = board.create('point', [-7 / 8, -.5], pointOption3);
var point3B = board.create('point', [-5 / 8, -.5], pointOption3);
var point3C = board.create('point', [-3 / 8, -.5], pointOption3);
var point3D = board.create('point', [-1 / 8, -.5], pointOption3);
var point3E = board.create('point', [1 / 8, -.5], pointOption3);
var point3F = board.create('point', [3 / 8, -.5], pointOption3);
var point3G = board.create('point', [5 / 8, -.5], pointOption3);
var point3H = board.create('point', [7 / 8, -.5], pointOption3);
//erzeuge Segmente f&#252;r die Kanten im Entscheidungsbaum und beschrifte diese mit bedingten Wahrscheinlichkeiten
var segment2A3A = board.create('segment', [point2A, point3A], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2A3B = board.create('segment', [point2A, point3B], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2B3C = board.create('segment', [point2B, point3C], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2B3D = board.create('segment', [point2B, point3D], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2C3E = board.create('segment', [point2C, point3E], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2C3F = board.create('segment', [point2C, point3F], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2D3G = board.create('segment', [point2D, point3G], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2D3H = board.create('segment', [point2D, point3H], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var textSegment2A3A = board.create('text', [0, 0, "q"], { anchor: segment2A3A, fixed: true, highlight: false });
var textSegment2A3B = board.create('text', [0, 0, "1-q"], { anchor: segment2A3B, fixed: true, highlight: false });
var textSegment2B3C = board.create('text', [0, 0, "q"], { anchor: segment2B3C, fixed: true, highlight: false });
var textSegment2B3D = board.create('text', [0, 0, "1-q"], { anchor: segment2B3D, fixed: true, highlight: false });
var textSegment2C3E = board.create('text', [0, 0, "q"], { anchor: segment2C3E, fixed: true, highlight: false });
var textSegment2C3F = board.create('text', [0, 0, "1-q"], { anchor: segment2C3F, fixed: true, highlight: false });
var textSegment2D3G = board.create('text', [0, 0, "q"], { anchor: segment2D3G, fixed: true, highlight: false });
var textSegment2D3H = board.create('text', [0, 0, "1-q"], { anchor: segment2D3H, fixed: true, highlight: false });




//erzeuge bewegbare Punkte und beschrifte diese mit Ereignissen
var pointDragOption3 = {
attractors: [point3A, point3B, point3C, point3D, point3E, point3F, point3G, point3H],
attractorDistance: .1,
snatchDistance: .2,
color: 'red',
face: 'square',
highlight: false,
name: "",
opacity: .5,
showInfobox: false,
size: 15
};
var pointDrag3A = board.create('point', [-7 / 8, -.9], pointDragOption3);
var pointDrag3B = board.create('point', [-5 / 8, -.9], pointDragOption3);
var pointDrag3C = board.create('point', [-3 / 8, -.9], pointDragOption3);
var pointDrag3D = board.create('point', [-1 / 8, -.9], pointDragOption3);
var pointDrag3E = board.create('point', [1 / 8, -.9], pointDragOption3);
var pointDrag3F = board.create('point', [3 / 8, -.9], pointDragOption3);
var pointDrag3G = board.create('point', [5 / 8, -.9], pointDragOption3);
var pointDrag3H = board.create('point', [7 / 8, -.9], pointDragOption3);
var textPointDrag3A = board.create('text', [-.02, 0, "G"], { anchor: pointDrag3A, fixed: true, highlight: false });
var textPointDrag3B = board.create('text', [-.02, 0, "G"], { anchor: pointDrag3B, fixed: true, highlight: false });
var textPointDrag3C = board.create('text', [-.02, 0, "G"], { anchor: pointDrag3C, fixed: true, highlight: false });
var textPointDrag3D = board.create('text', [-.02, 0, "G"], { anchor: pointDrag3D, fixed: true, highlight: false });
var textPointDrag3E = board.create('text', [-.025, 0, "Z1"], { anchor: pointDrag3E, fixed: true, highlight: false });
var textPointDrag3F = board.create('text', [-.025, 0, "Z1"], { anchor: pointDrag3F, fixed: true, highlight: false });
var textPointDrag3G = board.create('text', [-.025, 0, "Z2"], { anchor: pointDrag3G, fixed: true, highlight: false });
var textPointDrag3H = board.create('text', [-.025, 0, "Z2"], { anchor: pointDrag3H, fixed: true, highlight: false });




//binde die bewegbaren Punkte an STACK
stack_jxg.bind_point(ans2ARef, pointDrag2A);
stack_jxg.bind_point(ans2BRef, pointDrag2B);
stack_jxg.bind_point(ans2CRef, pointDrag2C);
stack_jxg.bind_point(ans2DRef, pointDrag2D);
stack_jxg.bind_point(ans3ARef, pointDrag3A);
stack_jxg.bind_point(ans3BRef, pointDrag3B);
stack_jxg.bind_point(ans3CRef, pointDrag3C);
stack_jxg.bind_point(ans3DRef, pointDrag3D);
stack_jxg.bind_point(ans3ERef, pointDrag3E);
stack_jxg.bind_point(ans3FRef, pointDrag3F);
stack_jxg.bind_point(ans3GRef, pointDrag3G);
stack_jxg.bind_point(ans3HRef, pointDrag3H);




board.update();
[[/jsxgraph]]
    </div>
    <div style="display:none;">
        <p>
            [[input:ansA11]] [[validation:ansA11]]
        </p>
        <p>
            [[input:ansA12]] [[validation:ansA12]]
        </p>
        <p>
            [[input:ansA13]] [[validation:ansA13]]
        </p>
        <p>
            [[input:ansA14]] [[validation:ansA14]]
        </p>
        <p>
            [[input:ansA21]] [[validation:ansA21]]
        </p>
        <p>
            [[input:ansA22]] [[validation:ansA22]]
        </p>
        <p>
            [[input:ansA23]] [[validation:ansA23]]
        </p>
        <p>
            [[input:ansA24]] [[validation:ansA24]]
        </p>
        <p>
            [[input:ansA25]] [[validation:ansA25]]
        </p>
        <p>
            [[input:ansA26]] [[validation:ansA26]]
        </p>
        <p>
            [[input:ansA27]] [[validation:ansA27]]
        </p>
        <p>
            [[input:ansA28]] [[validation:ansA28]]
        </p>
    </div>
    <div id="oer-aufgabe-8b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            Berechnen Sie mithilfe des Multiplikationssatzes die Gewinnwahrscheinlichkeit in Abh&#228;ngigkeit von \(p\) und \(q\). Zur Erinnerung ist oben noch einmal das Baumdiagramm aus (a) abgebildet.
        </p>
        <p>
            Die Gewinnwahrscheinlichkeit betr&#228;gt [[input:ansB]]. [[validation:ansB]] [[feedback:prtB]]
        </p>
    </div>
    <div id="oer-aufgabe-8c" style="display:none;">
        <p>
            <b>
                (c)
            </b>
            In der vorherigen Teilaufgabe haben wir die Gewinnwahrscheinlichkeit \(\frac{2-q}{3}\) berechnet, wobei \(q\) die Wahrscheinlichkeit ist, mit der die Kandidatin bei der zuerst gew&#228;hlten T&#252;r bleibt. Welche dieser drei Spielstrategien besitzt somit die gr&#246;&#223;te Gewinnwahrscheinlichkeit?
            <ul>
                <li>
                    Strategie 1: Die Kandidatin bleibt auf jeden Fall bei der zuerst gew&#228;hlten T&#252;r.
                </li>
                <li>
                    Strategie 2: Die Kandidatin wechselt auf jeden Fall zur anderen noch verschlossenen T&#252;r.
                </li>
                <li>
                    Strategie 3: Die Kandidatin wirft eine faire M&#252;nze und entscheidet sich je nach Ergebnis f&#252;r eine der beiden verbleibenden T&#252;ren.
                </li>
            </ul>
        </p>
        <p>
            Die gr&#246;&#223;te Gewinnwahrscheinlichkeit besitzt [[input:ansC]]. [[validation:ansC]] [[feedback:prtC]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
<span class="multilang" lang="en">
    <div id="oer-aufgabe-8-info">
        <details>
            <summary>
                Content and scope of this task
            </summary>
            <p>
                In this task you will learn to solve the goat problem stochastic. It consists of the following three subtasks.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    represent the situation in the goat problem with a tree diagram.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    calculate the probability of winning depending on the selected game strategy.
                </li>
                <li>
                    In
                    <b>
                        (c)
                    </b>
                    compare different game strategies and decide which of them maximizes profit probability.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-8">
        <p>
            The
            <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='auch bekannt unter "Drei-T&#252;ren-Problem"'>
                Goat problem
            </span>
            the following situation: In a game show, the candidate must choose one of three doors. Behind a door there is a win, for example a luxury car, behind the two other doors there is a goat. The profit is randomly placed behind one of the three doors by the producers of the show, and only the showmaster knows which door he is behind. After the election of the candidate, he first opens one of the two unelected doors behind which there is definitely a goat. He then introduces the candidate to the choice of staying at the first chosen door or changing to the remaining door. The object behind it finally goes into the possession of the candidate. How should the candidate decide, i.e. what game strategy is the biggest chance of winning?
        </p>
        <hr/>
    </div>
    <div id="oer-aufgabe-8a">
        <p>
            <b>
                (a)
            </b>
            To solve the problem, we represent all possible game patterns in a tree diagram.
            <ul>
                <li>
                    In the first step, the candidate selects one of the three doors. Since the win was previously randomly placed, she randomly selects her door \(1/3\) out. We call the door behind which the profit is, with \(G\) and the doors behind which the goats are located, with \(Z_1\) and \(Z_2\).
                </li>
                <li>
                    In the second step, the showmaster opens one of the two unelected doors, behind which there is definitely a goat. If the candidate is the door \(G\) the showmaster will be selected with probability \(p\in[0,1]\) the door \(Z_1\) or with probability \(1-p\) the door \(Z_2\) open. If the candidate has selected one of the two doors with a goat behind it, the showmaster will probably \(1\) open the other door with a goat behind it.
                </li>
                <li>
                    In the third step, the candidate remains with a pre-elected probability \(q\in[0,1]\) at the first chosen door or change with probability \(1-q\) to the other door still closed. The value of \(q\) thus defines their game strategy.
                </li>
            </ul>
            Complete the tree chart by moving the blue, diamond-shaped fields and the red, square fields to the right nodes.
        </p>
        <p>
            [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-8-JSXGraph">
        [[jsxgraph
input-ref-ansA11="ans2ARef"
input-ref-ansA12="ans2BRef"
input-ref-ansA13="ans2CRef"
input-ref-ansA14="ans2DRef"
input-ref-ansA21="ans3ARef"
input-ref-ansA22="ans3BRef"
input-ref-ansA23="ans3CRef"
input-ref-ansA24="ans3DRef"
input-ref-ansA25="ans3ERef"
input-ref-ansA26="ans3FRef"
input-ref-ansA27="ans3GRef"
input-ref-ansA28="ans3HRef"
]]
var board = JXG.JSXGraph.initBoard(divid, {
boundingbox: [-1.1, 1.1, 1.1, -1.1],
axis: false,
drag: { enabled: true },
grid: false,
showCopyright: false,
showFullscreen: false,
showInfobox: false,
showNavigation: false,
showScreenshot: false
});




//Schritt 0: Spielbeginn
var point0 = board.create('point', [0, 1], { color: 'grey', fixed: true, highlight: false, name: "", opacity: .5, showInfobox: false, size: 5 });




//Schritt 1: w&#228;hle eine der drei T&#252;ren G, Z1, Z2
//erzeuge Punkte f&#252;r die Knoten im Entscheidungsbaum und beschrifte diese mit Ereignissen
var pointOption1 = {
color: 'grey',
face: 'circle',
fixed: true,
highlight: false,
name: "",
opacity: .5,
showInfobox: false,
size: 15
};
var point1A = board.create('point', [-.5, .5], pointOption1);
var point1B = board.create('point', [0, .5], pointOption1);
var point1C = board.create('point', [.5, .5], pointOption1);
var textPoint1A = board.create('text', [-.02, 0, "G"], { anchor: point1A, fixed: true, highlight: false });
var textPoint1B = board.create('text', [-.025, 0, "Z1"], { anchor: point1B, fixed: true, highlight: false });
var textPoint1C = board.create('text', [-.025, 0, "Z2"], { anchor: point1C, fixed: true, highlight: false });
//erzeuge Segmente f&#252;r die Kanten im Entscheidungsbaum und beschrifte diese mit Wahrscheinlichkeiten
var segment01A = board.create('segment', [point0, point1A], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment01B = board.create('segment', [point0, point1B], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment01C = board.create('segment', [point0, point1C], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var textSegment01A = board.create('text', [0, 0, "1/3"], { anchor: segment01A, fixed: true, highlight: false });
var textSegment01B = board.create('text', [0, 0, "1/3"], { anchor: segment01B, fixed: true, highlight: false });
var textSegment01C = board.create('text', [0, 0, "1/3"], { anchor: segment01C, fixed: true, highlight: false });




//Schritt 2: Showmaster &#246;ffnet eine T&#252;r mit einer Ziege dahinter
//erzeuge Punkte f&#252;r die Knoten im Entscheidungsbaum
var pointOption2 = {
color: 'blue',
face: 'diamond',
fixed: true,
highlight: false,
name: "",
opacity: .1,
showInfobox: false,
size: 20
};
var point2A = board.create('point', [-.75, 0], pointOption2);
var point2B = board.create('point', [-.25, 0], pointOption2);
var point2C = board.create('point', [.25, 0], pointOption2);
var point2D = board.create('point', [.75, 0], pointOption2);
//erzeuge Segmente f&#252;r die Kanten im Entscheidungsbaum und beschrifte diese mit bedingten Wahrscheinlichkeiten
var segment1A2A = board.create('segment', [point1A, point2A], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment1A2B = board.create('segment', [point1A, point2B], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment1B2C = board.create('segment', [point1B, point2C], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment1C2D = board.create('segment', [point1C, point2D], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var textSegment1A2A = board.create('text', [0, 0, "p"], { anchor: segment1A2A, fixed: true, highlight: false });
var textSegment1A2B = board.create('text', [0, 0, "1-p"], { anchor: segment1A2B, fixed: true, highlight: false });
var textSegment1B2C = board.create('text', [0, 0, "1"], { anchor: segment1B2C, fixed: true, highlight: false });
var textSegment1C2D = board.create('text', [0, 0, "1"], { anchor: segment1C2D, fixed: true, highlight: false });




//erzeuge bewegbare Punkte und beschrifte diese mit Ereignissen
var pointDragOption2 = {
attractors: [point2A, point2B, point2C, point2D],
attractorDistance: .1,
snatchDistance: .2,
color: 'blue',
face: 'diamond',
highlight: false,
name: "",
opacity: .5,
showInfobox: false,
size: 20
};
var pointDrag2A = board.create('point', [-3 / 4, -.7], pointDragOption2);
var pointDrag2B = board.create('point', [-1 / 4, -.7], pointDragOption2);
var pointDrag2C = board.create('point', [1 / 4, -.7], pointDragOption2);
var pointDrag2D = board.create('point', [3 / 4, -.7], pointDragOption2);
var textPointDrag2A = board.create('text', [-.025, 0, "Z1"], { anchor: pointDrag2A, fixed: true, highlight: false });
var textPointDrag2B = board.create('text', [-.025, 0, "Z1"], { anchor: pointDrag2B, fixed: true, highlight: false });
var textPointDrag2C = board.create('text', [-.025, 0, "Z2"], { anchor: pointDrag2C, fixed: true, highlight: false });
var textPointDrag2D = board.create('text', [-.025, 0, "Z2"], { anchor: pointDrag2D, fixed: true, highlight: false });




//Schritt 3: bleibe bei der zuerst gew&#228;hlten T&#252;r oder wechsle zur anderen noch verschlossenen T&#252;r
//erzeuge Punkte f&#252;r die Knoten im Entscheidungsbaum
var pointOption3 = {
color: 'red',
face: 'square',
fixed: true,
highlight: false,
name: "",
opacity: .1,
showInfobox: false,
size: 15
};
var point3A = board.create('point', [-7 / 8, -.5], pointOption3);
var point3B = board.create('point', [-5 / 8, -.5], pointOption3);
var point3C = board.create('point', [-3 / 8, -.5], pointOption3);
var point3D = board.create('point', [-1 / 8, -.5], pointOption3);
var point3E = board.create('point', [1 / 8, -.5], pointOption3);
var point3F = board.create('point', [3 / 8, -.5], pointOption3);
var point3G = board.create('point', [5 / 8, -.5], pointOption3);
var point3H = board.create('point', [7 / 8, -.5], pointOption3);
//erzeuge Segmente f&#252;r die Kanten im Entscheidungsbaum und beschrifte diese mit bedingten Wahrscheinlichkeiten
var segment2A3A = board.create('segment', [point2A, point3A], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2A3B = board.create('segment', [point2A, point3B], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2B3C = board.create('segment', [point2B, point3C], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2B3D = board.create('segment', [point2B, point3D], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2C3E = board.create('segment', [point2C, point3E], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2C3F = board.create('segment', [point2C, point3F], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2D3G = board.create('segment', [point2D, point3G], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var segment2D3H = board.create('segment', [point2D, point3H], { color: 'black', fixed: true, highlight: false, strokeOpacity: .1 });
var textSegment2A3A = board.create('text', [0, 0, "q"], { anchor: segment2A3A, fixed: true, highlight: false });
var textSegment2A3B = board.create('text', [0, 0, "1-q"], { anchor: segment2A3B, fixed: true, highlight: false });
var textSegment2B3C = board.create('text', [0, 0, "q"], { anchor: segment2B3C, fixed: true, highlight: false });
var textSegment2B3D = board.create('text', [0, 0, "1-q"], { anchor: segment2B3D, fixed: true, highlight: false });
var textSegment2C3E = board.create('text', [0, 0, "q"], { anchor: segment2C3E, fixed: true, highlight: false });
var textSegment2C3F = board.create('text', [0, 0, "1-q"], { anchor: segment2C3F, fixed: true, highlight: false });
var textSegment2D3G = board.create('text', [0, 0, "q"], { anchor: segment2D3G, fixed: true, highlight: false });
var textSegment2D3H = board.create('text', [0, 0, "1-q"], { anchor: segment2D3H, fixed: true, highlight: false });




//erzeuge bewegbare Punkte und beschrifte diese mit Ereignissen
var pointDragOption3 = {
attractors: [point3A, point3B, point3C, point3D, point3E, point3F, point3G, point3H],
attractorDistance: .1,
snatchDistance: .2,
color: 'red',
face: 'square',
highlight: false,
name: "",
opacity: .5,
showInfobox: false,
size: 15
};
var pointDrag3A = board.create('point', [-7 / 8, -.9], pointDragOption3);
var pointDrag3B = board.create('point', [-5 / 8, -.9], pointDragOption3);
var pointDrag3C = board.create('point', [-3 / 8, -.9], pointDragOption3);
var pointDrag3D = board.create('point', [-1 / 8, -.9], pointDragOption3);
var pointDrag3E = board.create('point', [1 / 8, -.9], pointDragOption3);
var pointDrag3F = board.create('point', [3 / 8, -.9], pointDragOption3);
var pointDrag3G = board.create('point', [5 / 8, -.9], pointDragOption3);
var pointDrag3H = board.create('point', [7 / 8, -.9], pointDragOption3);
var textPointDrag3A = board.create('text', [-.02, 0, "G"], { anchor: pointDrag3A, fixed: true, highlight: false });
var textPointDrag3B = board.create('text', [-.02, 0, "G"], { anchor: pointDrag3B, fixed: true, highlight: false });
var textPointDrag3C = board.create('text', [-.02, 0, "G"], { anchor: pointDrag3C, fixed: true, highlight: false });
var textPointDrag3D = board.create('text', [-.02, 0, "G"], { anchor: pointDrag3D, fixed: true, highlight: false });
var textPointDrag3E = board.create('text', [-.025, 0, "Z1"], { anchor: pointDrag3E, fixed: true, highlight: false });
var textPointDrag3F = board.create('text', [-.025, 0, "Z1"], { anchor: pointDrag3F, fixed: true, highlight: false });
var textPointDrag3G = board.create('text', [-.025, 0, "Z2"], { anchor: pointDrag3G, fixed: true, highlight: false });
var textPointDrag3H = board.create('text', [-.025, 0, "Z2"], { anchor: pointDrag3H, fixed: true, highlight: false });




//binde die bewegbaren Punkte an STACK
stack_jxg.bind_point(ans2ARef, pointDrag2A);
stack_jxg.bind_point(ans2BRef, pointDrag2B);
stack_jxg.bind_point(ans2CRef, pointDrag2C);
stack_jxg.bind_point(ans2DRef, pointDrag2D);
stack_jxg.bind_point(ans3ARef, pointDrag3A);
stack_jxg.bind_point(ans3BRef, pointDrag3B);
stack_jxg.bind_point(ans3CRef, pointDrag3C);
stack_jxg.bind_point(ans3DRef, pointDrag3D);
stack_jxg.bind_point(ans3ERef, pointDrag3E);
stack_jxg.bind_point(ans3FRef, pointDrag3F);
stack_jxg.bind_point(ans3GRef, pointDrag3G);
stack_jxg.bind_point(ans3HRef, pointDrag3H);




board.update();
[[/jsxgraph]]
    </div>
    <div style="display:none;">
        <p>
            [[input:ansA11]] [[validation:ansA11]]
        </p>
        <p>
            [[input:ansA12]] [[validation:ansA12]]
        </p>
        <p>
            [[input:ansA13]] [[validation:ansA13]]
        </p>
        <p>
            [[input:ansA14]] [[validation:ansA14]]
        </p>
        <p>
            [[input:ansA21]] [[validation:ansA21]]
        </p>
        <p>
            [[input:ansA22]] [[validation:ansA22]]
        </p>
        <p>
            [[input:ansA23]] [[validation:ansA23]]
        </p>
        <p>
            [[input:ansA24]] [[validation:ansA24]]
        </p>
        <p>
            [[input:ansA25]] [[validation:ansA25]]
        </p>
        <p>
            [[input:ansA26]] [[validation:ansA26]]
        </p>
        <p>
            [[input:ansA27]] [[validation:ansA27]]
        </p>
        <p>
            [[input:ansA28]] [[validation:ansA28]]
        </p>
    </div>
    <div id="oer-aufgabe-8b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            Calculate profit probability depending on the multiplication rate \(p\) and \(q\). For memory, the tree diagram from (a) is shown again at the top.
        </p>
        <p>
            The profit probability is [[input:ansB]]. [[validation:ansB]] [[feedback:prtB]]
        </p>
    </div>
    <div id="oer-aufgabe-8c" style="display:none;">
        <p>
            <b>
                (c)
            </b>
            In the previous subtask we have the probability of profit \(\frac{2-q}{3}\) where \(q\) the probability with which the candidate remains at the first chosen door. Which of these three game strategies therefore has the greatest probability of winning?
            <ul>
                <li>
                    Strategy 1: The candidate will definitely stay at the first door.
                </li>
                <li>
                    Strategy 2: The candidate will definitely move to the other door still closed.
                </li>
                <li>
                    Strategy 3: The candidate casts a fair coin and decides for one of the two remaining doors, depending on the outcome.
                </li>
            </ul>
        </p>
        <p>
            The largest profit probability [[input:ansC]]. [[validation:ansC]] [[feedback:prtC]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="moodle_auto_format">
            <text/>
        </generalfeedback>
        <defaultgrade>3.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
Ziegenproblem

wurde entwickelt von
  
Christian M&#252;ller <c.mueller(at)hhu.de>

nach einer Idee von

Axel B&#252;cher <axel.buecher(at)hhu.de>
Peter Kern <kern(at)hhu.de>

an der Heinrich-Heine-Universit&#228;t D&#252;sseldorf.

Dieses Werk ist lizenziert unter einer Creative Commons Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie https://creativecommons.org/licenses/by-sa/4.0/.

SPDX-License-Identifier: CC-BY-SA-4.0

Technische Informationen:

Diese Aufgabe bindet das Skript 'stackselbstlern.js' von Michael Kallweit f&#252;r die Aufgabennavigation ein.
*/

/* korrekte Koordinaten der Felder */
TAnsSet2Z1: set([-0.75, 0], [0.75, 0])
TAnsSet2Z2: set([-0.25, 0], [0.25, 0])
TAnsSet3G: set([-0.875, -0.5], [-0.375, -0.5], [0.375, -0.5], [0.875, -0.5])
TAnsSet3Z1: set([-0.125, -0.5], [0.125, -0.5])
TAnsSet3Z2: set([-0.625, -0.5], [0.625, -0.5])

/* L&#246;sungen */
TansB: (2 - q) / 3
TansC: [[1, false, "Strategie 1"], [2, true, "Strategie 2"], [3, false, "Strategie 3"]]]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text/>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ansA11</name>
            <type>algebraic</type>
            <tans>[-0.75,0]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA12</name>
            <type>algebraic</type>
            <tans>[0.75,0]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA13</name>
            <type>algebraic</type>
            <tans>[-0.25,0]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA14</name>
            <type>algebraic</type>
            <tans>[0.25,0]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA21</name>
            <type>algebraic</type>
            <tans>[-0.875,-0.5]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA22</name>
            <type>algebraic</type>
            <tans>[-0.375,-0.5]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA23</name>
            <type>algebraic</type>
            <tans>[0.375,-0.5]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA24</name>
            <type>algebraic</type>
            <tans>[0.875,-0.5]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA25</name>
            <type>algebraic</type>
            <tans>[-0.125,-0.5]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA26</name>
            <type>algebraic</type>
            <tans>[0.125,-0.5]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA27</name>
            <type>algebraic</type>
            <tans>[-0.625,-0.5]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA28</name>
            <type>algebraic</type>
            <tans>[0.625,-0.5]</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansB</name>
            <type>algebraic</type>
            <tans>TansB</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansC</name>
            <type>dropdown</type>
            <tans>TansC</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prtA</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>f(x) := is(second(x) = 0)
g(x) := is(second(x) = -0.5)
h(x,y ):= is(x = y)

SAnsSet2Z1: set(ansA11, ansA12)
SAnsSet2Z2: set(ansA13, ansA14)
SAnsSet2: union(SAnsSet2Z1, SAnsSet2Z2)
B2: every(f, SAnsSet2) and is(cardinality(SAnsSet2) = 4)
B2Z1: h(SAnsSet2Z1, TAnsSet2Z1)
B2Z2: h(SAnsSet2Z2, TAnsSet2Z2)

SAnsSet3G: set(ansA21, ansA22, ansA23, ansA24)
SAnsSet3Z1: set(ansA25, ansA26)
SAnsSet3Z2: set(ansA27, ansA28)
SAnsSet3: union(SAnsSet3G, SAnsSet3Z1, SAnsSet3Z2)
B3: every(g, SAnsSet3) and is(cardinality(SAnsSet3) = 8)
B3G: h(SAnsSet3G, TAnsSet3G)
B3Z1: h(SAnsSet3Z1, TAnsSet3Z1)
B3Z2: h(SAnsSet3Z2, TAnsSet3Z2)

B: h([B2Z1, B2Z2, B3G, B3Z1, B3Z2], [true, true, true, true, true])</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>B2</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prtA-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtA-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben nicht jedem der im zweiten Schritt m&#246;glichen Ereignisse ein blaues, rautenf&#246;rmiges Feld zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        They have not assigned a blue, diamond-shaped field to each of the events possible in the second step.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>B3</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prtA-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtA-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben nicht jedem der im dritten Schritt m&#246;glichen Ereignisse ein rotes, quadratisches Feld zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        They have not assigned a red square field to each of the events possible in the third step.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>B2Z1</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prtA-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die blauen, rautenf&#246;rmigen und mit Z1 beschrifteten Felder sind den im zweiten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The blue, rhomb-shaped fields labeled with Z1 are correctly assigned to events possible in the second step.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prtA-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Nicht alle der blauen, rautenf&#246;rmigen und mit Z1 beschrifteten Felder sind den im zweiten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Not all of the blue, diamond-shaped fields marked with Z1 are correctly assigned to the events possible in the second step.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>B2Z1</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prtA-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die blauen, rautenf&#246;rmigen und mit Z2 beschrifteten Felder sind den im zweiten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The blue, rhomb-shaped fields labeled with Z2 are correctly assigned to events possible in the second step.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prtA-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Nicht alle der blauen, rautenf&#246;rmigen und mit Z2 beschrifteten Felder sind den im zweiten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Not all of the blue, diamond-shaped fields marked with Z2 are correctly assigned to the events possible in the second step.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>B3G</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>5</truenextnode>
                <trueanswernote>prtA-5-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die roten, quadratischen und mit G beschrifteten Felder sind den im dritten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The red, square and G-labeled fields are correctly assigned to events possible in the third step.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>5</falsenextnode>
                <falseanswernote>prtA-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Nicht alle der roten, quadratischen und mit G beschrifteten Felder sind den im dritten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Not all of the red, square and G-labeled fields are correctly assigned to events possible in the third step.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>B3Z1</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>6</truenextnode>
                <trueanswernote>prtA-6-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die roten, quadratischen und mit Z1 beschrifteten Felder sind den im dritten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The red, square fields marked with Z1 are correctly assigned to events possible in the third step.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>6</falsenextnode>
                <falseanswernote>prtA-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Nicht alle der roten, quadratischen und mit Z1 beschrifteten Felder sind den im dritten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Not all of the red, square and Z1 labeled fields are correctly assigned to events possible in the third step.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>6</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>B3Z2</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.2000000</truescore>
                <truepenalty/>
                <truenextnode>7</truenextnode>
                <trueanswernote>prtA-7-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die roten, quadratischen und mit Z2 beschrifteten Felder sind den im dritten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The red, square fields marked with Z2 are correctly assigned to events possible in the third step.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>7</falsenextnode>
                <falseanswernote>prtA-7-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Nicht alle der roten, quadratischen und mit Z2 beschrifteten Felder sind den im dritten Schritt m&#246;glichen Ereignissen richtig zugeordnet.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Not all of the red, square and Z2 labeled fields are correctly assigned to events possible in the third step.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>7</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>B</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtA-8-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-8a')
;show('oer-aufgabe-8b');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-8a')
;show('oer-aufgabe-8b');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtA-8-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtB</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansB</sans>
                <tans>TansB</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtB-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Es gibt vier m&#246;gliche Pfade, die zum Gewinn \(G\) f&#252;hren: \(G\rightarrow Z_1\rightarrow G\) und \(G\rightarrow Z_2\rightarrow G\) und \(Z_1\rightarrow Z_2\rightarrow G\) und \(Z_2\rightarrow Z_1\rightarrow G\). Indem wir mit dem Multiplikationssatz die bedingten Wahrscheinlichkeiten entlang der jeweiligen Kanten multiplizieren und diese vier Zwischenergebnisse summieren, erhalten wir die Gewinnwahrscheinlichkeit
\begin{equation*}
\frac{1}{3}\cdot p\cdot q+\frac{1}{3}\cdot(1-p)\cdot q+\frac{1}{3}\cdot1\cdot(1-q)+\frac{1}{3}\cdot1\cdot(1-q)=\frac{2-q}{3}.
\end{equation*}
Der Parameter \(p\) des Showmasters hat also keinen Einfluss auf die Gewinnwahrscheinlichkeit, sondern diese h&#228;ngt allein von der Spielstrategie der Kandidatin ab, kodiert durch den Parameter \(q\). Mit diesem Ergebnis k&#246;nnen wir nun die bestm&#246;gliche Spielstrategie angeben.
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-8b');
hide('oer-aufgabe-8-JSXGraph');
show('oer-aufgabe-8c');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        There are four possible paths to win \(G\) lead: \(G\rightarrow Z_1\rightarrow G\) and \(G\rightarrow Z_2\rightarrow G\) and \(Z_1\rightarrow Z_2\rightarrow G\) and \(Z_2\rightarrow Z_1\rightarrow G\). By multiplying the conditional probabilities along the respective edges with the multiplication set and summing these four intermediate results, we get the probability of profit
\begin{equation*}
\frac{1}{3}\cdot p\cdot q+\frac{1}{3}\cdot(1-p)\cdot q+\frac{1}{3}\cdot1\cdot(1-q)+\frac{1}{3}\cdot1\cdot(1-q)=\frac{2-q}{3}.
\end{equation*}The parameter \(p\) the showmaster thus has no influence on the probability of winning, but this depends solely on the candidate's strategy of playing, encoded by the parameter \(q\). With this result, we can now provide the best possible game strategy.
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-8b');
hide('oer-aufgabe-8-JSXGraph');
show('oer-aufgabe-8c');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtB-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Bestimmen Sie alle Pfade innerhalb des Baumdiagramms, die zum Gewinn \(G\) f&#252;hren, und berechnen Sie mithilfe des Multiplikationssatzes deren Wahrscheinlichkeiten. Summieren Sie anschlie&#223;end alle Teilergebnisse auf, um die Gewinnwahrscheinlichkeit zu erhalten.
    </p>
    <p>
        Probieren Sie es erneut und klicken Sie dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Determine all the paths within the tree chart that are used to win \(G\) guide and calculate their probabilities using the multiplication set. Then add all partial results to get the probability of winning.
    </p>
    <p>
        Try it again and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtC</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtC-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Strategie 1 entspricht \(q=1\) mit der Gewinnwahrscheinlichkeit \(\frac{2-q}{3}=\frac{2-1}{3}=\frac{1}{3}\). Da die monoton fallende Funktion \(q\mapsto\frac{2-q}{3}\) auf dem Intervall \([0,1]\) das Minimum in \(q=1\) annimmt, ist dies jedoch die schlechteste Strategie. Die Kandidatin sollte also auf keinen Fall bei der zuerst gew&#228;hlten T&#252;r bleiben.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Strategy 1 \(q=1\) with profit probability \(\frac{2-q}{3}=\frac{2-1}{3}=\frac{1}{3}\). Since the monotonously falling function \(q\mapsto\frac{2-q}{3}\) on the interval \([0,1]\) the minimum in \(q=1\) However, this is the worst strategy. So the candidate should not stay at the first door.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prtC-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtC-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Strategie 2 entspricht \(q=0\) mit der Gewinnwahrscheinlichkeit \(\frac{2-q}{3}=\frac{2-0}{3}=\frac{2}{3}\). Da die monoton fallende Funktion \(q\mapsto\frac{2-q}{3}\) auf dem Intervall \([0,1]\) das Maximum in \(q=0\) annimmt, ist dies sogar die bestm&#246;gliche Strategie. Die Kandidatin sollte also auf jeden Fall zur anderen noch verschlossenen T&#252;r wechseln.
    </p>
    <p>
        Die schlechteste Strategie erh&#228;lt man f&#252;r \(q=1\), d.h. wenn die Kandidatin mit Strategie 1 spielt und auf jeden Fall bei der zuerst gew&#228;hlten T&#252;r bleibt. Die Gewinnwahrscheinlichkeit betr&#228;gt dann \(\frac{2-q}{3}=\frac{2-1}{3}=\frac{1}{3}\), ist also nur halb so gro&#223; wie bei Strategie 2.
    </p>
    <p>
        Strategie 3 dagegen entspricht \(q=\frac{1}{2}\) mit der Gewinnwahrscheinlichkeit \(\frac{2-q}{3}=\frac{2-\frac{1}{2}}{3}=\frac{1}{2}\) und bietet somit nur eine 50:50 Chance auf den Gewinn.
    </p>
    <hr/>
    <p>
        Sie haben die Aufgabe erfolgreich gel&#246;st! Jetzt sollten Sie das Ziegenproblem sicher erkl&#228;ren k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Strategy 2 \(q=0\) with profit probability \(\frac{2-q}{3}=\frac{2-0}{3}=\frac{2}{3}\). Since the monotonously falling function \(q\mapsto\frac{2-q}{3}\) on the interval \([0,1]\) the maximum in \(q=0\) This is even the best possible strategy. So the candidate should definitely switch to the other door still closed.
    </p>
    <p>
        The worst strategy is given to \(q=1\), i.e. if the candidate plays with strategy 1 and in any case remains with the first chosen door. The profit probability is then \(\frac{2-q}{3}=\frac{2-1}{3}=\frac{1}{3}\)So it's only half as big as strategy 2.
    </p>
    <p>
        Strategy 3 \(q=\frac{1}{2}\) with profit probability \(\frac{2-q}{3}=\frac{2-\frac{1}{2}}{3}=\frac{1}{2}\) and thus offers only a 50:50 chance of winning.
    </p>
    <hr/>
    <p>
        You have successfully solved the task! Now you should be able to explain the goat problem safely.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prtC-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prtC-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Strategie 3 entspricht \(q=\frac{1}{2}\) mit der Gewinnwahrscheinlichkeit \(\frac{2-q}{3}=\frac{2-\frac{1}{2}}{3}=\frac{1}{2}\) und bietet somit eine 50:50 Chance auf den Gewinn. Dies ist jedoch nicht die bestm&#246;gliche Strategie.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Strategy 3 corresponds to \(q=\frac{1}{2}\) with profit probability \(\frac{2-q}{3}=\frac{2-\frac{1}{2}}{3}=\frac{1}{2}\) and thus offers a 50:50 chance of winning. However, this is not the best possible strategy.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prtC-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>