<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>2023-03-07 Urnenmodell vs. F&#228;chermodell</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <div id="oer-aufgabe-6-info">
        <details>
            <summary>
                Inhalt und Umfang dieser Aufgabe
            </summary>
            <p>
                In dieser Aufgabe lernen Sie, mit kombinatorischen Modellen zu rechnen. Sie besteht aus den folgenden drei Teilaufgaben.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    l&#246;sen Sie ein kombinatorisches Problem mit ununterscheidbaren Objekten.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    l&#246;sen Sie ein kombinatorisches Problem mit ununterscheidbaren Objekten und einer Nebenbedingung.
                </li>
                <li>
                    In
                    <b>
                        (c)
                    </b>
                    l&#246;sen Sie ein kombinatorisches Problem mit unterscheidbaren Objekten.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-6a">
        <p>
            <b>
                (a)
            </b>
            In einer Schublade eines Apothekerschranks gibt es {@k@} F&#228;cher zur Aufbewahrung von Medikamentenpackungen. Wie viele M&#246;glichkeiten gibt es, {@l@} gleich aussehende Packungen eines Medikaments in die Schublade einzusortieren, wenn in jedes Fach beliebig viele Packungen passen?
        </p>
    </div>
    <div id="oer-aufgabe-6aAns">
        <p>
            Die Anzahl an M&#246;glichkeiten ist [[input:ansA]].
            <i>
                Geben Sie das Ergebnis als Formel ein.
            </i>
            [[validation:ansA]] [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-6a11" style="display:none;">
        <hr/>
        <p>
            <b>
                (a1)
            </b>
            Durch welches F&#228;chermodell kann die Situation sinnvoll beschrieben werden?
        </p>
        <p>
            [[input:ansA11]] [[validation:ansA11]] [[feedback:prtA11]]
        </p>
    </div>
    <div id="oer-aufgabe-6a12" style="display:none;">
        <hr/>
        <p>
            <b>
                (a1)
            </b>
            Durch welches Urnenmodell kann die Situation sinnvoll beschrieben werden?
        </p>
        <p>
            [[input:ansA12]] [[validation:ansA12]] [[feedback:prtA12]]
        </p>
    </div>
    <div id="oer-aufgabe-6a2" style="display:none;">
        <hr/>
        <p>
            <b>
                (a2)
            </b>
            Wir bezeichnen die Anzahl der F&#228;cher mit \(N\) und die Anzahl der Packungen mit \(n\). Wie lautet dann die Formel zur Berechnung der Anzahl an Elementen im gew&#228;hlten kombinatorischen Modell "{@mZoR_Fach@}" bzw. "{@mZoR_Urne@}"?
        </p>
        <p>
            [[input:ansA2]] [[validation:ansA2]] [[feedback:prtA2]]
        </p>
    </div>
    <div id="oer-aufgabe-6b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            Wie viele M&#246;glichkeiten gibt es, {@l@} gleich aussehende Packungen eines Medikaments in eine Schublade mit {@k@} F&#228;chern einzusortieren, wenn in jedes Fach beliebig viele Packungen passen, aber kein Fach leer bleiben soll?
        </p>
    </div>
    <div id="oer-aufgabe-6bAns" style="display:none;">
        <p>
            Die Anzahl an M&#246;glichkeiten ist [[input:ansB]].
            <i>
                Geben Sie das Ergebnis als Formel ein.
            </i>
            [[validation:ansB]] [[feedback:prtB]]
        </p>
    </div>
    <div id="oer-aufgabe-6c" style="display:none;">
        <p>
            <b>
                (c)
            </b>
            Der Apothekerschrank hat {@n@} Schubladen und somit insgesamt \({@n@}\cdot{@k@}={@n*k@}\) F&#228;cher. Damit jedes Medikament m&#246;glichst schnell wiedergefunden werden kann, soll jedes Fach nur mit (ggf. mehreren Packungen von) jeweils einem Medikament belegt werden. Dabei darf ein Fach auch leer bleiben. Wie viele M&#246;glichkeiten gibt es, {@m@} verschiedene Medikamente in den Apothekerschrank einzusortieren?
        </p>
    </div>
    <div id="oer-aufgabe-6cAns" style="display:none;">
        <p>
            Die Anzahl an M&#246;glichkeiten ist [[input:ansC]].
            <i>
                Geben Sie das Ergebnis als Formel ein.
            </i>
            [[validation:ansC]] [[feedback:prtC]]
        </p>
    </div>
    <div id="oer-aufgabe-6c11" style="display:none;">
        <hr/>
        <p>
            <b>
                (c1)
            </b>
            Durch welches F&#228;chermodell kann die Situation sinnvoll beschrieben werden?
        </p>
        <p>
            [[input:ansC11]] [[validation:ansC11]] [[feedback:prtC11]]
        </p>
    </div>
    <div id="oer-aufgabe-6c12" style="display:none;">
        <hr/>
        <p>
            <b>
                (c1)
            </b>
            Durch welches Urnenmodell kann die Situation sinnvoll beschrieben werden?
        </p>
        <p>
            [[input:ansC12]] [[validation:ansC12]] [[feedback:prtC12]]
        </p>
    </div>
    <div id="oer-aufgabe-6c2" style="display:none;">
        <hr/>
        <p>
            <b>
                (c2)
            </b>
            Wir bezeichnen die Anzahl der F&#228;cher mit \(N\) und die Anzahl der Medikamente mit \(n\). Wie lautet dann die Formel zur Berechnung der Anzahl an Elementen im gew&#228;hlten kombinatorischen Modell "{@oZmR_Fach@}" bzw. "{@oZmR_Urne@}"?
        </p>
        <p>
            [[input:ansC2]] [[validation:ansC2]] [[feedback:prtC2]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
<span class="multilang" lang="en">
    <div id="oer-aufgabe-6-info">
        <details>
            <summary>
                Content and scope of this task
            </summary>
            <p>
                In this task you will learn to count on combinatorial models. It consists of the following three subtasks.
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    solve a combinatorial problem with distinctive objects.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    solve a combinatorial problem with distinctive objects and a condition.
                </li>
                <li>
                    In
                    <b>
                        (c)
                    </b>
                    solve a combinatorial problem with distinguishable objects.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <div id="oer-aufgabe-6a">
        <p>
            <b>
                (a)
            </b>
            In a drawer of a pharmacy cabinet there is {@k@} compartments for storing medication packages. How many possibilities are there, {@l@} to sort the same-looking packages of a drug into the drawer if any number of packages fit into each compartment?
        </p>
    </div>
    <div id="oer-aufgabe-6aAns">
        <p>
            The number of possibilities is [[input:ansA]].
            <i>
                Enter the result as formula.
            </i>
            [[validation:ansA]] [[feedback:prtA]]
        </p>
    </div>
    <div id="oer-aufgabe-6a11" style="display:none;">
        <hr/>
        <p>
            <b>
                (a1)
            </b>
            What fan model can the situation be described meaningfully?
        </p>
        <p>
            [[input:ansA11]] [[validation:ansA11]] [[feedback:prtA11]]
        </p>
    </div>
    <div id="oer-aufgabe-6a12" style="display:none;">
        <hr/>
        <p>
            <b>
                (a1)
            </b>
            What urn model can the situation be described meaningfully?
        </p>
        <p>
            [[input:ansA12]] [[validation:ansA12]] [[feedback:prtA12]]
        </p>
    </div>
    <div id="oer-aufgabe-6a2" style="display:none;">
        <hr/>
        <p>
            <b>
                (a2)
            </b>
            We call the number of subjects with \(N\) and the number of packages with \(n\). What is the formula for calculating the number of elements in the selected combinatorial model "{@mZoR_Fach@}" or "{@mZoR_Urne@}"?
        </p>
        <p>
            [[input:ansA2]] [[validation:ansA2]] [[feedback:prtA2]]
        </p>
    </div>
    <div id="oer-aufgabe-6b" style="display:none;">
        <p>
            <b>
                (b)
            </b>
            How many possibilities are there, {@l@} identical-looking packs of a drug in a drawer with {@k@} Sort compartments if any number of packages fit into each compartment, but no compartment should remain empty?
        </p>
    </div>
    <div id="oer-aufgabe-6bAns" style="display:none;">
        <p>
            The number of possibilities is [[input:ansB]].
            <i>
                Enter the result as formula.
            </i>
            [[validation:ansB]] [[feedback:prtB]]
        </p>
    </div>
    <div id="oer-aufgabe-6c" style="display:none;">
        <p>
            <b>
                (c)
            </b>
            The pharmacy cabinet has {@n@} drawers and thus overall \({@n@}\cdot{@k@}={@n*k@}\) Subjects. In order for each medication to be found as quickly as possible, each compartment should only be occupied with (possibly several packs of) one medication. A compartment may also remain empty. How many possibilities are there, {@m@} sorting various medicines into the pharmacy cabinet?
        </p>
    </div>
    <div id="oer-aufgabe-6cAns" style="display:none;">
        <p>
            The number of possibilities is [[input:ansC]].
            <i>
                Enter the result as formula.
            </i>
            [[validation:ansC]] [[feedback:prtC]]
        </p>
    </div>
    <div id="oer-aufgabe-6c11" style="display:none;">
        <hr/>
        <p>
            <b>
                (c1)
            </b>
            What fan model can the situation be described meaningfully?
        </p>
        <p>
            [[input:ansC11]] [[validation:ansC11]] [[feedback:prtC11]]
        </p>
    </div>
    <div id="oer-aufgabe-6c12" style="display:none;">
        <hr/>
        <p>
            <b>
                (c1)
            </b>
            What urn model can the situation be described meaningfully?
        </p>
        <p>
            [[input:ansC12]] [[validation:ansC12]] [[feedback:prtC12]]
        </p>
    </div>
    <div id="oer-aufgabe-6c2" style="display:none;">
        <hr/>
        <p>
            <b>
                (c2)
            </b>
            We call the number of subjects with \(N\) and the number of drugs with \(n\). What is the formula for calculating the number of elements in the selected combinatorial model "{@oZmR_Fach@}" or "{@oZmR_Urne@}"?
        </p>
        <p>
            [[input:ansC2]] [[validation:ansC2]] [[feedback:prtC2]]
        </p>
    </div>
    <style>
        .buttonweiter {
                border: 1px solid black;
                background-color: #A9F5BC; 
}
.buttonweiter:hover {
                background-color: #E0F8E6;
}
.buttonzurueck {
                border: 1px solid black;
                background-color: #DCB25E; 
}
.buttonzurueck:hover {
                background-color: #E7D2A7;
}
.buttontipp {
                border: 1px solid black;
                background-color: #c9d3f6; 
}
.buttontipp:hover {
                background-color: #dbe1f9;
}
.divtipp {
                background-color:#dbe1f9;
                border:1px black solid;
                border-radius: 4px;
                padding:20px;
                width:60%;
                text-align:center;
}
.herv {
                color: purple;
}
.erinnerungsbox {
                border: 2px rgb(152, 202, 62) solid;
                padding: 5px;
}
.list-group > .indent {
                margin-left: 15%;
}
span.stack-tooltip {
                text-decoration-line: underline;
                text-decoration-style: dotted;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>3.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
Urnenmodell vs. F&#228;chermodell

wurde entwickelt von
  
Christian M&#252;ller <c.mueller(at)hhu.de>

nach einer Idee von

Axel B&#252;cher <axel.buecher(at)hhu.de>
Peter Kern <kern(at)hhu.de>

an der Heinrich-Heine-Universit&#228;t D&#252;sseldorf.

Dieses Werk ist lizenziert unter einer Creative Commons Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie https://creativecommons.org/licenses/by-sa/4.0/.

SPDX-License-Identifier: CC-BY-SA-4.0

Technische Informationen:

Diese Aufgabe bindet das Skript 'stackselbstlern.js' von Michael Kallweit f&#252;r die Aufgabennavigation ein.
*/

/* &#252;berpr&#252;fe, ob ans als Wahrscheinlichkeit interpretiert werden kann */
probabilityp(ans) := is(numberp(float(ans)) and ans >= 0 and ans <= 1)
/* runde x auf 3 Nachkommastellen */
dp(x) := decimalplaces(x, 3)
/* stelle x in wissenschaftlicher Notation dar */
SCI(x) := scientific_notation(x, 2)
/* Rundungsgenauigkeit bei numerischen Antworttests */
epsilon: 0.001
/* vereinfache ganzzahlige Fakult&#228;ten nur bis factlim */
factlim: 10

/* Anzahl der Schubladen im Schrank */
n: rand_with_step(10, 13, 1)
/* Anzahl der F&#228;cher pro Schublade */
k: rand_with_step(10, 13, 1)
/* Anzahl der Packungen */
l: rand(10) + k +2
/* Anzahl der Medikamente */
m: rand_with_step(n + 1, n * k - 1, 1)

/* Formeln f&#252;r F&#228;cher- bzw. Urnenmodelle */
mZmR(N, n) := N^n
oZmR(N, n) := N! / (N - n)!
mZoR(N, n) := binomial(N + n - 1, n)
oZoR(N, n) := binomial(N, n)

/* Abk&#252;rzungen f&#252;r Formeln */
mZmR_Formel: "\\(N^n\\)"
oZmR_Formel: "\\(\\dfrac{N!}{(N - n)!}\\)"
mZoR_Formel: "\\(\\dbinom{N + n - 1}{n}\\)"
oZoR_Formel: "\\(\\dbinom{N}{n}\\)"
tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
/* Abk&#252;rzungen f&#252;r F&#228;chermodelle */
mZmR_Fach: tr("Verteilen unterscheidbarer Objekte mit Mehrfachbelegung","Distribution of distinguishable objects with multiple occupancy")
oZmR_Fach: tr("Verteilen unterscheidbarer Objekte ohne Mehrfachbelegung","Distribution of distinguishable objects without multiple occupancy")
mZoR_Fach: tr("Verteilen ununterscheidbarer Objekte mit Mehrfachbelegung","Distribution of indistinguishable objects with multiple occupancy")
oZoR_Fach: tr("Verteilen ununterscheidbarer Objekte ohne Mehrfachbelegung","Distribution of indistinguishable objects without multiple occupancy")

/* Abk&#252;rzungen f&#252;r Urnenmodelle */
mZmR_Urne: tr("Ziehen mit Zur&#252;cklegen und mit Reihenfolge","pull with put back and with order")
oZmR_Urne: tr("Ziehen ohne Zur&#252;cklegen und mit Reihenfolge","pull without put back and with order")
mZoR_Urne: tr("Ziehen mit Zur&#252;cklegen und ohne Reihenfolge","pull with put back and without order")
oZoR_Urne: tr("Ziehen ohne Zur&#252;cklegen und ohne Reihenfolge","pull without put back and without order")

/* L&#246;sungen */
TansA: mZoR(k, l)
TansA11: [[1, false, mZmR_Fach], [2, false, oZmR_Fach], [3, true, mZoR_Fach], [4, false, oZoR_Fach]]
TansA12: [[1, false, mZmR_Urne], [2, false, oZmR_Urne], [3, true, mZoR_Urne], [4, false, oZoR_Urne]]
TansA2: [[1, false, mZmR_Formel], [2, false, oZmR_Formel], [3, true, mZoR_Formel], [4, false, oZoR_Formel]]
TansB: mZoR(k, l - k)
TansC: oZmR(n * k, m)
/* Hinweis: Nur f&#252;r n * k <= 170 kann der numerische Wert von (n * k)! problemlos dargestellt werden. */
TansC11: [[1, false, mZmR_Fach], [2, true, oZmR_Fach], [3, false, mZoR_Fach], [4, false, oZoR_Fach]]
TansC12: [[1, false, mZmR_Urne], [2, true, oZmR_Urne], [3, false, mZoR_Urne], [4, false, oZoR_Urne]]
TansC2: [[1, false, mZmR_Formel], [2, true, oZmR_Formel], [3, false, mZoR_Formel], [4, false, oZoR_Formel]]]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(n={@n@}\), \(k={@k@}\), \(l={@l@}\), \(m={@m@}\)</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ansA</name>
            <type>numerical</type>
            <tans>TansA</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansA11</name>
            <type>radio</type>
            <tans>TansA11</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ansA12</name>
            <type>radio</type>
            <tans>TansA12</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ansA2</name>
            <type>radio</type>
            <tans>TansA2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ansB</name>
            <type>numerical</type>
            <tans>TansB</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansC</name>
            <type>numerical</type>
            <tans>TansC</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansC11</name>
            <type>radio</type>
            <tans>TansC11</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ansC12</name>
            <type>radio</type>
            <tans>TansC12</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ansC2</name>
            <type>radio</type>
            <tans>TansC2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <prt>
            <name>prtA</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA</sans>
                <tans>TansA</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die Situation wird durch das F&#228;chermodell "{@mZoR_Fach@}" bzw. durch das Urnenmodell "{@mZoR_Urne@}" sinnvoll beschrieben. Somit gibt es \[\binom{{@k@}+{@l@}-1}{{@l@}}=\binom{{@k+l-1@}}{{@l@}} \approx {@SCI(TansA)@}\] M&#246;glichkeiten.
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6a');
hide('oer-aufgabe-6aAns');
show('oer-aufgabe-6b');
show('oer-aufgabe-6bAns');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The situation is caused by the fan model "{@mZoR_Fach@}" or through the urn model "{@mZoR_Urne@}" reasonably described. So there are \[\binom{{@k@}+{@l@}-1}{{@l@}}=\binom{{@k+l-1@}}{{@l@}} \approx {@SCI(TansA)@}\] Opportunities.
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6a');
hide('oer-aufgabe-6aAns');
show('oer-aufgabe-6b');
show('oer-aufgabe-6bAns');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <div id="oer-aufgabe-6a-PRT1">
        <p>
            Bevor Sie die Frage erneut versuchen, schauen wir uns zun&#228;chst das zugrunde liegende kombinatorische Modell an. W&#228;hlen Sie zwischen der Sichtweise eines
            <button class="buttonweiter" onclick="
hide('oer-aufgabe-6aAns');
show('oer-aufgabe-6a11');
" type="button">
                F&#228;chermodells
            </button>
            oder eines
            <button class="buttonweiter" onclick="
hide('oer-aufgabe-6aAns');
show('oer-aufgabe-6a12');
" type="button">
                Urnenmodells
            </button>
            .
        </p>
    </div>
    <div id="oer-aufgabe-6a-PRT2" style="display:none;">
        <p>
            <b>
                Bearbeiten Sie mithilfe des Zwischenschritts nun noch einmal die Ausgangsfrage.
            </b>
        </p>
        <p>
            Sie haben gerade gelernt, dass das F&#228;chermodell "{@mZoR_Fach@}" bzw. das Urnenmodell "{@mZoR_Urne@}" zur Beschreibung der Situation geeignet ist, und dass die Anzahl der Elemente in diesem kombinatorischen Modell durch die Formel {@mZoR_Formel@} berechnet wird. Identifizieren Sie nun noch die richtigen Werte von \(N\) und \(n\) und beantworten Sie dann die Frage erneut.
        </p>
    </div>
</span>
<span class="multilang" lang="en">
    <div id="oer-aufgabe-6a-PRT1">
        <p>
            Before you try again, we first look at the underlying combinatorial model. Choose between the view of one
            <button class="buttonweiter" onclick="
hide('oer-aufgabe-6aAns');
show('oer-aufgabe-6a11');
" type="button">
                Fan model
            </button>
            or
            <button class="buttonweiter" onclick="
hide('oer-aufgabe-6aAns');
show('oer-aufgabe-6a12');
" type="button">
                Urnenmodells
            </button>
            .
        </p>
    </div>
    <div id="oer-aufgabe-6a-PRT2" style="display:none;">
        <p>
            <b>
                Edit the initial question with the help of the intermediate step.
            </b>
        </p>
        <p>
            You just learned that the fan model "{@mZoR_Fach@}" or the urn model "{@mZoR_Urne@}" is suitable for describing the situation and that the number of elements in this combinatorial model by the formula {@mZoR_Formel@} is calculated. Identify the correct values of \(N\) and \(n\) and then answer the question again.
        </p>
    </div>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtA11</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA11</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1A1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das F&#228;chermodell "{@mZmR_Fach@}" entspricht nicht der vorliegenden Situation, da die Packungen gleich aussehen. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The fan model "{@mZmR_Fach@}" does not correspond to the present situation, since the packages look the same. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1A1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA11</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1A1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das F&#228;chermodell "{@oZmR_Fach@}" entspricht nicht der vorliegenden Situation, da in jedes Fach beliebig viele Packungen passen und die Packungen gleich aussehen. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The fan model "{@oZmR_Fach@}" does not correspond to the present situation, since any number of packages fit into each compartment and the packages look the same. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1A1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA11</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1A1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie haben ein sinnvolles F&#228;chermodell ausgew&#228;hlt. Aus der Sicht eines Urnenmodells entspricht das &#252;brigens dem "{@mZoR_Urne@}".
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6a11');
show('oer-aufgabe-6a2');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You have selected a meaningful fan model. From the viewpoint of a urn model, this corresponds to the "{@mZoR_Urne@}".
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6a11');
show('oer-aufgabe-6a2');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt1A1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA11</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1A1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das F&#228;chermodell "{@oZoR_Fach@}" entspricht nicht der vorliegenden Situation, da in jedes Fach beliebig viele Packungen passen. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The fan model "{@oZoR_Fach@}" does not correspond to the present situation, since any number of packages fit into each compartment. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1A1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtA12</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA12</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1A2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das Urnenmodell "{@mZmR_Urne@}" entspricht nicht der vorliegenden Situation, da die Packungen gleich aussehen. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The urn model "{@mZmR_Urne@}" does not correspond to the present situation, since the packages look the same. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1A2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA12</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1A2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das Urnenmodell "{@oZmR_Urne@}" entspricht nicht der vorliegenden Situation, da in jedes Fach beliebig viele Packungen passen und die Packungen gleich aussehen. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The urn model "{@oZmR_Urne@}" does not correspond to the present situation, since any number of packages fit into each compartment and the packages look the same. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1A2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA12</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1A2-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie haben ein sinnvolles Urnenmodell ausgew&#228;hlt. Aus der Sicht eines F&#228;chermodells entspricht das &#252;brigens dem "{@mZoR_Fach@}".
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6a12');
show('oer-aufgabe-6a2');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You have selected a meaningful urn model. From the point of view of a fan model, the "{@mZoR_Fach@}".
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6a12');
show('oer-aufgabe-6a2');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt1A2-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA12</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1A2-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das Urnenmodell "{@oZoR_Urne@}" entspricht nicht der vorliegenden Situation, da in jedes Fach beliebig viele Packungen passen. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The urn model "{@oZoR_Urne@}" does not correspond to the present situation, since any number of packages fit into each compartment. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1A2-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtA2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA2</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1B-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Diese Formel gilt f&#252;r das kombinatorische Modell "{@mZmR_Fach@}" bzw. "{@mZmR_Urne@}". Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        This formula applies to the combinatorial model "{@mZmR_Fach@}" or "{@mZmR_Urne@}". Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1B-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA2</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1B-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Diese Formel gilt f&#252;r das kombinatorische Modell "{@oZmR_Fach@}" bzw. "{@oZmR_Urne@}". Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        This formula applies to the combinatorial model "{@oZmR_Fach@}" or "{@oZmR_Urne@}". Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt1B-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA2</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1B-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie haben die richtige Formel f&#252;r das kombinatorische Modell "{@mZoR_Fach@}" bzw. "{@mZoR_Urne@}" ausgew&#228;hlt.
    </p>
    <hr/>
    <p>
        Jetzt k&#246;nnen Sie die Ausgangsfrage noch einmal beantworten. Klicken Sie dazu auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6a2');
show('oer-aufgabe-6aAns');
hide('oer-aufgabe-6a-PRT1');
show('oer-aufgabe-6a-PRT2');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You have the right formula for the combinatorial model "{@mZoR_Fach@}" or "{@mZoR_Urne@}" selected.
    </p>
    <hr/>
    <p>
        Now you can answer the starting question again. Click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6a2');
show('oer-aufgabe-6aAns');
hide('oer-aufgabe-6a-PRT1');
show('oer-aufgabe-6a-PRT2');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt1B-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansA2</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1B-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Diese Formel gilt f&#252;r das kombinatorische Modell "{@oZoR_Fach@}" bzw. "{@oZoR_Urne@}". Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        This formula applies to the combinatorial model "{@oZoR_Fach@}" or "{@oZoR_Urne@}". Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1B-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtB</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansB</sans>
                <tans>TansB</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die Nebenbedingung, dass kein Fach leer bleiben soll, k&#246;nnen wir erf&#252;llen, indem wir zuerst in jedes Fach eine Packung legen. Danach sind noch {@l-k@} Packungen &#252;brig, die in die {@k@} F&#228;cher einsortiert werden sollen. Diese Situation wird wie in (a) durch das F&#228;chermodell "{@mZoR_Fach@}" bzw. durch das Urnenmodell "{@mZoR_Urne@}" sinnvoll beschrieben. Somit gibt es \[\binom{{@k@}+{@l-k@}-1}{{@l-k@}}=\binom{{@l-1@}}{{@l-k@}}={@TansB@}\] M&#246;glichkeiten.
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6b');
hide('oer-aufgabe-6bAns');
show('oer-aufgabe-6c');
show('oer-aufgabe-6cAns');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The condition that no compartment should remain empty can be fulfilled by placing a package first in each compartment. After that, {@l-k@} In the {@k@} It is necessary to sort in compartments. This situation is described as in (a) by the fan model "{@mZoR_Fach@}" or through the urn model "{@mZoR_Urne@}" reasonably described. So there are \[\binom{{@k@}+{@l-k@}-1}{{@l-k@}}=\binom{{@l-1@}}{{@l-k@}}={@TansB@}\] Opportunities.
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6b');
hide('oer-aufgabe-6bAns');
show('oer-aufgabe-6c');
show('oer-aufgabe-6cAns');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Vermutlich bereitet Ihnen die Nebenbedingung, dass kein Fach leer bleiben soll, Probleme. Diese k&#246;nnen wir erf&#252;llen, indem wir zuerst in jedes Fach eine Packung legen. Danach sind noch {@l-k@} Packungen &#252;brig, die in die {@k@} F&#228;cher einsortiert werden sollen. Erinnern Sie sich an den Aufgabenteil (a), in dem Sie eine &#228;hnliche Fragestellung mit dem F&#228;chermodell "{@mZoR_Fach@}" bzw. dem Urnenmodell "{@mZoR_Urne@}" gel&#246;st haben.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Presumably, the condition that no compartment should remain empty will cause you problems. We can fulfill these by placing a pack in each compartment first. After that, {@l-k@} In the {@k@} It is necessary to sort in compartments. Remember the task (a) in which you have a similar question with the fan model "{@mZoR_Fach@}" or the urn model "{@mZoR_Urne@}" have solved.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtC</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC</sans>
                <tans>TansC</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die Situation wird durch das F&#228;chermodell "{@oZmR_Fach@}" bzw. durch das Urnenmodell "{@oZmR_Urne@}" sinnvoll beschrieben. Somit gibt es \[\frac{{@n*k@}!}{({@n*k@}-{@m@})!}=\frac{{@n*k@}!}{{@n*k-m@}!} \approx {@SCI(TansC)@}\] M&#246;glichkeiten.
    </p>
    <hr/>
    <p>
        Sie haben die Aufgabe erfolgreich gel&#246;st! Jetzt sollten Sie mit kombinatorischen Modellen sicher rechnen k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The situation is caused by the fan model "{@oZmR_Fach@}" or through the urn model "{@oZmR_Urne@}" reasonably described. So there are \[\frac{{@n*k@}!}{({@n*k@}-{@m@})!}=\frac{{@n*k@}!}{{@n*k-m@}!} \approx {@SCI(TansC)@}\] Opportunities.
    </p>
    <hr/>
    <p>
        You have successfully solved the task! Now you should be able to count on combinatorial models safely.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <div id="oer-aufgabe-6c-PRT1">
        <p>
            Bevor Sie die Frage erneut versuchen, schauen wir uns zun&#228;chst das zugrunde liegende kombinatorische Modell an. W&#228;hlen Sie zwischen der Sichtweise eines
            <button class="buttonweiter" onclick="
hide('oer-aufgabe-6cAns');
show('oer-aufgabe-6c11');
" type="button">
                F&#228;chermodells
            </button>
            oder eines
            <button class="buttonweiter" onclick="
hide('oer-aufgabe-6cAns');
show('oer-aufgabe-6c12');
" type="button">
                Urnenmodells
            </button>
            .
        </p>
    </div>
    <div id="oer-aufgabe-6c-PRT2" style="display:none;">
        <p>
            <b>
                Bearbeiten Sie mithilfe des Zwischenschritts nun noch einmal die Ausgangsfrage.
            </b>
        </p>
        <p>
            Sie haben gerade gelernt, dass das F&#228;chermodell "{@oZmR_Fach@}" bzw. das Urnenmodell "{@oZmR_Urne@}" zur Beschreibung der Situation geeignet ist, und dass die Anzahl der Elemente in diesem kombinatorischen Modell durch die Formel {@oZmR_Formel@} berechnet wird. Identifizieren Sie nun noch die richtigen Werte von \(N\) und \(n\) und beantworten Sie dann die Frage erneut.
        </p>
    </div>
</span>
<span class="multilang" lang="en">
    <div id="oer-aufgabe-6c-PRT1">
        <p>
            Before you try again, we first look at the underlying combinatorial model. Choose between the view of one
            <button class="buttonweiter" onclick="
hide('oer-aufgabe-6cAns');
show('oer-aufgabe-6c11');
" type="button">
                Fan model
            </button>
            or
            <button class="buttonweiter" onclick="
hide('oer-aufgabe-6cAns');
show('oer-aufgabe-6c12');
" type="button">
                Urnenmodells
            </button>
            .
        </p>
    </div>
    <div id="oer-aufgabe-6c-PRT2" style="display:none;">
        <p>
            <b>
                Edit the initial question with the help of the intermediate step.
            </b>
        </p>
        <p>
            You just learned that the fan model "{@oZmR_Fach@}" or the urn model "{@oZmR_Urne@}" is suitable for describing the situation and that the number of elements in this combinatorial model by the formula {@oZmR_Formel@} is calculated. Identify the correct values of \(N\) and \(n\) and then answer the question again.
        </p>
    </div>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtC11</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC11</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3A1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das F&#228;chermodell "{@mZmR_Fach@}" entspricht nicht der vorliegenden Situation, da jedes Fach nur mit jeweils einem Medikament belegt werden soll. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The fan model "{@mZmR_Fach@}" does not correspond to the present situation, since each compartment is to be occupied with only one drug. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3A1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC11</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3A1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie haben ein sinnvolles F&#228;chermodell ausgew&#228;hlt. Aus der Sicht eines Urnenmodells entspricht das &#252;brigens dem "{@oZmR_Urne@}".
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6c11');
show('oer-aufgabe-6c2');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You have selected a meaningful fan model. From the viewpoint of a urn model, this corresponds to the "{@oZmR_Urne@}".
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6c11');
show('oer-aufgabe-6c2');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt3A1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC11</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3A1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das F&#228;chermodell "{@mZoR_Fach@}" entspricht nicht der vorliegenden Situation, da jedes Fach nur mit jeweils einem Medikament belegt werden soll und die Medikamente verschieden sind. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The fan model "{@mZoR_Fach@}" does not correspond to the present situation, since each subject is to be occupied only with one drug and the drugs are different. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt3A1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC11</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3A1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das F&#228;chermodell "{@oZoR_Fach@}" entspricht nicht der vorliegenden Situation, da die Medikamente verschieden sind. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The fan model "{@oZoR_Fach@}" does not correspond to the present situation, since the drugs are different. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3A1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtC12</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC12</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3A2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das Urnenmodell "{@mZmR_Urne@}" entspricht nicht der vorliegenden Situation, da jedes Fach nur mit jeweils einem Medikament belegt werden soll. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The urn model "{@mZmR_Urne@}" does not correspond to the present situation, since each compartment is to be occupied with only one drug. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3A2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC12</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3A2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie haben ein sinnvolles Urnenmodell ausgew&#228;hlt. Aus der Sicht eines F&#228;chermodells entspricht das &#252;brigens dem "{@oZmR_Fach@}".
    </p>
    <hr/>
    <p>
        Um zum n&#228;chsten Aufgabenteil zu gelangen, klicken Sie auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6c12');
show('oer-aufgabe-6c2');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You have selected a meaningful urn model. From the point of view of a fan model, the "{@oZmR_Fach@}".
    </p>
    <hr/>
    <p>
        To get to the next task, click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6c12');
show('oer-aufgabe-6c2');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt3A2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC12</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3A2-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das Urnenmodell "{@mZoR_Urne@}" entspricht nicht der vorliegenden Situation, da jedes Fach nur mit jeweils einem Medikament belegt werden soll und die Medikamente verschieden sind. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The urn model "{@mZoR_Urne@}" does not correspond to the present situation, since each subject is to be occupied only with one drug and the drugs are different. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt3A2-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC12</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3A2-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Das Urnenmodell "{@oZoR_Urne@}" entspricht nicht der vorliegenden Situation, da die Medikamente verschieden sind. Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        The urn model "{@oZoR_Urne@}" does not correspond to the present situation, since the drugs are different. Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3A2-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prtC2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC2</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3B-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Diese Formel gilt f&#252;r das kombinatorische Modell "{@mZmR_Fach@}" bzw. "{@mZmR_Urne@}". Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        This formula applies to the combinatorial model "{@mZmR_Fach@}" or "{@mZmR_Urne@}". Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt3B-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC2</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3B-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie haben die richtige Formel f&#252;r das kombinatorische Modell "{@oZmR_Fach@}" bzw. "{@oZmR_Urne@}" ausgew&#228;hlt.
    </p>
    <hr/>
    <p>
        Jetzt k&#246;nnen Sie die Ausgangsfrage noch einmal beantworten. Klicken Sie dazu auf
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6c2');
show('oer-aufgabe-6cAns');
hide('oer-aufgabe-6c-PRT1');
show('oer-aufgabe-6c-PRT2');
" type="button">
            Weiter
        </button>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You have the right formula for the combinatorial model "{@oZmR_Fach@}" or "{@oZmR_Urne@}" selected.
    </p>
    <hr/>
    <p>
        Now you can answer the starting question again. Click on
        <button class="buttonweiter" onclick="
hide('oer-aufgabe-6c2');
show('oer-aufgabe-6cAns');
hide('oer-aufgabe-6c-PRT1');
show('oer-aufgabe-6c-PRT2');
" type="button">
            Next
        </button>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt3B-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC2</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3B-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Diese Formel gilt f&#252;r das kombinatorische Modell "{@mZoR_Fach@}" bzw. "{@mZoR_Urne@}". Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        This formula applies to the combinatorial model "{@mZoR_Fach@}" or "{@mZoR_Urne@}". Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt3B-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansC2</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3B-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
    <p>
        Diese Formel gilt f&#252;r das kombinatorische Modell "{@oZoR_Fach@}" bzw. "{@oZoR_Urne@}". Versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
    <p>
        This formula applies to the combinatorial model "{@oZoR_Fach@}" or "{@oZoR_Urne@}". Try it again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3B-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>321442369</deployedseed>
        <deployedseed>2035281425</deployedseed>
        <deployedseed>707220013</deployedseed>
        <deployedseed>423096658</deployedseed>
        <deployedseed>1197141518</deployedseed>
        <deployedseed>693669683</deployedseed>
        <deployedseed>1386410612</deployedseed>
        <deployedseed>426025044</deployedseed>
        <deployedseed>1052553960</deployedseed>
        <deployedseed>355561115</deployedseed>
        <deployedseed>68</deployedseed>
        <deployedseed>381295442</deployedseed>
        <deployedseed>1849239961</deployedseed>
        <deployedseed>1462595149</deployedseed>
        <deployedseed>1583413784</deployedseed>
        <deployedseed>419243088</deployedseed>
        <deployedseed>1178969919</deployedseed>
        <deployedseed>491071754</deployedseed>
        <deployedseed>770664406</deployedseed>
        <deployedseed>1268785554</deployedseed>
        <deployedseed>1343298884</deployedseed>
        <deployedseed>2000267913</deployedseed>
    </question>
    

</quiz>