<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Brettspiel</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <style>
        .buttonweiter {
                                border: 1px solid black;
                                background-color: #a9f5bc;
                }

    
    
    
    .buttonweiter:hover {
                                background-color: #e0f8e6;
                }

    
    
    
    .buttonzurueck {
                                border: 1px solid black;
                                background-color: #dcb25e;
                }

    
    
    
    .buttonzurueck:hover {
                                background-color: #e7d2a7;
                }

    
    
    
    .buttontipp {
                                border: 1px solid black;
                                background-color: #c9d3f6;
                }

    
    
    
    .buttontipp:hover {
                                background-color: #dbe1f9;
                }

    
    
    
    .divtipp {
                                background-color: #dbe1f9;
                                border: 1px black solid;
                                border-radius: 4px;
                                padding: 20px;
                                width: 60%;
                                text-align: center;
                }

    
    
    
    .tabelle_wiederholung {
                                border: solid 2px rgb(152, 202, 62);
                                padding: 6px;
                }

    
    
    
    .erinnerungsbox {
                                border: 2px rgb(152, 202, 62) solid;
                                padding: 5px;
                                margin-bottom: 2ex;
                }

    
    
    
    .herv {
                                color: purple;
                }

    
    
    
    .que.stack .formulation details {
                                border: 1px solid #aaa;
                                border-radius: 4px;
                                padding: .5em .5em 0;
                }

    
    
    
    .que.stack .formulation summary {
                                font-weight: bold;
                                margin: -.5em -.5em 0;
                                padding: .5em;
                }

    
    
    
    .que.stack .formulation details[open] {
                                padding: .5em;
                }

    
    
    
    .que.stack .formulation details[open] summary {
                                border-bottom: 1px solid #aaa;
                                margin-bottom: .5em;
                }

    
    
    
    span.stack-tooltip {
                                text-decoration-line:underline;
                                text-decoration-style:dotted;
                                cursor: help;
                }
    </style>
    <div>
        <details>
            <summary>
                Inhalt und Umfang dieser Aufgabe
            </summary>
            <p>
                Dieser Aufgabe thematisiert ein Brettspiel. In diesem Kontext bestimmen Sie verschiedene Wahrscheinlichkeiten sowie einen Erwartungswert und eine Varianz. Die Aufgabe besteht aus den folgenden f&#252;nf Teilaufgaben:
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    bestimmen Sie eine Wahrscheinlichkeit.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    geben Sie die Verteilung einer Zufallsvariablen an.
                </li>
                <li>
                    In
                    <b>
                        (c)
                    </b>
                    geben Sie einen Parameter der Verteilung der Zufallsvariablen an.
                </li>
                <li>
                    In
                    <b>
                        (d)
                    </b>
                    geben Sie den Erwartungswert der Zufallsvariablen an.
                </li>
                <li>
                    In
                    <b>
                        (e)
                    </b>
                    geben Sie die Varianz der Zufallsvariablen an.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <!-- Aufgabenstellung -->
    <div id="oer-aufgabe-4_aufgabenstellung" style="display: box">
        <p>
            Beim Brettspiel "Mensch &#228;rgere dich nicht" darf man zu Beginn dreimal w&#252;rfeln und
                                kann starten, wenn man dabei mindestens eine 6 w&#252;rfelt. Ansonsten muss
                                man warten, bis man das n&#228;chste Mal an der Reihe ist und dann wieder
                                dreimal w&#252;rfeln darf.
        </p>
    </div>
    <!-- Erinnerung an alle Dinge, die aus vorherigen Aufgabenteilen bekannt sind -->
    <div class="erinnerungsbox" id="oer-aufgabe-4_div-bekanntes" style="display: none">
        <p>
            <b>
                Folgendes wissen wir bereits aus den vorherigen Teilaufgaben:
            </b>
        </p>
        <ul id="oer-aufgabe-4_div-bekanntes_innen">
        </ul>
    </div>
    <!-- Aufgabe (a) -->
    <div id="oer-aufgabe-4_aufgabe1" style="display: box">
        <p>
            <b>
                (a)
            </b>
            Wie gro&#223; ist die Wahrscheinlichkeit, bei den drei Versuchen zu
                                Beginn mindestens eine 6 zu erhalten?
        </p>
        <p>
            <em>
                Bitte geben Sie einen Bruch oder eine auf mindestens zwei Nachkommastellen gerundete Dezimalzahl ein.
            </em>
        </p>
        <p>
            Antwort: Die Wahrscheinlichkeit betr&#228;gt [[input:ans1]]
                                [[validation:ans1]]
        </p>
        [[feedback:prt1]]
    </div>
    <!-- Aufgabe (b) - Verteilung -->
    <div id="oer-aufgabe-4_aufgabe2" style="display: none">
        <p>
            Wir bezeichnen mit \(X\) die Anzahl der Runden, die man warten muss,
                                bevor man starten kann.
        </p>
        <p>
            <b>
                (b)
            </b>
            Bestimmen Sie die Verteilung von \(X\)
        </p>
        <p>
            Antwort: \(X\) ist [[input:ans2]] [[validation:ans2]]
        </p>
        [[feedback:prt2]]
    </div>
    <!-- Aufgabe (c) - Parameter der Verteilung -->
    <div id="oer-aufgabe-4_aufgabe2_2" style="display: none">
        <p>
            Sie haben soeben festgestellt, dass \(X\) geometrisch verteilt ist.
        </p>
        <p>
            <b>
                (c)
            </b>
            Bitte geben Sie den Parameter der Verteilung an.
        </p>
        <p>
            <em>
                Bitte geben Sie einen Bruch oder eine auf mindestens zwei Nachkommastellen gerundete Dezimalzahl ein.
            </em>
        </p>
        <p>
            Antwort: \(X\) ist geometrisch verteilt mit Parameter [[input:ans3]]
                                [[validation:ans3]]
        </p>
        [[feedback:prt3]]
    </div>
    <!-- Aufgabe d) - &#220;berleitungstext nach Schleife f&#252;r EW -->
    <div id="oer-aufgabe-4_text_nach_schleife_3" style="display: none">
        <p class="herv">
            Beantworten Sie nun noch einmal die folgende Frage. Nutzen Sie dazu die
                                Informationen in der obigen Box.
        </p>
    </div>
    <!-- Aufgabe (d) (Angebot, doch noch den zweiten Zwischenschritt zu nutzen -->
    <div class="herv" id="oer-aufgabe-4_aufgabe3_nochmal-inSchleife" style="display: none">
        <p>
            Falls Sie doch noch einen weiteren Zeischenschritt zur Berechnung von
                                \(\text{E}(X)\) w&#252;nschen, klicken Sie bitte auf den folgenden Button:
            <button class="buttonweiter" onclick="
		    		    		    		    show('oer-aufgabe-4_zwischen3b');
		    		    		    		    hide('oer-aufgabe-4_aufgabe3');
		    		    		    		    hide('oer-aufgabe-4_text_nach_schleife_3');
		    		    		    		    hide('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
                Zum zweiten Zwischenschritt
            </button>
        </p>
    </div>
    <!-- Aufgabe (d) - Zwischenschritt 1 (W-Funktion) -->
    <div id="oer-aufgabe-4_zwischen3a" style="display: none">
        <p>
            Sie haben zuletzt festgestellt, dass \(X\) geometrisch verteilt ist. Dieser Zwischenschritt soll Ihnen dabei helfen, den Erwartungswert von \(X\) zu bestimmen.
        </p>
        <p>
            <b>
                (d.1)
            </b>
            Sei \(Z\) nun eine (ggf. andere) geometrisch verteilte Zufallsvariable
                                mit Parameter \(p\). Geben Sie unter Verwendung von \(p\) die
            <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
                Wahrscheinlichkeitsfunktion
            </span>
            von \(Z\) an.
        </p>
        <p>
            Antwort: Die
            <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
                Wahrscheinlichkeitsfunktion
            </span>
            von \(Z\) in Abh&#228;ngigkeit von
                                \(p\) lautet \(P(Z=k)=\) [[input:ans4_z1]] [[validation:ans4_z1]]
        </p>
        [[feedback:prt4_z1]]
    </div>
    <!-- Aufgabe (d) - Zwischenschritt 2 (Formel EW) -->
    <div id="oer-aufgabe-4_zwischen3b" style="display: none">
        <p>
            Dieser Zwischenschritt soll Ihnen dabei helfen, den Erwartungswert von \(X\) zu bestimmen
        </p>
        <p>
            <b>
                (d.2)
            </b>
            Geben Sie nun noch die Formel f&#252;r \(\text{E}(Z)\) in Abh&#228;ngigkeit von
                                \(p\) ein.
        </p>
        <p>
            Antwort: Die Formel lautet \(\text{E}(Z)=\) [[input:ans4_z2]]
                                [[validation:ans4_z2]]
        </p>
        [[feedback:prt4_z2]]
    </div>
    <!-- Aufgabe (d) - EW -->
    <div id="oer-aufgabe-4_aufgabe3" style="display: none">
        <p>
            <b>
                (d)
            </b>
            Bestimmen Sie den Erwartungswert von \(X\).
        </p>
        <p>
            <em>
                Bitte geben Sie einen Bruch oder eine auf mindestens zwei Nachkommastellen gerundete Dezimalzahl ein.
            </em>
        </p>
        <p>
            Antwort: \(\text{E}(X)=\) [[input:ans4]] [[validation:ans4]]
        </p>
        [[feedback:prt4]]
    </div>
    <!-- Aufgabe (e) - Var -->
    <div id="oer-aufgabe-4_aufgabe4" style="display: none">
        <p>
            <b>
                (e)
            </b>
            Bestimmen Sie die Varianz von \(X\).
        </p>
        <p>
            <em>
                Bitte geben Sie einen Bruch oder eine auf mindestens zwei Nachkommastellen gerundete Dezimalzahl ein.
            </em>
        </p>
        <p>
            Antwort: \(\text{Var}(X)=\) [[input:ans5]] [[validation:ans5]]
        </p>
        [[feedback:prt5]]
    </div>
    <p>
        Bitte klicken Sie, nachdem Sie geantwortet haben, auf "Pr&#252;fen".
    </p>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
<span class="multilang" lang="en">
    <style>
        .buttonweiter {
                                border: 1px solid black;
                                background-color: #a9f5bc;
                }

    
    
    
    .buttonweiter:hover {
                                background-color: #e0f8e6;
                }

    
    
    
    .buttonzurueck {
                                border: 1px solid black;
                                background-color: #dcb25e;
                }

    
    
    
    .buttonzurueck:hover {
                                background-color: #e7d2a7;
                }

    
    
    
    .buttontipp {
                                border: 1px solid black;
                                background-color: #c9d3f6;
                }

    
    
    
    .buttontipp:hover {
                                background-color: #dbe1f9;
                }

    
    
    
    .divtipp {
                                background-color: #dbe1f9;
                                border: 1px black solid;
                                border-radius: 4px;
                                padding: 20px;
                                width: 60%;
                                text-align: center;
                }

    
    
    
    .tabelle_wiederholung {
                                border: solid 2px rgb(152, 202, 62);
                                padding: 6px;
                }

    
    
    
    .erinnerungsbox {
                                border: 2px rgb(152, 202, 62) solid;
                                padding: 5px;
                                margin-bottom: 2ex;
                }

    
    
    
    .herv {
                                color: purple;
                }

    
    
    
    .que.stack .formulation details {
                                border: 1px solid #aaa;
                                border-radius: 4px;
                                padding: .5em .5em 0;
                }

    
    
    
    .que.stack .formulation summary {
                                font-weight: bold;
                                margin: -.5em -.5em 0;
                                padding: .5em;
                }

    
    
    
    .que.stack .formulation details[open] {
                                padding: .5em;
                }

    
    
    
    .que.stack .formulation details[open] summary {
                                border-bottom: 1px solid #aaa;
                                margin-bottom: .5em;
                }

    
    
    
    span.stack-tooltip {
                                text-decoration-line:underline;
                                text-decoration-style:dotted;
                                cursor: help;
                }
    </style>
    <div>
        <details>
            <summary>
                Content and scope of this task
            </summary>
            <p>
                This task addresses a board game. In this context, you define various probabilities as well as an expected value and a variance. The task consists of the following five subtasks:
            </p>
            <ul>
                <li>
                    In
                    <b>
                        (a)
                    </b>
                    determine a probability.
                </li>
                <li>
                    In
                    <b>
                        (b)
                    </b>
                    specify the distribution of a random variable.
                </li>
                <li>
                    In
                    <b>
                        (c)
                    </b>
                    specify a parameter of the distribution of random variables.
                </li>
                <li>
                    In
                    <b>
                        (d)
                    </b>
                    specify the expected value of the random variables.
                </li>
                <li>
                    In
                    <b>
                        (e)
                    </b>
                    specify the variance of random variables.
                </li>
            </ul>
        </details>
        <hr/>
    </div>
    <!-- Aufgabenstellung -->
    <div id="oer-aufgabe-4_aufgabenstellung" style="display: box">
        <p>
            At the board game "People don't annoy you" you can cube three times at the beginning and
can start by diceing at least one 6. Otherwise,
you wait until you get the next time and then again
three times.
        </p>
    </div>
    <!-- Erinnerung an alle Dinge, die aus vorherigen Aufgabenteilen bekannt sind -->
    <div class="erinnerungsbox" id="oer-aufgabe-4_div-bekanntes" style="display: none">
        <p>
            <b>
                We already know the following from the previous subtasks:
            </b>
        </p>
        <ul id="oer-aufgabe-4_div-bekanntes_innen">
        </ul>
    </div>
    <!-- Aufgabe (a) -->
    <div id="oer-aufgabe-4_aufgabe1" style="display: box">
        <p>
            <b>
                (a)
            </b>
            How great is the probability of the three attempts
Start getting at least one 6?
        </p>
        <p>
            <em>
                Please enter a fraction or a decimal number rounded to at least two decimal places.
            </em>
        </p>
        <p>
            Answer: The probability is [[input:ans1]][[validation:ans1]]
        </p>
        [[feedback:prt1]]
    </div>
    <!-- Aufgabe (b) - Verteilung -->
    <div id="oer-aufgabe-4_aufgabe2" style="display: none">
        <p>
            We call \(X\) the number of rounds you have to wait,
before you can start.
        </p>
        <p>
            <b>
                (b)
            </b>
            Determine the distribution of \(X\)
        </p>
        <p>
            Answer: \(X\) is [[input:ans2]] [[validation:ans2]]
        </p>
        [[feedback:prt2]]
    </div>
    <!-- Aufgabe (c) - Parameter der Verteilung -->
    <div id="oer-aufgabe-4_aufgabe2_2" style="display: none">
        <p>
            You have just stated that \(X\) is geometrically distributed.
        </p>
        <p>
            <b>
                (c)
            </b>
            Please specify the parameter of the distribution.
        </p>
        <p>
            <em>
                Please enter a fraction or a decimal number rounded to at least two decimal places.
            </em>
        </p>
        <p>
            Answer: \(X\) is geometrically distributed with parameters [[input:ans3]][[validation:ans3]]
        </p>
        [[feedback:prt3]]
    </div>
    <!-- Aufgabe d) - &#220;berleitungstext nach Schleife f&#252;r EW -->
    <div id="oer-aufgabe-4_text_nach_schleife_3" style="display: none">
        <p class="herv">
            Answer the following question again. Use this
Information in the above box.
        </p>
    </div>
    <!-- Aufgabe (d) (Angebot, doch noch den zweiten Zwischenschritt zu nutzen -->
    <div class="herv" id="oer-aufgabe-4_aufgabe3_nochmal-inSchleife" style="display: none">
        <p>
            If you still have another zeaning step to calculate
    \(\text{E}(X)\) please click the following button:
            <button class="buttonweiter" onclick="
		    		    		    		    show('oer-aufgabe-4_zwischen3b');
		    		    		    		    hide('oer-aufgabe-4_aufgabe3');
		    		    		    		    hide('oer-aufgabe-4_text_nach_schleife_3');
		    		    		    		    hide('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
                The second intermediate step
            </button>
        </p>
    </div>
    <!-- Aufgabe (d) - Zwischenschritt 1 (W-Funktion) -->
    <div id="oer-aufgabe-4_zwischen3a" style="display: none">
        <p>
            You recently found that \(X\) is geometrically distributed. This intermediate step will help you to achieve the expected value of \(X\) to determine.
        </p>
        <p>
            <b>
                (d)
            </b>
            Be careful \(Z\) now a geometrically distributed random variable
with parameters \(p\). Enter using \(p\) the
            <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
                Probability function
            </span>
            of \(Z\) on.
        </p>
        <p>
            Answer: The
            <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
                Probability function
            </span>
            of \(Z\) depending on
    \(p\) read: \(P(Z=k)=\) [[input:ans4_z1]] [[validation:ans4_z1]]
        </p>
        [[feedback:prt4_z1]]
    </div>
    <!-- Aufgabe (d) - Zwischenschritt 2 (Formel EW) -->
    <div id="oer-aufgabe-4_zwischen3b" style="display: none">
        <p>
            This intermediate step will help you to achieve the expected value of \(X\) to determine
        </p>
        <p>
            <b>
                (d.2)
            </b>
            Now give the formula for \(\text{E}(Z)\) depending on
    \(p\) one.
        </p>
        <p>
            Answer: The formula is \(\text{E}(Z)=\) [[input:ans4_z2]][[validation:ans4_z2]]
        </p>
        [[feedback:prt4_z2]]
    </div>
    <!-- Aufgabe (d) - EW -->
    <div id="oer-aufgabe-4_aufgabe3" style="display: none">
        <p>
            <b>
                (d)
            </b>
            Determine the expected value of \(X\).
        </p>
        <p>
            <em>
                Please enter a fraction or a decimal number rounded to at least two decimal places.
            </em>
        </p>
        <p>
            Answer: \(\text{E}(X)=\) [[input:ans4]] [[validation:ans4]]
        </p>
        [[feedback:prt4]]
    </div>
    <!-- Aufgabe (e) - Var -->
    <div id="oer-aufgabe-4_aufgabe4" style="display: none">
        <p>
            <b>
                (e)
            </b>
            Determine the variance of \(X\).
        </p>
        <p>
            <em>
                Please enter a fraction or a decimal number rounded to at least two decimal places.
            </em>
        </p>
        <p>
            Answer: \(\text{Var}(X)=\) [[input:ans5]] [[validation:ans5]]
        </p>
        [[feedback:prt5]]
    </div>
    <p>
        Please click "Check".
    </p>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>5.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
  Brettspiel

  wurde entwickelt von
  
    Jonas Lache <jonas.lache[at]ruhr-uni-bochum.de>

  nach einer Idee von
  
    Herold Dehling <herold.dehling[at]ruhr-uni-bochum.de>

  an der Ruhr-Universit&#228;t Bochum.

  Dieses Werk ist lizenziert unter einer Creative Commons
  Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International
  Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie
  http://creativecommons.org/licenses/by-sa/4.0/.

  SPDX-License-Identifier: CC-BY-SA-4.0

  Technische Informationen:

  Diese Aufgabe bindet das Skript &#8216;stackselbstlern.js&#8217; von Michael
  Kallweit f&#252;r die Aufgabennavigation ein.
*/

/*Funktion, je nach Wert auf zwei Nachkomma- oder siginifikante Stellen rundet:*/
rnd(x):= if (log(abs(x))/log(10))>=0 then decimalplaces(x,2) else significantfigures(x,2);

/*Korrekte L&#246;sung f&#252;r Teil 1:*/
TA1:1-(5/6)^3;

/*Optionen f&#252;r Auswahl der Verteilung:*/
Liste2: [
    [1,true,"geometrisch verteilt"],
    [2,false,"hypergeometrisch verteilt"],
    [3,false,"Bernoulli-verteilt"],
    [4,false,"Laplace-verteilt"],
    [5,false,"Poisson-verteilt"],
    [6,false,"binomialverteilt"]
];
Liste2: random_permutation(Liste2);

/*Wahrscheinlichkeitsfunktion (geometrische Verteilung):*/
p_ex(k):=(1-float(TA1))^k*float(TA1);
p_ger(k):=(1-TA1)^k*TA1;

/*Erwartungswert:*/
TA3: (1-TA1)/TA1;

/*Wahrscheinlichkeitsfunktion Formel:*/
p(k):=(1-p)^k*p;

/*Erwartungswert Formel:*/
EWformel: (1-p)/p;

/*Erwartungswert Formel (Z&#228;hler und Nenner vertauscht):*/
EWformel_f: p/(1-p);

/*Varianz:*/
TA4: (1-TA1)/(TA1^2);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@TA1@}</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>TA1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>dropdown</type>
            <tans>Liste2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>TA1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>TA3</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4_z1</name>
            <type>algebraic</type>
            <tans>p(k)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4_z2</name>
            <type>algebraic</type>
            <tans>EWformel</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans5</name>
            <type>algebraic</type>
            <tans>TA4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans1</sans>
                <tans>TA1</tans>
                <testoptions>0.005</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <i>
            Ein
        </i>
        m&#246;glicher Weg die Aufgabe zu l&#246;sen, ist &#252;ber die Binomialverteilung:
    </p>
    <p>
        Bezeichnen wir die Anzahl der in den drei Versuchen gew&#252;rfelten 6-en mit \(Y\), dann hat \(Y\) eine \(\text{Bin}(3,\frac{1}{6})\)-Verteilung. Mit der Wahrscheinlichkeit des Gegenereignisses gilt:
                \[ P(Y\ge 1)=1-P(Y=0)=1-\left(\frac{5}{6} \right)^3 = {@TA1@}\approx {@rnd(TA1)@} \]
    </p>
    <p>
        Alternativ kann man die Aufgabe auch l&#246;sen, indem man das Zufallsexperiment als Laplace-Experiment auffasst und die Anzahl der g&#252;nstigen Ergebnisse durch die Anzahl aller Ergebnisse teilt. Zudem ist es m&#246;glich, auszunutzen, dass es sich bei der Augenzahl des W&#252;rfels in den drei W&#252;rfen um unabh&#228;ngige Ereignisse \(A_1\), \(A_2\) und \(A_3\) handelt, und die Regel \(P(A_1\cap A_2\cap A_3)=P(A_1)\cdot P(A_2)\cdot P(A_3)\) anzuwenden.
    </p>
    <hr/>
    <p>
        Bitte klicken Sie auf den folgenden Button, um zur n&#228;chsten Teilaufgabe zu gelangen:
        <button class="buttonweiter" onclick="
                                hide('oer-aufgabe-4_aufgabe1');
                                show('oer-aufgabe-4_aufgabe2');
                                show('oer-aufgabe-4_div-bekanntes');
                " type="button">
            Weiter
        </button>
    </p>
    <script>
        (function(){
				const ansDiv = document.querySelector("#oer-aufgabe-4_div-bekanntes_innen");
				//ansDiv.style.display = "block";
				const paragraph = document.createElement("li");
				paragraph.append("Die Wahrscheinlichkeit, bei den drei Versuchen zu Beginn mindestens eine 6 zu erhalten, betr&#228;gt "+"\\"+"(p="+{#tex1(TA1)#} +"\\"+"approx"+ {#tex1(rnd(TA1))#} +"\\"+").");
				ansDiv.append(paragraph);
})();
    </script>
</span>
<span class="multilang" lang="en">
    <p>
        <i>
            A
        </i>
        possible way to solve the task is via the binomial distribution:
    </p>
    <p>
        Let us record the number of six in the three attempts with \(Y\), then has \(Y\) a \(\text{Bin}(3,\frac{1}{6})\)- Distribution. The likelihood of the counter-event applies:
    \[ P(Y\ge 1)=1-P(Y=0)=1-\left(\frac{5}{6} \right)^3 = {@TA1@}\approx {@rnd(TA1)@} \]
    </p>
    <p>
        Alternatively, it is also possible to solve the problem by recording the random experiment as a Laplace experiment and dividing the number of favorable results by the number of all results. In addition, it is possible to exploit the fact that the number of dice in the three litters in order to achieve independent events \(A_1\), \(A_2\) and \(A_3\) and the rule \(P(A_1\cap A_2\cap A_3)=P(A_1)\cdot P(A_2)\cdot P(A_3)\) to be applied.
    </p>
    <hr/>
    <p>
        Please click the following button to get to the next part task:
        <button class="buttonweiter" onclick="
                                hide('oer-aufgabe-4_aufgabe1');
                                show('oer-aufgabe-4_aufgabe2');
                                show('oer-aufgabe-4_div-bekanntes');
                " type="button">
            Next
        </button>
    </p>
    <script>
        (function(){
				const ansDiv = document.querySelector("#oer-aufgabe-4_div-bekanntes_innen");
				//ansDiv.style.display = "block";
				const paragraph = document.createElement("li");
				paragraph.append("Die Wahrscheinlichkeit, bei den drei Versuchen zu Beginn mindestens eine 6 zu erhalten, betr&#228;gt "+"\\"+"(p="+{#tex1(TA1)#} +"\\"+"approx"+ {#tex1(rnd(TA1))#} +"\\"+").");
				ansDiv.append(paragraph);
})();
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    [[ if test='is(ans1&lt;0 or ans1&gt;1)' ]]
    <p>
        Ihre Antwort liegt nicht im Intervall \([0,1]\). Wahrscheinlichkeiten m&#252;ssen aber immer einen Wert zwischen 0 und 1 annehmen.
    </p>
    [[/ if ]]
    <p>
        Bitte probieren Sie es erneut und klicken dann noch einmal auf "Pr&#252;fen".
    </p>
    <div id="oer-aufgabe-4_tipp_prt1_button">
        <p>
            Wenn Sie nicht weiterkommen, k&#246;nnen Sie auf den folgenden Button
								klicken:
            <button class="buttontipp" onclick="hide('oer-aufgabe-4_tipp_prt1_button'); show('oer-aufgabe-4_tipp_prt1_div');" type="button">
                Zeige den Tipp
            </button>
        </p>
    </div>
    <div class="divtipp" id="oer-aufgabe-4_tipp_prt1_div" style="display:none">
        <p>
            Tipp:
        </p>
        <p>
            Nutzen Sie die Wahrscheinlichkeit des Gegenereignisses, d. h. \(P\left(A^C\right)=1-P(A)\)
        </p>
    </div>
</span>
<span class="multilang" lang="en">
    [[ if test='is(ans1&lt;0 or ans1&gt;1)' ]]
    <p>
        Your answer is not in interval \([0,1]\). Probabilities must always assume a value between 0 and 1.
    </p>
    [[/ if ]]
    <p>
        Please try it again and then click "Check".
    </p>
    <div id="oer-aufgabe-4_tipp_prt1_button">
        <p>
            If you do not proceed, you can click on the following button
click:
            <button class="buttontipp" onclick="hide('oer-aufgabe-4_tipp_prt1_button'); show('oer-aufgabe-4_tipp_prt1_div');" type="button">
                Show the tip
            </button>
        </p>
    </div>
    <div class="divtipp" id="oer-aufgabe-4_tipp_prt1_div" style="display:none">
        <p>
            Tip:
        </p>
        <p>
            Use the probability of the counter-event, i.e. \(P\left(A^C\right)=1-P(A)\)
        </p>
    </div>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/* === Erkl&#228;rung der Knoten: Wurde ... ausgew&#228;hlt?===
1. Knoten: geometrisch verteilt
2. Knoten: hypergeometrisch verteilt
3. Knoten: Bernoulli-verteilt
4. Knoten: Laplace-verteilt
5. Knoten: Poisson-verteilt
6. Knoten: binomialverteilt
*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Wird ein Zufallsexperiment mehrfach unabh&#228;ngig durchgef&#252;hrt und so lange
                wiederholt, bis ein bestimmtes Ereignis \(A\) erstmals eingetreten ist
                (etwa ein Erfolg bei einem Bernoulli-Experiment, eine 6 beim W&#252;rfeln, oder
                ein 3-er Pasch beim Werfen mit drei W&#252;rfeln), so hat die Anzahl der
                Fehlversuche vor dem ersten Eintreten des Ereignisses eine geometrische
                Verteilung mit Parameter \(p = P(A)\).
    </p>
    <hr/>
    <p>
        Bitte klicken Sie auf den folgenden Button, um zum n&#228;chsten Aufgabenteil zu
                gelangen:
        <button class="buttonweiter" onclick="
                    show('oer-aufgabe-4_aufgabe2_2');
                    hide('oer-aufgabe-4_aufgabe2');" type="button">
            Weiter
        </button>
    </p>
    <script>
        (function () {
                                const ansDiv = document.querySelector(
                                                "#oer-aufgabe-4_div-bekanntes_innen"
                                );
                                const paragraph = document.createElement("li");
                                paragraph.append(
                                                "\\" +
                                                                "(X" +
                                                                "\\" +
                                                                ") (die Anzahl der Runden, die man warten muss, bevor man starten kann) ist geometrisch verteilt."
                                );
                                ansDiv.append(paragraph);
                })();
    </script>
</span>
<span class="multilang" lang="en">
    <p>
        If a random experiment is carried out repeatedly independently and for so long
repeated until a certain event \(A\) entered for the first time
(a success in a Bernoulli experiment, a 6 in dice, or
a 3 pash at throwing with three cubes), the number of
incorrect attempts before the first occurrence of the event a geometric
Distribution with parameters \(p = P(A)\).
    </p>
    <hr/>
    <p>
        Please click on the following button to get to the next task
to:
        <button class="buttonweiter" onclick="
                    show('oer-aufgabe-4_aufgabe2_2');
                    hide('oer-aufgabe-4_aufgabe2');" type="button">
            Next
        </button>
    </p>
    <script>
        (function () {
                                const ansDiv = document.querySelector(
                                                "#oer-aufgabe-4_div-bekanntes_innen"
                                );
                                const paragraph = document.createElement("li");
                                paragraph.append(
                                                "\\" +
                                                                "(X" +
                                                                "\\" +
                                                                ") (die Anzahl der Runden, die man warten muss, bevor man starten kann) ist geometrisch verteilt."
                                );
                                ansDiv.append(paragraph);
                })();
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        \(X\) bezeichnet die Anzahl der Fehlversuche vor dem ersten Erfolg. Schauen Sie in Ihren Unterlagen, im Internet oder in der Literatur nach, welche Verteilung f&#252;r die Modellierung eines solchen Sachverhalts verwendet wird.
    </p>
    <p>
        W&#228;hlen Sie dann diese L&#246;sung aus und klicken Sie erneut auf "Pr&#252;fen".
    </p>
    <hr/>
    <p>
        Zu Ihrer ausgew&#228;hlten Option:
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        \(X\) refers to the number of failures before the first success. In your documents, on the Internet or in the literature, check which distribution is used for modeling such a situation.
    </p>
    <p>
        Then select this solution and click "Check" again.
    </p>
    <hr/>
    <p>
        To your selected option:
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die hypergeometrische Verteilung tritt auf, wenn eine Stichprobe (der Gr&#246;&#223;e \(n\)) aus einer endlichen Grundgesamtheit (der M&#228;chtigkeit \(N\)), von der ein bestimmter Teil (\(R\) Elemente) eine gewisse Eigenschaft besitzt, ohne Zur&#252;cklegen gezogen wird. Die Anzahl der Elemente in der Stichprobe mit der gesuchten Eigenschaft hat dann eine hypergeometrische Verteilung (mit Parametern \((n,R,N)\)).
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The hypergeometric distribution occurs when a sample (size) \(n\)) from a finite population (the power \(N\)), of which a certain part (\(R\) Elements) has a certain property without being drawn back. The number of elements in the sample with the desired property then has a hypergeometric distribution (with parameters) \((n,R,N)\)).
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Bernoulli-verteilte Zufallsvariablen k&#246;nnen nur die Werte 0 oder 1 annehmen. Die Bernoulli-Verteilung hat den Parameter \(p = P(X = 1)\). Solche Zufallsvariablen treten im Zusammenhang mit Experimenten auf, die nur zwei m&#246;gliche Ergebnisse haben (Kopf oder Zahl beim M&#252;nzwurf, Schaden oder kein Schaden bei einer Versicherung), oder bei denen man sich nur daf&#252;r interessiert, ob ein bestimmtes Ereignis eintritt oder nicht (eine 6 beim W&#252;rfeln, ein 3-er Pasch beim 3-fachen W&#252;rfeln). Oft verwendet man die Begriffe "Erfolg" und "Misserfolg" f&#252;r die beiden Ergebnisse, obwohl diese Bezeichnung in vielen Anwendungen unpassend oder gar irref&#252;hrend ist.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Bernoulli distributed random variables can only assume the values 0 or 1. The Bernoulli distribution has the parameter \(p = P(X = 1)\). Such random variables occur in connection with experiments which have only two possible results (head or number in coin throwing, damage or no damage in an insurance), or in which one is only interested in whether a particular event occurs or not (a 6 in the dice, a 3 piece in the 3-fold dice). Often the terms "success" and "mistake" are used for the two results, although this term is inappropriate or even misleading in many applications.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt2-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die Laplace-Verteilung (auch diskrete Gleichverteilung genannt) wird dann verwendet, wenn eine Zufallsvariable \(X\) einen endlichen Wertebereich hat und jeder m&#246;gliche Wert mit derselben Wahrscheinlichkeit angenommen wird.  Bezeichnet \(N\) die Anzahl der m&#246;glichen Werte, so gilt \(p(k)=\frac{1}{N}\) f&#252;r jedes \(k\) aus dem Wertebereich von \(X\). Beispiele sind die Augenzahl beim einmaligen Werfen eines fairen W&#252;rfels oder die Nummer, die beim einmaligen Ziehen aus einer Urne mit \(N\) nummerierten Kugeln gezogen wird.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The Laplace distribution (also called discrete equal distribution) is used when a random variable \(X\) has a finite value range and any possible value is assumed with the same probability. Reference \(N\) the number of possible values, \(p(k)=\frac{1}{N}\) for each \(k\) from the value range of \(X\). Examples are the number of eyes when throwing a fair cube once or the number that, when pulling a single urn with \(N\) numbered balls are pulled.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prt2-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>5</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-5-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Mit der Poisson-Verteilung kann eine Binomialverteilung approximiert werden, wenn die Anzahl der Experimente \(n\) gro&#223; und die Erfolgswahrscheinlichkeit \(p\) sehr klein ist (Gesetz der seltenen Ereignisse). Ist \(np\) sehr gro&#223;, so nimmt man statt der Poisson-Approximation eher die Normalapproximation.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        With the Poisson distribution, a binomial distribution can be approximated if the number of experiments \(n\) large and the probability of success \(p\) is very small (law of rare events). I \(np\) very large, so instead of the Poisson approximation, normal approximation is rather taken.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>5</falsenextnode>
                <falseanswernote>prt2-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>6</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-6-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Mit der Binomialverteilung wird die Anzahl der Erfolge in (\(n\)) unabh&#228;ngigen Bernoulli-Experimenten (mit Erfolgswahrscheinlichkeit \(p\)) modelliert. (Die Anzahl der Erfolge hat dann eine Binomialverteilung mit Parametern \((n, p)\).)
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        With the binomial distribution, the number of successes in (\(n\)) independent Bernoulli experiments (with probability of success) \(p\)) modelled. (The number of successes then has a binomial distribution with parameters \((n, p)\).
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Bitte w&#228;hlen Sie eine der Optionen aus
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Please select one of the options
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/* === Erkl&#228;rung der Knoten: ===
1. Knoten: Ist die Antwort eine Zahl?
2. Knoten: Ist die Antwort richtig?
*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>numberp(ans3)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihre Antwort sollte eine Zahl sein, ist es aber nicht. Bitte geben Sie eine Zahl ein und klicken Sie dann auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your answer should be a number, but it is not. Please enter a number and then click "Check".
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans3</sans>
                <tans>TA1</tans>
                <testoptions>0.005</testoptions>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>1</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Diese Wahrscheinlichkeit \(p\approx{@rnd(TA1)@}\) entspricht der von Ihnen im
                allerersten Aufgabenteil berechneten Wahrscheinlichkeit.
    </p>
    <hr/>
    <p>
        Bitte klicken Sie auf den folgenden Button, um zur n&#228;chsten Teilaufgabe zu
                gelangen:
        <button class="buttonweiter" onclick="
                    show('oer-aufgabe-4_aufgabe3');
                    hide('oer-aufgabe-4_aufgabe2_2');" type="button">
            Weiter
        </button>
    </p>
    <script>
        (function(){
    	    	    	    	const ansDiv = document.querySelector("#oer-aufgabe-4_div-bekanntes_innen");
    	    	    	    	const paragraph = document.createElement("li");
    	    	    	    	paragraph.append("Der Parameter der Verteilung lautet "+"\\"+"(p="+ {#tex1(TA1)#} +"\\"+"approx"+ {#tex1(rnd(TA1))#} +"\\"+")");
    	    	    	    	ansDiv.append(paragraph);
                })();
    </script>
</span>
<span class="multilang" lang="en">
    <p>
        This probability \(p\approx{@rnd(TA1)@}\) in the
first part of the task calculated probability.
    </p>
    <hr/>
    <p>
        Please click on the following button to enter the next part task
to:
        <button class="buttonweiter" onclick="
                    show('oer-aufgabe-4_aufgabe3');
                    hide('oer-aufgabe-4_aufgabe2_2');" type="button">
            Next
        </button>
    </p>
    <script>
        (function(){
    	    	    	    	const ansDiv = document.querySelector("#oer-aufgabe-4_div-bekanntes_innen");
    	    	    	    	const paragraph = document.createElement("li");
    	    	    	    	paragraph.append("Der Parameter der Verteilung lautet "+"\\"+"(p="+ {#tex1(TA1)#} +"\\"+"approx"+ {#tex1(rnd(TA1))#} +"\\"+")");
    	    	    	    	ansDiv.append(paragraph);
                })();
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Tipp: Diese Wahrscheinlichkeit kennen Sie bereits!
    </p>
    <p>
        Bitte probieren Sie es erneut und klicken dann noch einmal auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Tip: You already know this probability!
    </p>
    <p>
        Please try it again and then click "Check".
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/* === Erkl&#228;rung der Knoten: ===
1. Knoten: Ist die Antwort richtig?
2. Knoten: Wurde stattdessen die Varianz angegeben?
*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans4</sans>
                <tans>TA3</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Erwartungswert einer geometrisch verteilten Zufallsvariable \(X\) mit
                Parameter \(p\) kann mit der Formel \( \text{E}(X)= \frac{1-p}{p}\)
                berechnet werden. Daraus ergibt sich in diesem Fall \[ \text{E}(X)=
                \frac{1-{@rnd(TA1)@}}{{@rnd(TA1)@}} \approx {@rnd(TA3)@} \]
    </p>
    <hr/>
    <p>
        Bitte klicken Sie auf den folgenden Button, um zur n&#228;chsten Teilaufgabe zu gelangen:
        <button class="buttonweiter" onclick="
                hide('oer-aufgabe-4_aufgabe3');
                show('oer-aufgabe-4_aufgabe4');
                hide('oer-aufgabe-4_text_nach_schleife_3');
                hide('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
            Weiter
        </button>
    </p>
    <script>
        (function(){
    	    	    	    	const ansDiv = document.querySelector("#oer-aufgabe-4_div-bekanntes_innen");
    	    	    	    	const paragraph = document.createElement("li");
    	    	    	    	paragraph.append("Es gilt "+"\\"+"("+"\\"+"text{E}(X)="+ {#tex1(TA3)#} +"\\"+"approx"+{#tex1(rnd(TA3))#} +"\\"+")");
    	    	    	    	ansDiv.append(paragraph);
                })();
    </script>
</span>
<span class="multilang" lang="en">
    <p>
        The expected value of a geometrically distributed random variable \(X\) with
Parameters \(p\) with the formula \( \text{E}(X)= \frac{1-p}{p}\)are calculated. This results in this case \[ \text{E}(X)=
                \frac{1-{@rnd(TA1)@}}{{@rnd(TA1)@}} \approx {@rnd(TA3)@} \]
    </p>
    <hr/>
    <p>
        Please click the following button to get to the next part task:
        <button class="buttonweiter" onclick="
                hide('oer-aufgabe-4_aufgabe3');
                show('oer-aufgabe-4_aufgabe4');
                hide('oer-aufgabe-4_text_nach_schleife_3');
                hide('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
            Next
        </button>
    </p>
    <script>
        (function(){
    	    	    	    	const ansDiv = document.querySelector("#oer-aufgabe-4_div-bekanntes_innen");
    	    	    	    	const paragraph = document.createElement("li");
    	    	    	    	paragraph.append("Es gilt "+"\\"+"("+"\\"+"text{E}(X)="+ {#tex1(TA3)#} +"\\"+"approx"+{#tex1(rnd(TA3))#} +"\\"+")");
    	    	    	    	ansDiv.append(paragraph);
                })();
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans4</sans>
                <tans>TA4</tans>
                <testoptions>0.04</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben die Varianz von \(X\) anstatt des Erwartungswerts eingegeben.
    </p>
    <p>
        Bitte geben Sie den Erwartungswert von \(X\) ein und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have the variance of \(X\) instead of the expected value entered.
    </p>
    <p>
        Please give the expected value of \(X\) and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <div id="oer-aufgabe-4_aufgabe3_ausSchleife" style="display: none">
        <p>
            Bitte korrigieren Sie Ihre vorherige Antwort und klicken Sie dann auf
                                "Pr&#252;fen".
        </p>
    </div>
    <div id="oer-aufgabe-4_aufgabe3_inSchleife">
        <p>
            Dass Sie die Aufgabe nicht auf Anhieb l&#246;sen konnten, ist kein Problem! Im Folgenden wird die Aufgabe in kleinere Teilaufgaben unterteilt. Wenn Sie diese Aufgaben alle gel&#246;st haben, k&#246;nnen Sie die urspr&#252;ngliche Aufgabe sicherlich l&#246;sen.
        </p>
        <p>
            Bitte klicken Sie zum Fortfahren auf den folgenden Button:
            <button class="buttonweiter" onclick="
                                                show('oer-aufgabe-4_zwischen3a');
                                                hide('oer-aufgabe-4_aufgabe3');
                                                hide('oer-aufgabe-4_aufgabe3_inSchleife');
                                                show('oer-aufgabe-4_aufgabe3_ausSchleife')" type="button">
                Weiter
            </button>
        </p>
        <p>
            Falls Sie sich nur vertippt haben oder eine Idee haben, was der
                                Erwartungswert von \(X\) sein k&#246;nnte, k&#246;nnen Sie das Feld aber
                                auch bereits jetzt neu ausf&#252;llen und auf "Pr&#252;fen" klicken.
        </p>
    </div>
</span>
<span class="multilang" lang="en">
    <div id="oer-aufgabe-4_aufgabe3_ausSchleife" style="display: none">
        <p>
            Please correct your previous reply and then click on
"Check."
        </p>
    </div>
    <div id="oer-aufgabe-4_aufgabe3_inSchleife">
        <p>
            That you couldn't solve the task right away is no problem! In the following, the task is subdivided into smaller subtasks. If you have all solved these tasks, you can certainly solve the original task.
        </p>
        <p>
            Please click on the following button to continue:
            <button class="buttonweiter" onclick="
                                                show('oer-aufgabe-4_zwischen3a');
                                                hide('oer-aufgabe-4_aufgabe3');
                                                hide('oer-aufgabe-4_aufgabe3_inSchleife');
                                                show('oer-aufgabe-4_aufgabe3_ausSchleife')" type="button">
                Next
            </button>
        </p>
        <p>
            If you've just stumbled or have an idea what the
expected value of \(X\) you could be the field but
already refill and click on "Check".
        </p>
    </div>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4_z1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text>/* === Erkl&#228;rung der Knoten: ===
1. Knoten: Ist die Antwort richtig?
2. Knoten: Ist die Antwort eine Zahl?
2. Knoten: Enth&#228;lt die Antwort den Parameter p?
*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4_z1</sans>
                <tans>p(k)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4_z1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="correct">
        <span style="font-size: 1.5em; color: green">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Die
        <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
            Wahrscheinlichkeitsfunktion
        </span>
        einer geometrisch verteilten Zufallsvariable
                \(Z\) in Abh&#228;ngigkeit eines Parameters \(p\) lautet \(P(Z=k)={@p(k)@}\).
    </p>
    <hr/>
    <p>
        Falls Sie einen weiteren Zwischenschritt w&#252;nschen, klicken Sie bitte auf den
                folgenden Button:
        <button class="buttonweiter" onclick="
                hide('oer-aufgabe-4_zwischen3a');
                show('oer-aufgabe-4_zwischen3b');
                hide('oer-aufgabe-4_text_nach_schleife_3');
                hide('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
            Zu einem weiteren Zwischenschritt
        </button>
    </p>
    <p>
        Falls Sie nun erneut versuchen mlchten, den Erwartungswert zu bestimmen,
                klcken Sie bitte auf diesen Button:
        <button class="buttonzurueck" onclick="
                hide('oer-aufgabe-4_zwischen3a');
                show('oer-aufgabe-4_aufgabe3');
                show('oer-aufgabe-4_text_nach_schleife_3');
                show('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
            Zur&#252;ck zur Berechnung von \(\text{E}(X)\)
        </button>
    </p>
    <script>
        (function () {
                                const ansDiv = document.querySelector(
                                                "#oer-aufgabe-4_div-bekanntes_innen"
                                );
                                const paragraph = document.createElement("li");
                                paragraph.append(
                                                "F&#252;r eine geometrisch verteilte Zufallsvariable " +
                                                                "\\" +
                                                                "(Z" +
                                                                "\\" +
                                                                ") mit Parameter " +
                                                                "\\" +
                                                                "(p" +
                                                                "\\" +
                                                                ") gilt " +
                                                                "\\" +
                                                                "(P(Z=k)=p" +
                                                                "\\" +
                                                                "cdot(1-p)^k" +
                                                                "\\" +
                                                                ")"
                                );
                                ansDiv.append(paragraph);
                })();
    </script>
</span>
<span class="multilang" lang="en">
    <p class="correct">
        <span style="font-size: 1.5em; color: green">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        The
        <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
            Probability function
        </span>
        a geometrically distributed random variable
    \(Z\) as a function of a parameter \(p\) read: \(P(Z=k)={@p(k)@}\).
    </p>
    <hr/>
    <p>
        If you want another step, please click on the
the following button:
        <button class="buttonweiter" onclick="
                hide('oer-aufgabe-4_zwischen3a');
                show('oer-aufgabe-4_zwischen3b');
                hide('oer-aufgabe-4_text_nach_schleife_3');
                hide('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
            To a further intermediate step
        </button>
    </p>
    <p>
        If you try again to determine the expected value,
please click on this button:
        <button class="buttonzurueck" onclick="
                hide('oer-aufgabe-4_zwischen3a');
                show('oer-aufgabe-4_aufgabe3');
                show('oer-aufgabe-4_text_nach_schleife_3');
                show('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
            Back to the calculation of \(\text{E}(X)\)
        </button>
    </p>
    <script>
        (function () {
                                const ansDiv = document.querySelector(
                                                "#oer-aufgabe-4_div-bekanntes_innen"
                                );
                                const paragraph = document.createElement("li");
                                paragraph.append(
                                                "F&#252;r eine geometrisch verteilte Zufallsvariable " +
                                                                "\\" +
                                                                "(Z" +
                                                                "\\" +
                                                                ") mit Parameter " +
                                                                "\\" +
                                                                "(p" +
                                                                "\\" +
                                                                ") gilt " +
                                                                "\\" +
                                                                "(P(Z=k)=p" +
                                                                "\\" +
                                                                "cdot(1-p)^k" +
                                                                "\\" +
                                                                ")"
                                );
                                ansDiv.append(paragraph);
                })();
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt4_z1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="incorrect">
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
</span>
<span class="multilang" lang="en">
    <p class="incorrect">
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>numberp(float(ans4_z1))</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4_z1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihre Eingabe ist eine Zahl, Sie sollten aber die Wahrscheinlichkeits
        <i>
            funktion
        </i>
        \(P(Z=k)\) unter der Verwendung eines Parameters \(p\) angeben.
    </p>
    <p>
        Bitte tun Sie dies und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your input is a number, but you should have the probability
        <i>
            Function
        </i>
        \(P(Z=k)\) using a parameter \(p\) indicate.
    </p>
    <p>
        Please do this and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt4_z1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>member(p,setify(listofvars(ans4_z1)))</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4_z1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie k&#246;nnen gerne in Ihren Unterlagen, der Literatur oder im Internet nachschauen, was die
        <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
            Wahrscheinlichkeitsfunktion
        </span>
        einer geometrisch verteilten Zufallsvariable ist.
    </p>
    <p>
        Geben Sie die L&#246;sung dann in das Feld ein und klicken Sie anschlie&#223;end erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You can check in your documents, literature or on the Internet what the
        <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
            Probability function
        </span>
        is a geometrically distributed random variable.
    </p>
    <p>
        Then enter the solution in the field and then click "Check".
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4_z1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihre Funktion enth&#228;lt nicht den Parameter \(p\). Eventuell haben Sie einen numerischen Wert f&#252;r \(p\) eingesetzt, was aber
        <i>
            nicht
        </i>
        gefordert war.
    </p>
    <p>
        Sie k&#246;nnen gerne in Ihren Unterlagen, der Literatur oder im Internet nachschauen, was die
        <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
            Wahrscheinlichkeitsfunktion
        </span>
        einer geometrisch verteilten Zufallsvariable ist.
    </p>
    <p>
        Geben Sie die L&#246;sung dann in das Feld ein und klicken Sie anschlie&#223;end erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your function does not include the parameter \(p\). You may have a numerical value for \(p\) but what
        <i>
            not
        </i>
        was required.
    </p>
    <p>
        You can check in your documents, literature or on the Internet what the
        <span class="stack-tooltip" data-placement="top" data-toggle="tooltip" title='Auch bekannt unter "Z&#228;hldichte"'>
            Probability function
        </span>
        is a geometrically distributed random variable.
    </p>
    <p>
        Then enter the solution in the field and then click "Check".
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt4_z2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text>/* === Erkl&#228;rung der Knoten: ===
Knoten 1: Ist die Formel korrekt?
Knoten 2: Wurden Z&#228;hler und Nenner vertauscht?
Knoten 3: Ist die Antwort eine Zahl?
*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4_z2</sans>
                <tans>EWformel</tans>
                <testoptions>0.005</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4_z2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="correct">
        <span style="font-size: 1.5em; color: green">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Es gilt \(\text{E}(Z)={@EWformel@}\). Nun sollten Sie bereit sein, die
                urspr&#252;ngliche Frage nach dem Erwartungswert von \(X\) zu beantworten.
    </p>
    <hr/>
    <p>
        Bitte klicken Sie auf den folgenden Button, um zur Berechnung von
                \(\text{E}(X)\) zur&#252;ckzukehren:
        <button class="buttonzurueck" onclick="
                hide('oer-aufgabe-4_zwischen3b');
                show('oer-aufgabe-4_aufgabe3');
                show('oer-aufgabe-4_bekanntes-3b');
                show('oer-aufgabe-4_text_nach_schleife_3');
                hide('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
            Zur&#252;ck zur urspr&#252;nglichen Aufgabe
        </button>
    </p>
    <script>
        (function () {
                                const ansDiv = document.querySelector(
                                                "#oer-aufgabe-4_div-bekanntes_innen"
                                );
                                const paragraph = document.createElement("li");
                                paragraph.append(
                                                "F&#252;r eine geometrisch verteilte Zufallsvariable " +
                                                                "\\" +
                                                                "(Z" +
                                                                "\\" +
                                                                ") mit Parameter " +
                                                                "\\" +
                                                                "(p" +
                                                                "\\" +
                                                                ") gilt " +
                                                                "\\" +
                                                                "(" +
                                                                "\\" +
                                                                "text{E}(Z)=" +
                                                                "\\" +
                                                                "frac{1-p}{p}" +
                                                                "\\" +
                                                                ")"
                                );
                                ansDiv.append(paragraph);
                })();
    </script>
</span>
<span class="multilang" lang="en">
    <p class="correct">
        <span style="font-size: 1.5em; color: green">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        The following shall apply: \(\text{E}(Z)={@EWformel@}\). Now you should be ready to
original question of the expected value of \(X\) to answer.
    </p>
    <hr/>
    <p>
        Please click on the following button to calculate
    \(\text{E}(X)\) to return:
        <button class="buttonzurueck" onclick="
                hide('oer-aufgabe-4_zwischen3b');
                show('oer-aufgabe-4_aufgabe3');
                show('oer-aufgabe-4_bekanntes-3b');
                show('oer-aufgabe-4_text_nach_schleife_3');
                hide('oer-aufgabe-4_aufgabe3_nochmal-inSchleife');" type="button">
            Back to the original task
        </button>
    </p>
    <script>
        (function () {
                                const ansDiv = document.querySelector(
                                                "#oer-aufgabe-4_div-bekanntes_innen"
                                );
                                const paragraph = document.createElement("li");
                                paragraph.append(
                                                "F&#252;r eine geometrisch verteilte Zufallsvariable " +
                                                                "\\" +
                                                                "(Z" +
                                                                "\\" +
                                                                ") mit Parameter " +
                                                                "\\" +
                                                                "(p" +
                                                                "\\" +
                                                                ") gilt " +
                                                                "\\" +
                                                                "(" +
                                                                "\\" +
                                                                "text{E}(Z)=" +
                                                                "\\" +
                                                                "frac{1-p}{p}" +
                                                                "\\" +
                                                                ")"
                                );
                                ansDiv.append(paragraph);
                })();
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt4_z2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="incorrect">
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
</span>
<span class="multilang" lang="en">
    <p class="incorrect">
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4_z2</sans>
                <tans>EWformel_f</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4_z2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Kann es sein, dass Sie den Z&#228;hler und den Nenner vertauscht haben?
    </p>
    <p>
        Bitte korrigieren Sie Ihre L&#246;sung und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Can it be that you swapped the counter and the denominator?
    </p>
    <p>
        Please correct your solution and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt4_z2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>numberp(float(ans4_z2))</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt4_z2-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihre Eingabe ist eine Zahl, Sie sollten aber den Parameter \(p\) verwenden. Eventuell haben Sie f&#252;r \(p\) eine Zahl eingesetzt, dies war aber explizit
        <i>
            nicht
        </i>
        gefordert.
    </p>
    <p>
        Bitte geben Sie die allgemeine Formel f&#252;r \(\text{E}(Z)\) ein und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your input is a number, but you should \(p\) use. You may have \(p\) a number, but this was explicitly
        <i>
            not
        </i>
        demanded.
    </p>
    <p>
        Please give the general formula for \(\text{E}(Z)\) and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt4_z2-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie k&#246;nnen in Ihren Unterlagen, der Literatur oder dem Internet nach einer Formel suchen, in der der Erwartungswert einer geometrisch verteilten Zufallsvariablen als Funktion des Parameters der Verteilung dargestellt wird.
    </p>
    <p>
        Geben Sie die L&#246;sung dann in das Feld ein und klicken Sie anschlie&#223;end erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You can search in your documents, literature or the Internet for a formula in which the expected value of a geometrically distributed random variable is displayed as a function of the parameter of the distribution.
    </p>
    <p>
        Then enter the solution in the field and then click "Check".
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt5</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/* === Erkl&#228;rung der Knoten: ===
1. Knoten: Ist die Antwort richtig?
2. Knoten: Wurde stattdessen der Erwartungswert angegeben?
*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans5</sans>
                <tans>TA4</tans>
                <testoptions>0.04</testoptions>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die Varianz einer geometrisch verteilten Zufallsvariable \(X\) mit Parameter
                \(p\) kann mit der Formel \( \text{Var}(X)= \frac{1-p}{p^2}\) berechnet
                werden. Daraus ergibt sich in diesem Fall \[ \text{Var}(X)=
                \frac{1-{@rnd(TA1)@}}{{@rnd(TA1)@}^2} \approx {@rnd(TA4)@} \]
    </p>
    <hr/>
    <p>
        Herzlichen Gl&#252;ckwunsch, Sie sind am Ende der Aufgabe angekommen! Sie sollten nun erkennen k&#246;nnen, wenn eine Zufallsvariable geometrisch verteilt ist und den Parameter der Verteilung sowie den Erwartungswert und die Varianz der Zufallsvariablen berechnen k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The variance of a geometrically distributed random variable \(X\) with parameters
    \(p\) with the formula \( \text{Var}(X)= \frac{1-p}{p^2}\) Calculator
,. This results in this case \[ \text{Var}(X)=
                \frac{1-{@rnd(TA1)@}}{{@rnd(TA1)@}^2} \approx {@rnd(TA4)@} \]
    </p>
    <hr/>
    <p>
        Congratulations, you have arrived at the end of the task! You should now be able to recognize if a random variable is geometrically distributed and can calculate the parameter of the distribution as well as the expected value and the variance of the random variables.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt5-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans><![CDATA[is(ans5>=0)]]></sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt5-4-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Ihre Antwort ist eine negative Zahl. Varianzen m&#252;ssen aber immer gr&#246;&#223;er oder gleich null sein.
    </p>
    <p>
        Bitte versuchen Sie es noch einmal und klicken dann auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Your answer is a negative number. Variances must always be greater or equal to zero.
    </p>
    <p>
        Please try again and then click "Check".
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans5</sans>
                <tans>TA3</tans>
                <testoptions>0.01</testoptions>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt5-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben den Erwartungswert von \(X\) anstatt der Varianz eingegeben.
    </p>
    <p>
        Bitte geben Sie die Varianz von \(X\) ein und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
    <p>
        Sie k&#246;nnen in Ihren Unterlagen, der Literatur oder dem Internet nach einer Formel suchen, in der die Varianz einer geometrisch verteilten Zufallsvariablen als Funktion des Parameters der Verteilung dargestellt wird.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have the expected value of \(X\) instead of entering the variance.
    </p>
    <p>
        Please give the variance of \(X\) and then click "Check" again.
    </p>
    <p>
        You can search in your documents, literature or the Internet for a formula in which the variance of a geometrically distributed random variable is displayed as a function of the parameter of the distribution.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt5-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie k&#246;nnen in Ihren Unterlagen, der Literatur oder dem Internet nach einer Formel suchen, in der die Varianz einer geometrisch verteilten Zufallsvariablen als Funktion des Parameters der Verteilung dargestellt wird.
    </p>
    <p>
        Bitte geben Sie die richtige Antwort dann in das Eingabefeld ein und klicken dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You can search in your documents, literature or the Internet for a formula in which the variance of a geometrically distributed random variable is displayed as a function of the parameter of the distribution.
    </p>
    <p>
        Please enter the correct answer in the input field and then click "Check".
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>