<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Die t&#228;gliche Radfahrt zur Universit&#228;t (t-Test)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <style>
        button.top-next,
button.next {
                border: 1px solid black;
                background-color: #A9F5BC;
}




button.top-next:hover,
button.next:hover {
                background-color: #E0F8E6;
}




button.top,
button.prev {
                border: 1px solid black;
                background-color: #DCB25E;
}
button.top:hover,
button.prev:hover {
                background-color: #E7D2A7;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
hr.partsep {
                border-top: thin solid black;
                width: 80%;
}




.que.stack .formulation .memory .stackprtfeedback {
                border: 1px solid gray;
                border-radius: 10px;
                background: transparent;
                padding: 5px;
}
    </style>
    <details>
        <summary>
            Inhalt und Umfang dieser Aufgabe
        </summary>
        <p>
            In dieser Aufgabe lernen Sie, einen t-Test mit zwei Stichproben
                        durchzuf&#252;hren.  Dieser Test vergleicht den Erwartungswert zweier
                        Verteilungen.
        </p>
        <ul>
            <li>
                <b>
                    (a)
                </b>
                Geeignete Teststatistik ausw&#228;hlen
            </li>
            <li>
                <b>
                    (b)
                </b>
                Wert der Teststatistik bestimmen
            </li>
            <li>
                <b>
                    (c)
                </b>
                Kritischen Wert bestimmen
            </li>
            <li>
                <b>
                    (d)
                </b>
                Testentscheidung treffen
            </li>
        </ul>
    </details>
    <div class="question-view">
        <p>
            Ein Professor f&#228;hrt t&#228;glich mit dem Fahrrad die 8.8 km lange
                Strecke von seiner Wohnung zur Universit&#228;t. Im Januar 2020
                ist er an {# size_x #} Tagen gefahren; dabei war die mittlere
                Fahrzeit \(\bar{x} = {# mean_x #}\) min und die Stichprobenvarianz
                \(s_x^2 = {# var_x #}\) min\(^2\).  Im August 2020 ist er erneut
                an {# size_y #} Tagen gefahren; diesmal war die mittlere Fahrzeit
                \(\bar{y}={# mean_y #}\) min und die Stichprobenvarianz \(s_y^2 =
                {# var_y #}\) min
            <sup>
                2
            </sup>
            .  Wir gehen davon aus, dass die
                Fahrzeiten unabh&#228;ngig und normalverteilt sind und dass die
                Varianzen in beiden Monaten gleich waren.  Wir wollen die
                Hypothese testen, dass der Erwartungswert in beiden Monaten gleich
                war gegen die Alternative, dass der Erwartungswert im August 2020
                gr&#246;&#223;er war.
        </p>
        <hr class="partsep"/>
        <div class="question-view">
            <p>
                <b>
                    (a)
                </b>
                W&#228;hlen Sie eine geeignete Teststatistik f&#252;r dieses Testproblem aus.
            </p>
            <p>
                [[input:ans1]][[validation:ans1]]
            </p>
            [[feedback:prt_ans1]]
        </div>
        <div class="question-view">
            <div class="memory">
                [[feedback:vorinfos]]
            </div>
            <p>
                <b>
                    (b)
                </b>
                Bestimmen Sie den Wert Ihrer Teststatistik.
            </p>
            <p>
                Antwort: [[input:ans2]] [[validation:ans2]]
            </p>
            [[feedback:prt_ans2]]
            <p>
                <em>
                    Hinweis: Wenn Sie gerundete Werte eingeben, dann runden Sie bitte
        auf zwei Nachkommastellen.
                </em>
            </p>
        </div>
        <div class="question-view">
            <p>
                <b>
                    (c)
                </b>
                Bestimmen Sie den kritischen Wert zu Ihrer Teststatistik zum
                        Niveau {@ alpha = aspercent(niveau) @}.
            </p>
            <p>
                Antwort: [[input:ans3]] [[validation:ans3]]
            </p>
            [[feedback:prt_ans3]]
            <p>
                <em>
                    Hinweis: Wenn Sie gerundete Werte eingeben, dann runden Sie bitte
        auf zwei Nachkommastellen.
                </em>
            </p>
        </div>
        <div class="question-view">
            <div class="memory">
                [[feedback:ergebnisse]]
            </div>
            <p>
                <b>
                    (d)
                </b>
                Wie lautet Ihre Testentscheidung, wenn zum Niveau {@ alpha = aspercent(niveau) @} getestet werden soll?
            </p>
            <p>
                Antwort: [[input:ans4]] [[validation:ans4]]
            </p>
            [[feedback:prt_ans4]]
        </div>
        <script>
            /*
    * This program is free software: you can redistribute it and/or
    * modify it under the terms of the GNU General Public License as
    * published by the Free Software Foundation, either version 3 of the
    * License, or (at your option) any later version.
    *
    * This program is distributed in the hope that it will be useful, but
    * WITHOUT ANY WARRANTY; without even the implied warranty of
    * MERCHANTABILITY or FITNESS FOR A PARTICULAR PURPOSE.  See the GNU
    * General Public License for more details.
    *
    * You should have received a copy of the GNU General Public License
    * along with this program.  If not, see <https://www.gnu.org/licenses/>.
    *
    * Copyright (C) 2021 by Daniel Mei&#223;ner <daniel.meissner-i4k@ruhr-uni-bochum.de>
    * SPDX-License-Identifier: GPL-3.0-or-later
    *
    */




/* @license magnet:?xt=urn:btih:1f739d935676111cfff4b4693e3816e664797050&dn=gpl-3.0.txt GPL-v3-or-Later */




(function() {
/* Make sure to create QController instances only once even if the
            script is loaded multiple times.  This is so that the script can
            be included in more than one STACK question on the page and still
            function properly. */
if (window.QController === undefined) {
        window.QController = QController;
        document.addEventListener("DOMContentLoaded", event => {
                const responseform = document.getElementById("responseform");
                if (!responseform) {
                        throw Error("Element with id responseform not found.");
                }

    
    
    
    responseform.querySelectorAll(".stack.que").forEach(stackQue => {
                        const rootElement = stackQue.querySelector(".question-view");

      
      
      
      if (!rootElement) {
                                console.warn("QController: no elements with class question-view inside " + stackQue.id);
                                return;
                        }

      
      
      
      const controller = new QController(rootElement, stackQue.id);
                        controller.installHandlers();
                        controller.loadState();
                        controller.update();
                });
        });
}




/* This class handles the navigation of a single STACK question */
function QController(rootElement, questionId) {
        this.rootElement = rootElement;
        this.id = questionId;
        this.views = rootElement.querySelectorAll(".question-view");
        this.visited = (new Array(this.views.length)).fill(false);
        this.currentView = Array.prototype.findIndex.call(this.views, isLeaf);
        this.visited[this.currentView] = true;
}




QController.prototype.handleVisitedNodes = function(current) {
        var visitedSiblings = [];
        var el = current;
        while (el.nextElementSibling) {
                var index =
                                Array.prototype.indexOf.call(this.views, el.nextElementSibling);
                visitedSiblings.push(this.visited[index]);
                el = el.nextElementSibling;
        }

  
  
  
  var withoutHelp = visitedSiblings.every(x => x == false);
        var helpCompleted = visitedSiblings.every(x => x == true);

  
  
  
  current.querySelectorAll(".visited").forEach(function(element) {
                var index = element.getAttribute("data-qcontroller-help");
                if (!index || index === "none")
                        element.hidden = !withoutHelp;
                else if (index === "complete")
                        element.hidden = !helpCompleted;
                else if (index === "partial")
                        element.hidden = withoutHelp || helpCompleted;
                else
                        element.hidden = !visitedSiblings[index - 1];
        });
};




/* Update view according to current controller state */
QController.prototype.update = function() {
        this.views.forEach(hide);

  
  
  
  const current = this.views[this.currentView];
        if (isLeaf(current)) {
                showView(current);
                this.handleVisitedNodes(current);
        } else {
                this.showFirstView(current);
        }
};




/* Install event handlers for all buttons under root element */
QController.prototype.installHandlers = function() {
        const controller = this;
        this.rootElement.querySelectorAll("button").forEach(button => {
                button.addEventListener("click", ev => {
                        const view = button.closest(".question-view");

      
      
      
      if (!view) {
                                console.log("qcontroller: <button> not inside question view");
                                return;
                        }

      
      
      
      const classList = button.classList;
                        let forward = null;

      
      
      
      if (classList.contains("next")) {
                                forward = view.nextElementSibling;
                        } else if (classList.contains("top-next")) {
                                const parent = view.parentNode;
                                if (parent) {
                                        forward = parent.nextElementSibling;
                                } else {
                                        forward = null;
                                }
                        } else if (classList.contains("top")) {
                                forward = containingView(view);
                        } else if (classList.contains("prev")) {
                                forward = view.previousElementSibling;
                        }

      
      
      
      if (forward && forward.classList.contains("question-view")) {
                                controller.visit(forward);
                                controller.saveState();
                                controller.update();
                        }
                });
        });
};




QController.prototype.hasVisited = function(viewIndex) {
        return this.visited[viewIndex];
};




QController.prototype.visit = function(nextView) {
        const index = Array.prototype.findIndex.call(
                this.views,
                x => x === nextView);
        this.currentView = index;
        this.visited[index] = true;
};




/* Save controller state to current session */
QController.prototype.saveState = function() {
        const key = this.id + "-controller-state";
        const state = {
                currentView: this.currentView,
                visited: this.visited
        };
        window.sessionStorage.setItem(key, JSON.stringify(state));
};




/* Load controller state from current session */
QController.prototype.loadState = function() {
        const key = this.id + "-controller-state";
        const state = JSON.parse(window.sessionStorage.getItem(key));

  
  
  
  if (state) {
                this.currentView = state.currentView;
                this.visited = state.visited;
        }
};




/* Show the first leaf question view */
QController.prototype.showFirstView = function(questionView) {
        let view = questionView;

  
  
  
  while (view && !isLeaf(view)) {
                view = view.querySelector(".question-view");
        }

  
  
  
  showView(view);
        this.handleVisitedNodes(view);
}




function showView(view) {
        /* show all containing views */
        while (view) {
                show(view);
                view = containingView(view);
        }
}




/* Return the view containing VIEW */
function containingView(view) {
        let parent = view.parentNode;

  
  
  
  if (parent) {
                parent = parent.closest(".question-view");
        }

  
  
  
  return parent;
}




function show(view) {
        view.style.display = "block";
}




function isLeaf(view) {
        return view.querySelector(".question-view") == null;
}




function isVisible(view) {
        return view.style.display != "none";
}




function hide(view) {
        view.style.display = "none";
}




})();
/* @license-end */
        </script>
    </div>
</span>
<span class="multilang" lang="en">
    <style>
        button.top-next,
button.next {
                border: 1px solid black;
                background-color: #A9F5BC;
}




button.top-next:hover,
button.next:hover {
                background-color: #E0F8E6;
}




button.top,
button.prev {
                border: 1px solid black;
                background-color: #DCB25E;
}
button.top:hover,
button.prev:hover {
                background-color: #E7D2A7;
}
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
hr.partsep {
                border-top: thin solid black;
                width: 80%;
}




.que.stack .formulation .memory .stackprtfeedback {
                border: 1px solid gray;
                border-radius: 10px;
                background: transparent;
                padding: 5px;
}
    </style>
    <details>
        <summary>
            Content and scope of this task
        </summary>
        <p>
            In this task you will learn a t test with two samples
to be carried out. This test compares the expected value of two
Distributions.
        </p>
        <ul>
            <li>
                <b>
                    (a)
                </b>
                Select suitable test statistics
            </li>
            <li>
                <b>
                    (b)
                </b>
                Determine value of test statistics
            </li>
            <li>
                <b>
                    (c)
                </b>
                Determine critical value
            </li>
            <li>
                <b>
                    (d)
                </b>
                Making a decision
            </li>
        </ul>
    </details>
    <div class="question-view">
        <p>
            A professor drives the 8.8 km long by bike daily
Route from his apartment to university. In January 2020
he is {# size_x #} driving days; the middle
Driving time \(\bar{x} = {# mean_x #}\) min and sample variance
    \(s_x^2 = {# var_x #}\) min\(^2\). In August 2020 he is again
on {# size_y #} driving days; this time was the middle driving time
    \(\bar{y}={# mean_y #}\) min and sample variance \(s_y^2 =
                {# var_y #}\) min
            <sup>
                2
            </sup>
            . We assume that the
driving times are independent and normally distributed and that
Variances were the same in both months. We want
Hypothesis test that the expected value is equal in both months
was against the alternative that the expected value in August 2020
was bigger.
        </p>
        <hr class="partsep"/>
        <div class="question-view">
            <p>
                <b>
                    (a)
                </b>
                Select a suitable test statistics for this test problem.
            </p>
            <p>
                [[input:ans1]][[validation:ans1]]
            </p>
            [[feedback:prt_ans1]]
        </div>
        <div class="question-view">
            <div class="memory">
                [[feedback:vorinfos]]
            </div>
            <p>
                <b>
                    (b)
                </b>
                Determine the value of your test statistics.
            </p>
            <p>
                Answer: [[input:ans2]] [[validation:ans2]]
            </p>
            [[feedback:prt_ans2]]
            <p>
                <em>
                    Note: If you enter rounded values, please round
on two offsprings.
                </em>
            </p>
        </div>
        <div class="question-view">
            <p>
                <b>
                    (c)
                </b>
                Determine the critical value for your test statistics
Level {@ alpha = aspercent(niveau) @}.
            </p>
            <p>
                Answer: [[input:ans3]] [[validation:ans3]]
            </p>
            [[feedback:prt_ans3]]
            <p>
                <em>
                    Note: If you enter rounded values, please round
on two offsprings.
                </em>
            </p>
        </div>
        <div class="question-view">
            <div class="memory">
                [[feedback:ergebnisse]]
            </div>
            <p>
                <b>
                    (d)
                </b>
                What is your test decision if to the level {@ alpha = aspercent(niveau) @} to be tested?
            </p>
            <p>
                Answer: [[input:ans4]] [[validation:ans4]]
            </p>
            [[feedback:prt_ans4]]
        </div>
        <script>
            /*
    * This program is free software: you can redistribute it and/or
    * modify it under the terms of the GNU General Public License as
    * published by the Free Software Foundation, either version 3 of the
    * License, or (at your option) any later version.
    *
    * This program is distributed in the hope that it will be useful, but
    * WITHOUT ANY WARRANTY; without even the implied warranty of
    * MERCHANTABILITY or FITNESS FOR A PARTICULAR PURPOSE.  See the GNU
    * General Public License for more details.
    *
    * You should have received a copy of the GNU General Public License
    * along with this program.  If not, see <https://www.gnu.org/licenses/>.
*
* Copyright (C) 2021 by Daniel Mei&#223;ner <daniel.meissner-i4k@ruhr-uni-bochum.de>* SPDX-License Identifier: GPL-3.0-or-later
*
*/




/* @license magnet:?xt=urn:btih:1f739d935676111cfff4b4693e3816e664797050&dn=gpl-3.0.txt GPL-v3-orLater */




(function() {
/* Make sure to create QController instances only once even if the
            script is loaded multiple times.  This is so that the script can
            be included in more than one STACK question on the page and still
            function properly. */
if (window.QController === undefined) {
        window.QController = QController;
        document.addEventListener("DOMContentLoaded", event => {
                const responseform = document.getElementById("responseform");
                if (!responseform) {
                        throw Error("Element with id responseform not found.");
                }

    
    
    
    responseform.querySelectorAll(".stack.que").forEach(stackQue => {
                        const rootElement = stackQue.querySelector(".question-view");

      
      
      
      if (!rootElement) {
                                console.warn("QController: no elements with class question-view inside " + stackQue.id);
                                return;
                        }

      
      
      
      const controller = new QController(rootElement, stackQue.id);
                        controller.installHandlers();
                        controller.loadState();
                        controller.update();
                });
        });
}




/* This class handles the navigation of a single STACK question */
function QController(rootElement, questionId) {
        this.rootElement = rootElement;
        this.id = questionId;
        this.views = rootElement.querySelectorAll(".question-view");
        this.visited = (new Array(this.views.length)).fill(false);
        this.currentView = Array.prototype.findIndex.call(this.views, isLeaf);
        this.visited[this.currentView] = true;
}




QController.prototype.handleVisitedNodes = function(current) {
        var visitedSiblings = [];
        var el = current;
        while (el.nextElementSibling) {
                var index =
                                Array.prototype.indexOf.call(this.views, el.nextElementSibling);
                visitedSiblings.push(this.visited[index]);
                el = el.nextElementSibling;
        }

  
  
  
  var withoutHelp = visitedSiblings.every(x => x == false);
        var helpCompleted = visitedSiblings.every(x => x == true);

  
  
  
  current.querySelectorAll(".visited").forEach(function(element) {
                var index = element.getAttribute("data-qcontroller-help");
                if (!index || index === "none")
                        element.hidden = !withoutHelp;
                else if (index === "complete")
                        element.hidden = !helpCompleted;
                else if (index === "partial")
                        element.hidden = withoutHelp || helpCompleted;
                else
                        element.hidden = !visitedSiblings[index - 1];
        });
};




/* Update view according to current controller state */
QController.prototype.update = function() {
        this.views.forEach(hide);

  
  
  
  const current = this.views[this.currentView];
        if (isLeaf(current)) {
                showView(current);
                this.handleVisitedNodes(current);
        } else {
                this.showFirstView(current);
        }
};




/* Install event handlers for all buttons under root element */
QController.prototype.installHandlers = function() {
        const controller = this;
        this.rootElement.querySelectorAll("button").forEach(button => {
                button.addEventListener("click", ev => {
                        const view = button.closest(".question-view");

      
      
      
      if (!view) {
                                console.log("qcontroller: <button> not inside question view");
                                return;
                        }

      
      
      
      const classList = button.classList;
                        let forward = null;

      
      
      
      if (classList.contains("next")) {
                                forward = view.nextElementSibling;
                        } else if (classList.contains("top-next")) {
                                const parent = view.parentNode;
                                if (parent) {
                                        forward = parent.nextElementSibling;
                                } else {
                                        forward = null;
                                }
                        } else if (classList.contains("top")) {
                                forward = containingView(view);
                        } else if (classList.contains("prev")) {
                                forward = view.previousElementSibling;
                        }

      
      
      
      if (forward && forward.classList.contains("question-view")) {
                                controller.visit(forward);
                                controller.saveState();
                                controller.update();
                        }
                });
        });
};




QController.prototype.hasVisited = function(viewIndex) {
        return this.visited[viewIndex];
};




QController.prototype.visit = function(nextView) {
        const index = Array.prototype.findIndex.call(
                this.views,
                x => x === nextView);
        this.currentView = index;
        this.visited[index] = true;
};




/* Save controller state to current session */
QController.prototype.saveState = function() {
        const key = this.id + "-controller-state";
        const state = {
                currentView: this.currentView,
                visited: this.visited
        };
        window.sessionStorage.setItem(key, JSON.stringify(state));
};




/* Load controller state from current session */
QController.prototype.loadState = function() {
        const key = this.id + "-controller-state";
        const state = JSON.parse(window.sessionStorage.getItem(key));

  
  
  
  if (state) {
                this.currentView = state.currentView;
                this.visited = state.visited;
        }
};




/* Show the first leaf question view */
QController.prototype.showFirstView = function(questionView) {
        let view = questionView;

  
  
  
  while (view && !isLeaf(view)) {
                view = view.querySelector(".question-view");
        }

  
  
  
  showView(view);
        this.handleVisitedNodes(view);
}




function showView(view) {
        /* show all containing views */
        while (view) {
                show(view);
                view = containingView(view);
        }
}




/* Return the view containing VIEW */
function containingView(view) {
        let parent = view.parentNode;

  
  
  
  if (parent) {
                parent = parent.closest(".question-view");
        }

  
  
  
  return parent;
}




function show(view) {
        view.style.display = "block";
}




function isLeaf(view) {
        return view.querySelector(".question-view") == null;
}




function isVisible(view) {
        return view.style.display != "none";
}




function hide(view) {
        view.style.display = "none";
}




}),
/* @license-end */
        </script>
    </div>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>10.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/* This question uses the Maxima package distrib */

/* render number as percentages */
texaspercent (e) := block([num: first (args (e))], sconcat(tex1(ev(num*100, simp)), "\\,\\%"));
texput(aspercent, texaspercent);

texqt(e) := block([q: first(args(e)), df: second(args(e))],
  sconcat("t_{", tex1(df), "; ", tex1(q), "}"));
texput(qt, texqt);

/* bars for means of X and Y*/
texput(Xbar, "\\bar{X}");
texput(Ybar, "\\bar{Y}");

size_x: 20;
size_y: 20;

means: sort(rand_selection([26, 28, 29], 2));
mean_x: means[1];
mean_y: means[2];

var_x: 8;
var_y: 12;

niveau: 5/100;

simp:true;
var_comb: block([tt1, tt2],
  tt1: (size_x - 1) / (size_x + size_y - 2),
  tt2: (size_y - 1) / (size_x + size_y - 2),
  tt1 * var_x + tt2 * var_y);
statistic: (mean_y - mean_x) / sqrt((1/size_x + 1/size_y) * var_comb);
critical: quantile_student_t(1 - niveau, size_x + size_y - 2);
simp: false;

rejectionPlot: block([nu: ev(size_x + size_y - 2, simp)],
  plot(
    [[discrete, [critical, critical], [0, 1]],
    pdf_student_t(x, nu)],
    [x, mean_student_t(nu) - 5*var_student_t(nu), mean_student_t(nu) + 5*var_student_t(nu)],
    [xlabel, "x"],
    [ylabel, sconcat("t_{", nu, "}(x)")]));

choices: block([vals, correct, incorrect],
  correct: [T = (Ybar - Xbar) / sqrt((1/m + 1/n) *s_P^2)],
  incorrect: [
  T = sum(X[i] - mu[0], i, 1, n)/sqrt(n * s_X^2),
  T = (n-1)* s_X^2 / sigma_0^2,
  F = s_Y^2 / s_X^2,
  X = sum((N[i] - n*theta[i])^2 / (n*theta[i]), i, 1, k)],
  vals: makelist(i, i, ev(length(correct) + length(incorrect), simp)),
  zip_with(
    lambda([a, b],
    [a, if a <= length(correct) then true else false, stack_disp(b,"id")]),
  vals, append(correct, incorrect)));

test_result: [
  ["Hypothese verwerfen", is(statistic >= critical)],
  ["Hypothese nicht verwerfen", is(statistic < critical)]];

Tans1: random_permutation (choices);
Tans2: statistic;
Tans3: critical;
Tans4: random_permutation (test_result);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>\(\bar{y} = {# mean_y #}, \bar{x} = {# mean_x #}\)</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>0</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist nicht komplett falsch.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is not completely wrong.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed>i</variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>radio</type>
            <tans>Tans1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>decimalplaces(Tans2, 2)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>numerical</type>
            <tans>decimalplaces(Tans3, 2)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>dropdown</type>
            <tans>Tans4</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>ergebnisse</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>[ans2, ans3]</sans>
                <tans>[Tans2, Tans3]</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ergebnisse-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        In den vorherigen Aufgabenteilen haben Sie den Wert der
                                        Teststatistik und den kritischen Wert wie folgt bestimmt:

          
          
          
          \begin{align*}
                                        T &amp;= {@ans2@}\\
                                        t_{{@size_y + size_x - 2@};{@ev(niveau,float)@}} &amp;= {@ans3@}
                                        \end{align*}
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the previous tasks you have the value of
Test statistics and the critical value as follows:

 
 
 
 \begin{align*}
                                        T &amp;= {@ans2@}\\
                                        t_{{@size_y + size_x - 2@};{@ev(niveau,float)@}} &amp;= {@ans3@}
                                        \end{align*}
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ergebnisse-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt_ans1</name>
            <value>3.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>first(mcq_correct(Tans1))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die ausgew&#228;hlte Teststatistik berechnet den Unterschied im
arithmetischen Mittel der Stichproben und normiert das Ergebnis
mithilfe der
        <em>
            kombinierten Standardabweichung \({s_P}^2\)
        </em>
        .
Dies stellt sicher, dass gro&#223;e Werte der Teststatistik nicht auf eine
gro&#223;e Streuung in den Stichproben zur&#252;ckzuf&#252;hren ist.
    </p>
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        ,
um zum n&#228;chsten Aufgabenteil zu gelangen
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The selected test statistics calculate the difference
arithmetic means of the samples and normalizes the result
using:
        <em>
            combined standard deviation \({s_P}^2\)
        </em>
        .
This ensures that large values of test statistics are not applied to
large spread in the sample.
    </p>
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        ,
to get to the next task
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        [[if test="is(ans1=2)"]]




Dies ist ein Test, der den Erwartungswert einer Stichprobe mit einem
vorgegebenen Wert \(\mu_0\) vergleicht.  Lesen Sie noch einmal den
Einleitungstext und achten Sie darauf, wie viele Stichproben
vorkommen.




[[elif test="ans1=3"]]




Dies ist ein Test, der die Varianz einer Stichprobe mit einem
vorgegebenen Wert \(\sigma_0^2\) vergleicht.  Lesen Sie noch einmal
den Einleitungstext und achten Sie darauf, welcher Parameter in der
Hypothese und wie viele Stichproben vorkommen.




[[elif test="ans1=4"]]




Diese Teststatistik geh&#246;rt zum F-Test zum Vergleich der Varianzen
zweier Stichproben.  Dieser Test ist hier allerdings nicht angebracht.
Lesen Sie noch einmal den Einleitungstext und achten Sie darauf,
welcher Parameter in der Hypothese vorkommt.




[[elif test="ans1=5"]]




Der ausgew&#228;hlte Test testet Hypothesen &#252;ber die Verteilung einer
vorliegenden Stichprobe.  Lesen Sie noch einmal den Einleitungstext
und achten Sie darauf, welcher Parameter in der Hypothese vorkommt.




[[/if]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        [[if test="is(ans1=2)"]]This is a test that shows the expected value of a sample with a
predefined value \(\mu_0\) compares. Read again the
Introduction and make sure how many samples
to be found.




[[elif test="ans1=3"]]This is a test that shows the variance of a sample with a
predefined value \(\sigma_0^2\) compares. Read more
the introduction text and make sure that the parameters in the
Hypothesis and how many samples occur.




[[elif test="ans1=4"]]These test statistics are part of the F test for comparing the variances
two samples. However, this test is not appropriate here.
Read the introduction text again and make sure
which parameter occurs in the hypothesis.




[[elif test="ans1=5"]]The selected test tests hypotheses on the distribution of a
this sample. Read the introduction text again
and make sure which parameter is present in the hypothesis.




[[/if]]
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt_ans2</name>
            <value>3.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>2</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans2</sans>
                <tans>Tans2</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Wir setzen die in der Aufgabe gegebenen Angaben in die Formel f&#252;r
                        die passende Teststatistik ein:
    </p>
    \[
                {@ T = (Ybar - Xbar) / sqrt((1/m + 1/n) *s_P^2) @}
                = {@ (mean_y - mean_x) / sqrt((1/size_y + 1/size_x) * var_comb) @}
                = {@ ev(mean_y - mean_x, simp) / sqrt(ev((1/size_y + 1/size_x) * var_comb, simp)) @}
                = {@ ans2 @}
                \]
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        ,
um zum n&#228;chsten Aufgabenteil zu gelangen
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        We put the information given in the task in the formula for
the appropriate test statistics:
    </p>
    \[
                {@ T = (Ybar - Xbar) / sqrt((1/m + 1/n) *s_P^2) @}
                = {@ (mean_y - mean_x) / sqrt((1/size_y + 1/size_x) * var_comb) @}
                = {@ ev(mean_y - mean_x, simp) / sqrt(ev((1/size_y + 1/size_x) * var_comb, simp)) @}
                = {@ ans2 @}
                \]
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        ,
to get to the next task
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie m&#252;ssen die im Text gegebenen Daten f&#252;r die empirischen
Varianzen in die Formel f&#252;r die Teststatistik einsetzen.  Die richtige
Formel kennen Sie bereits:




\[
{@ T = (Ybar - Xbar) / sqrt((1/m + 1/n) * s_P^2) @}
\]
    </p>
    <p>
        Hier bezeichnet \({s_P}^2\) die
        <em>
            kombinierte
Stichprobenvarianz
        </em>
        . Sie ist gegeben durch:




\[
{s_P}^2 = \frac{m-1}{m+n-2} {s_X}^2 + \frac{n-1}{m+n-2} {s_Y}^2
\]
    </p>
    <p>
        Bitte versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You need the information given in the text for the empirical
Use variances in the formula for test statistics. The Right
You already know formula:




\[
{@ T = (Ybar - Xbar) / sqrt((1/m + 1/n) * s_P^2) @}
\]
    </p>
    <p>
        Reference \({s_P}^2\) the
        <em>
            Combined
Sample variance
        </em>
        . It is given by:




\[
{s_P}^2 = \frac{m-1}{m+n-2} {s_X}^2 + \frac{n-1}{m+n-2} {s_Y}^2
\]
    </p>
    <p>
        Please try again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt_ans3</name>
            <value>2.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>3</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans3</sans>
                <tans>Tans3</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der kritische Wert ist durch das {@aspercent(1-niveau)@}-Quantil der
                {@t[size_y + size_x -2]@}-Verteilung gegeben.
                \[
                {@ qt(decimalplaces(1-niveau, 2), size_y + size_x - 2) = ans3@}
                \]
    </p>
    <p>
        Sie k&#246;nnen diesen Wert entweder aus einer passenden Tabelle
                ablesen, oder den R-Befehl
        <a href="https://cran.r-project.org/doc/manuals/r-release/R-intro.html#R-as-a-set-of-statistical-tables" target="_blank">
            <code>
                qt({# decimalplaces(1-niveau, 2) #}, {#
                size_x + size_y - 2 #})
            </code>
        </a>
        verwenden.
    </p>
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        , um zum n&#228;chsten Aufgabenteil zu
                gelangen
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The critical value is due to {@aspercent(1-niveau)@}-Quality of
    {@t[size_y + size_x -2]@}- Distribution.
    \[
                {@ qt(decimalplaces(1-niveau, 2), size_y + size_x - 2) = ans3@}
                \]
    </p>
    <p>
        You can use this value either from a matching table
read or read the R command
        <a href="https://cran.r-project.org/doc/manuals/r-release/R-intro.html#R-as-a-set-of-statistical-tables" target="_blank">
            <code>
                q.({# decimalplaces(1-niveau, 2) #}, {#
                size_x + size_y - 2 #})
            </code>
        </a>
        use.
    </p>
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        to the next task
the
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Lesen Sie noch einmal den Einleitungstext und achten Sie besonders
        darauf, welche Alternative genannt wird.  Handelt es sich um eine
        einseitige oder zweiseitige Alternative?  Schauen Sie dann ggf. noch
        einmal ins Vorlesungsskript oder in die Literatur.
    </p>
    <p>
        Falls Sie mehr Hilfe bei der Antwort ben&#246;tigen, klicken Sie
        auf
        <button class="help" type="button">
            Hilfe
        </button>
        .
    </p>
    <div style="display:none">
        <p>
            Da es sich um eine einseitige Alternative handelt, ist der
                        kritische Wert durch das {@aspercent(1-niveau)@}-Quantil der
                        {@t[size_x + size_y-2]@}-Verteilung gegeben.  Diesen Wert k&#246;nnen Sie
                        aus der entsprechenden Tabelle ablesen oder mithilfe eines
                        Statistikprogrammes wie
            <a href="https://www.r-project.org">
                R
            </a>
            berechnen.
        </p>
        <p>
            Die folgende Grafik veranschaulicht den Ablehnungsbereich des
                Tests.  Wenn \(T\) Werte annimmt, die auf der x-Achse rechts der
                blauen Linie liegen, verwirft der Test die Hypothese.
        </p>
        {@ rejectionPlot @}
    </div>
    <script>
        (function() {
                const script = document.currentScript;
                const help = script.parentNode.querySelector('div[style="display:none"]');
                const btn = script.parentNode.querySelector('button[type="button"]');
                btn.addEventListener("click", ev => {
                        help.style.display = "block";
                });
                })();
    </script>
</span>
<span class="multilang" lang="en">
    <p>
        Read the introduction text again and pay particular attention
what alternative is called. It is about
one-sided or two-side alternative? See if necessary
once in the lecture script or in the literature.
    </p>
    <p>
        If you need more help in the answer, click
on
        <button class="help" type="button">
            Help
        </button>
        .
    </p>
    <div style="display:none">
        <p>
            This is a one-sided alternative.
critical value due to the {@aspercent(1-niveau)@}-Quality of
    {@t[size_x + size_y-2]@}- Distribution. This value you can
read from the relevant table or by means of
Statistics programmes such as
            <a href="https://www.r-project.org">
                R
            </a>
            calculate.
        </p>
        <p>
            The following graph illustrates the rejection area of the
Tests. If \(T\) assuming values on the x axis to the right of the
lying blue line, the test rejects the hypothesis.
        </p>
        {@ rejectionPlot @}
    </div>
    <script>
        (function() {
                const script = document.currentScript;
                const help = script.parentNode.querySelector('div[style="display:none"]');
                const btn = script.parentNode.querySelector('button[type="button"]');
                btn.addEventListener("click", ev => {
                        help.style.display = "block";
                });
                })();
    </script>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt_ans4</name>
            <value>2.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans4</sans>
                <tans>first(mcq_correct(Tans4))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt_ans4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Herzlichen Gl&#252;ckwunsch, Sie sind am Ende der Aufgabe angelangt!
    </p>
    <h5>
        Zusammenfassung
    </h5>
    <p>
        In dieser Aufgabe wurde der t-Test zum Vergleich der
                                        Erwartungswerte zweier Stichproben thematisiert.  Es sollte
                                        die Hypothese getestet werden, dass der Erwartungswert der
                                        Fahrzeiten in beiden Monaten gleich gro&#223; war, gegen die
                                        Alternative, dass der Erwartungswert im August gr&#246;&#223;er war.
                                        Im ersten Aufgabenteil haben Sie die daf&#252;r notwendige
                                        Teststatistik ausgew&#228;hlt:

  
  
  
  \[
            T = \frac{\bar{Y} - \bar{X}}{\sqrt{(\frac{1}{m} + \frac{1}{n}) {s_P}^2}}
        \]

  
  
  
  Dieser Test eignet sich f&#252;r die vorliegende Situation, da davon
        ausgegangen wird, dass die Varianzen gleich gro&#223; sind.
    </p>
    <p>
        Als n&#228;chstes mussten Sie den Wert dieser Teststatistik
        ausrechnen. Sie haben dazu die entsprechenden Zahlenwerte aus dem
        Einleitungstext herausgelesen und in die Formel eingesetzt:

  
  
  
  \[
        T = \frac{\bar{Y} - \bar{X}}{\sqrt{(\frac{1}{m} + \frac{1}{n}) {s_P}^2}}
                = {@ (mean_y - mean_x) / sqrt((1/size_y + 1/size_x) * var_comb) @}
                = {@ ans2 @}
        \]
    </p>
    <p>
        Um im letzten Aufgabenteil die Testentscheidung treffen zu
        k&#246;nnen, wird der kritische Wert ben&#246;tigt. Hierbei kommt es auf die
        Alternative an.  In unserer Aufgabe haben wir es mit einer
        <em>
            einseitigen
        </em>
        Alternative zu tun. Daher ist der kritische
        Wert durch das \(1-\alpha\)-Quantil der {@t[size_y + size_x
        -2]@}-Verteilung gegeben. Diesen k&#246;nnen Sie aus einer passenden
        Tabelle ablesen:

  
  
  
  \[
        t_{{@size_y+size_x-2@}; {@niveau@}} = {@ ans3 @}
        \]
    </p>
    <p>
        Am Ende verwerfen Sie also die Hypothese, falls {@T &gt;=
        ans3@} ist.

  
  
  
  [[if test="statistic &gt;= critical"]]

  
  
  
  Wir verwerfen also die Hypothese, dass die Erwartungswerte der
        beiden Stichproben gleich gro&#223; sind.  Da statistische Tests den
        Fehler 1. Art minimieren, k&#246;nnen wir uns bei dieser Entscheidung
        ziemlich sicher sein.

  
  
  
  [[else]]

  
  
  
  Wir behalten also die Hypothese, dass die Erwartungswerte der beiden
        Stichproben gleich gro&#223; sind, bei.  Wir m&#252;ssen uns aber im klaren
        sein, dass wir keine Kontrolle &#252;ber den Fehler 2. Art haben.  Daher
        m&#252;ssen wir die Testentscheidung mit Vorsicht betrachten.

  
  
  
  [[/if]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Congratulations, you have reached the end of the task!
    </p>
    <h5>
        Summary
    </h5>
    <p>
        In this task, the t-test was compared to the
expectation values of two samples are thematized. It should
the hypothesis that the expected value of the
driving times in both months were the same, against which
Alternative that the expected value was greater in August.
In the first part you have the necessary
Test statistics selected:

 
 
 
 \[
            T = \frac{\bar{Y} - \bar{X}}{\sqrt{(\frac{1}{m} + \frac{1}{n}) {s_P}^2}}
        \]This test is suitable for the present situation, since
it is assumed that the variances are equal.
    </p>
    <p>
        Next, you had the value of this test statistics
Calculate. You have the corresponding numerical values from
Introductory text read out and used in the formula:

 
 
 
 \[
        T = \frac{\bar{Y} - \bar{X}}{\sqrt{(\frac{1}{m} + \frac{1}{n}) {s_P}^2}}
                = {@ (mean_y - mean_x) / sqrt((1/size_y + 1/size_x) * var_comb) @}
                = {@ ans2 @}
        \]
    </p>
    <p>
        To make the test decision in the last part of the task
the critical value is required. In this case,
Alternative. In our task, we have a
        <em>
            single side
        </em>
        Alternative to do. Therefore, the critical
Value by \(1-\alpha\)-Quality of {@t[size_y + size_x
        -2]@}- Distribution. You can use this from a matching
Tables:

 
 
 
 \[
        t_{{@size_y+size_x-2@}; {@niveau@}} = {@ ans3 @}
        \]
    </p>
    <p>
        In the end you will discard the hypothesis if {@T &gt;=
        ans3@} is.

 
 
 
 [[if test="statistic &gt;= critical"]]We therefore reject the hypothesis that the expected values of the
both samples are equal. Since statistical tests
1. Kindly minimized, we can make this decision
be pretty sure.

 
 
 
 [[else]]We therefore retain the hypothesis that the expected values of the two
Sampling is the same size. We need to be clear
be that we have no control over the 2nd type error. Therefore,
we must consider the test decision with caution.

 
 
 
 [[/if]]
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt_ans4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Schauen Sie noch einmal in den
                                        Vorlesungsnotizen bzw. in der Literatur nach dem
                                        Verwerfungsbereich des Tests. Aus dem vorherigen
                                        Aufgabenteil wissen Sie bereits, dass es sich hier um eine
                                        einseitige Alternative handelt.
    </p>
    <p>
        Bitte versuchen Sie
                                        es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Check again in the
Lecture notes or in literature according to
Dismissal range of the test. From the previous
You already know that this is a
a single alternative.
    </p>
    <p>
        Please try
it again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>vorinfos</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>first(mcq_correct(Tans1))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im vorherigen Aufgabenteil haben Sie die folgende Teststatistik
ausgew&#228;hlt: \[{@ T = (Ybar - Xbar) / sqrt((1/m + 1/n) *s_P^2) @}\]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the previous section you have the following test statistics
selected: \[{@ T = (Ybar - Xbar) / sqrt((1/m + 1/n) *s_P^2) @}\]
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>69730508</deployedseed>
        <deployedseed>1882914288</deployedseed>
        <deployedseed>207746449</deployedseed>
        <deployedseed>703555420</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>first(mcq_correct(Tans1))</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>1.5</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>2.17</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>first(mcq_correct(Tans4))</value>
            </testinput>
            <expected>
                <name>prt_ans1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt_ans2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>2-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt_ans3</name>
                <expectedscore>0.5000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt_ans3-6-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt_ans4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt_ans4-6-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>