<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Kartenspiel</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <style>
        .buttonweiter {
                                border: 1px solid black;
                                background-color: #a9f5bc;
                }

    
    
    
    .buttonweiter:hover {
                                background-color: #e0f8e6;
                }

    
    
    
    .buttonzurueck {
                                border: 1px solid black;
                                background-color: #dcb25e;
                }

    
    
    
    .buttonzurueck:hover {
                                background-color: #e7d2a7;
                }

    
    
    
    .buttontipp {
                                border: 1px solid black;
                                background-color: #c9d3f6;
                }

    
    
    
    .buttontipp:hover {
                                background-color: #dbe1f9;
                }

    
    
    
    .erinnerungsbox {
                                border: 2px rgb(152, 202, 62) solid;
                                padding: 5px;
                                margin: 5px;
                }

    
    
    
    .que.stack .formulation details {
                                border: 1px solid #aaa;
                                border-radius: 4px;
                                padding: .5em .5em 0;
                }

    
    
    
    .que.stack .formulation summary {
                                font-weight: bold;
                                margin: -.5em -.5em 0;
                                padding: .5em;
                }

    
    
    
    .que.stack .formulation details[open] {
                                padding: .5em;
                }
    </style>
    <div id="summarybox">
        <details>
            <summary>
                Inhalt und Umfang dieser Aufgabe
            </summary>
            <p>
                In dieser Aufgabe lernen Sie, wie Sie die hypergeometrische Verteilung zur L&#246;sung kombinatorischer Fragestellungen nutzen k&#246;nnen.
            </p>
        </details>
        <hr/>
    </div>
    <div id="aufgabentitel" style="display:box;">
        <b>
            Kartenspiel
        </b>
        <div id="aufgabe" style="display:box;">
            <b>
                <br/>
            </b>
            <table style="width: 100%;">
                <thead>
                    <tr>
                        <th scope="col">
                            <span style="font-weight: normal;">
                                Ein gut durchgemischter Stapel mit je 16 roten und schwarzen Karten wird in zwei gleichgro&#223;e Stapel geteilt. Wie gro&#223; ist die Wahrscheinlichkeit, dass sich in dem linken Stapel \({@k@}\) rote Karten befinden?
                                <br/>
                                <i>
                                    Runden Sie Ihr Ergebnis auf 3 Nachkommastellen.
                                </i>
                            </span>
                        </th>
                    </tr>
                </thead>
                <tbody>
                </tbody>
            </table>
            <br/>
            <p>
                <u>
                    Antwort:
                </u>
                Die Wahrscheinlichkeit betr&#228;gt [[input:ans1]]. [[validation:ans1]]
                <br/>
            </p>
            Bitte klicken Sie nach dem Eintragen der L&#246;sung auf den Button "Pr&#252;fen".
            <br/>
            [[feedback:prt1]]
        </div>
        <div id="zwischen0b" style="display:none;">
            <span>
                <br/>
                <table style="border-width: 2px; border-style: solid; border-color: rgb(152, 202, 62);">
                    <thead>
                        <tr>
                            <th scope="col">
                                <span style="font-weight: normal;">
                                    <span>
                                        <u>
                                            Urspr&#252;ngliche Fragestellung:
                                        </u>
                                        Ein gut durchgemischter Stapel mit je 16 roten und schwarzen Karten wird in zwei gleichgro&#223;e Stapel geteilt. Wie gro&#223; ist die Wahrscheinlichkeit, dass sich in den kleineren Stapeln \({@k@}\)  rote und \({@n-k@}\) schwarze Karten befinden?
                                        <br/>
                                        <i>
                                            Runden Sie Ihr Ergebnis auf 3 Nachkommastellen.
                                        </i>
                                    </span>
                                </span>
                            </th>
                        </tr>
                    </thead>
                    <tbody>
                    </tbody>
                </table>
                <br/>
            </span>
            Aus dem urspr&#252;nglichen Stapel mit 32 Karten werden also zwei Stapel mit je 16 Karten gebildet und es gen&#252;gt aus Symmetriegr&#252;nden, nur einen Stapel zu betrachten, d.h. wir fragen uns, ob in diesem einen Stapel \({@k@}\) rote und \({@n-k@}\) schwarze Karten sind. Daf&#252;r wiederum
                ist es ausreichend, sich z.B. nur die Anzahl der roten Karten (in diesem kleineren Stapel) anzusehen. Damit liegt die klassische Situation f&#252;r eine Stichprobenentnahme (aus einer zweiteiligen Grundgesamtheit, hier: roten und schwarzen Karten) vor.
                Welche Verteilung ist daher geeignet, um die Situation zu modellieren?
            <br/>
            [[input:ans0b]] [[validation:ans0b]]
            <br/>
            <br/>
            Bitte klicken Sie nach dem Eintragen der L&#246;sung auf den Button "Pr&#252;fen".
            <br/>
            [[feedback:prt0b]]
        </div>
        <div id="zwischen0c" style="display:none;">
            <span>
                <br/>
                <table style="border-width: 2px; border-style: solid; border-color: rgb(152, 202, 62);">
                    <thead>
                        <tr>
                            <th scope="col">
                                <span style="font-weight: normal;">
                                    <span>
                                        <u>
                                            Urspr&#252;ngliche Fragestellung:
                                        </u>
                                        Ein gut durchgemischter Stapel mit je 16 roten und schwarzen Karten wird in zwei gleichgro&#223;e Stapel geteilt. Wie gro&#223; ist die Wahrscheinlichkeit, dass sich in den kleineren Stapeln \({@k@}\)  rote und \({@n-k@}\) schwarze Karten befinden?
                                        <br/>
                                        <i>
                                            Runden Sie Ihr Ergebnis auf 3 Nachkommastellen.
                                        </i>
                                    </span>
                                </span>
                            </th>
                        </tr>
                    </thead>
                    <tbody>
                    </tbody>
                </table>
                <br/>
            </span>
            Aus dem urspr&#252;nglichen Stapel mit 32 Karten werden also zwei Stapel mit je 16 Karten gebildet und es gen&#252;gt aus Symmetriegr&#252;nden, nur einen Stapel zu betrachten, d.h. wir fragen uns, ob in diesem einen Stapel  \({@k@}\)  rote und \({@n-k@}\) schwarze Karten sind. Daf&#252;r wiederum
                ist es ausreichend, sich z.B. nur die Anzahl der roten Karten (in diesem kleineren Stapel) anzusehen. Damit liegt die klassische Situation f&#252;r eine Stichprobenentnahme (aus einer zweiteiligen Grundgesamtheit, hier: roten und schwarzen Karten) vor,
                d.h. f&#252;r die hypergeometrische Verteilung \(P(Z=k):=H_{n,N,R}({k})=\frac{{R\choose k} {N-R\choose n-k}}{{N\choose n}}, 0\leq k\leq n\). Welche Aussage gilt in Bezug auf den Aufgabenkontext?
            <br/>
            [[input:ansZwischen0c]] [[validation:ansZwischen0c]]
            <br/>
            <br/>
            Bitte klicken Sie nach dem Eintragen der L&#246;sung auf den Button "Pr&#252;fen".
            <br/>
            [[feedback:ansZwischen0c]]
        </div>
        <div id="zwischen1" style="display:none;">
            <span>
                <br/>
                <table style="border-width: 2px; border-style: solid; border-color: rgb(152, 202, 62);">
                    <thead>
                        <tr>
                            <th scope="col">
                                <span style="font-weight: normal;">
                                    <span>
                                        <u>
                                            Urspr&#252;ngliche Fragestellung:
                                        </u>
                                        Ein gut durchgemischter Stapel mit je 16 roten und schwarzen Karten wird in zwei gleichgro&#223;e Stapel geteilt. Wie gro&#223; ist die Wahrscheinlichkeit, dass sich in den kleineren Stapeln \({@k@}\)  rote und \({@n-k@}\) schwarze Karten befinden?
                                        <br/>
                                        <i>
                                            Runden Sie Ihr Ergebnis auf 3 Nachkommastellen.
                                        </i>
                                    </span>
                                </span>
                            </th>
                        </tr>
                    </thead>
                    <tbody>
                    </tbody>
                </table>
                <br/>
            </span>
            Aus dem urspr&#252;nglichen Stapel mit 32 Karten werden also zwei Stapel mit je 16 Karten gebildet und es gen&#252;gt aus Symmetriegr&#252;nden, nur einen Stapel zu betrachten, d.h. wir fragen uns, ob in diesem einen Stapel  \({@k@}\)  rote und \({@n-k@}\) schwarze Karten sind. Daf&#252;r wiederum
                ist es ausreichend, sich z.B. nur die Anzahl der roten Karten (in diesem kleineren Stapel) anzusehen. Damit liegt die klassische Situation f&#252;r eine Stichprobenentnahme (aus einer zweiteiligen Grundgesamtheit, hier: roten und schwarzen Karten) vor,
                d.h. f&#252;r die hypergeometrische Verteilung \(P(Z=k):=H_{n,N,R}({k})=\frac{{R\choose k} {N-R\choose n-k}}{{N\choose n}}, 0\leq k\leq n\). Hierbei bezeichnet \(Z\) die Zufallsvariable, die die Anzahl der Erfolge (z.B. Ziehen von \(k=1\) schwarzen Karte)
                beim Ziehen von \(n\) Karten aus einer Grundgesamtheit von \(N\) Karten z&#228;hlt, von denen \(R\) Karten die gew&#252;nschte Eigenschaft haben (z.B. schwarze Karte).
            <br/>
            Die Anzahl aller Karten im Stapel bezeichnen wir ab jetzt mit \(N\).
            <br/>
            <br/>
            Bestimmen Sie zun&#228;chst \(n\), also den Stichprobenumfang.
            <br/>
            \(n=\) [[input:ans2]] [[validation:ans2]]
            <br/>
            <br/>
            Bitte klicken Sie nach dem Eintragen der L&#246;sung auf den Button "Pr&#252;fen".
            <br/>
            [[feedback:prt2]]
        </div>
        <div id="zwischen2" style="display:none;">
            <span>
                <br/>
                <table style="border-width: 2px; border-style: solid; border-color: rgb(152, 202, 62);">
                    <thead>
                        <tr>
                            <th scope="col">
                                <span style="font-weight: normal;">
                                    <span>
                                        <u>
                                            Urspr&#252;ngliche Fragestellung:
                                        </u>
                                        Ein gut durchgemischter Stapel mit je 16 roten und schwarzen Karten wird in zwei gleichgro&#223;e Stapel geteilt. Wie gro&#223; ist die Wahrscheinlichkeit, dass sich in den kleineren Stapeln \({@k@}\)  rote und \({@n-k@}\) schwarze Karten befinden?
                                        <br/>
                                        <i>
                                            Runden Sie Ihr Ergebnis auf 3 Nachkommastellen.
                                        </i>
                                    </span>
                                </span>
                            </th>
                        </tr>
                    </thead>
                    <tbody>
                    </tbody>
                </table>
                <br/>
            </span>
            Bestimmen Sie nun \(R\), also die Anzahl der roten Karten im Ursprungsstapel.
            <br/>
            <br/>
            \(R=\) [[input:ans3]] [[validation:ans3]]
            <br/>
            <br/>
            Bitte klicken Sie nach dem Eintragen der L&#246;sung auf den Button "Pr&#252;fen".
            <br/>
            [[feedback:prt3]]
        </div>
    </div>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
<span class="multilang" lang="en">
    <style>
        .buttonweiter {
                                border: 1px solid black;
                                background-color: #a9f5bc;
                }

    
    
    
    .buttonweiter:hover {
                                background-color: #e0f8e6;
                }

    
    
    
    .buttonzurueck {
                                border: 1px solid black;
                                background-color: #dcb25e;
                }

    
    
    
    .buttonzurueck:hover {
                                background-color: #e7d2a7;
                }

    
    
    
    .buttontipp {
                                border: 1px solid black;
                                background-color: #c9d3f6;
                }

    
    
    
    .buttontipp:hover {
                                background-color: #dbe1f9;
                }

    
    
    
    .erinnerungsbox {
                                border: 2px rgb(152, 202, 62) solid;
                                padding: 5px;
                                margin: 5px;
                }

    
    
    
    .que.stack .formulation details {
                                border: 1px solid #aaa;
                                border-radius: 4px;
                                padding: .5em .5em 0;
                }

    
    
    
    .que.stack .formulation summary {
                                font-weight: bold;
                                margin: -.5em -.5em 0;
                                padding: .5em;
                }

    
    
    
    .que.stack .formulation details[open] {
                                padding: .5em;
                }
    </style>
    <div id="summarybox">
        <details>
            <summary>
                Content and scope of this task
            </summary>
            <p>
                In this task, you will learn how to use the hypergeometric distribution to solve combinatorial questions.
            </p>
        </details>
        <hr/>
    </div>
    <div id="aufgabentitel" style="display:box;">
        <b>
            Card game
        </b>
        <div id="aufgabe" style="display:box;">
            <b>
                <br/>
            </b>
            <table style="width: 100%;">
                <thead>
                    <tr>
                        <th scope="col">
                            <span style="font-weight: normal;">
                                A well-mixed stack with 16 red and black cards is divided into two equally-sized stacks. How great is the probability that in the left stack \({@k@}\) red cards?
                                <br/>
                                <i>
                                    Round your result to 3 postage points.
                                </i>
                            </span>
                        </th>
                    </tr>
                </thead>
                <tbody>
                </tbody>
            </table>
            <br/>
            <p>
                <u>
                    Answer:
                </u>
                The probability is [[input:ans1]]. [[validation:ans1]]
                <br/>
            </p>
            Please click the "Check" button after entering the solution.
            <br/>
            [[feedback:prt1]]
        </div>
        <div id="zwischen0b" style="display:none;">
            <span>
                <br/>
                <table style="border-width: 2px; border-style: solid; border-color: rgb(152, 202, 62);">
                    <thead>
                        <tr>
                            <th scope="col">
                                <span style="font-weight: normal;">
                                    <span>
                                        <u>
                                            Original question:
                                        </u>
                                        A well-mixed stack with 16 red and black cards is divided into two equally-sized stacks. How great is the probability that in the smaller stacks \({@k@}\) red and red \({@n-k@}\) black cards?
                                        <br/>
                                        <i>
                                            Round your result to 3 postage points.
                                        </i>
                                    </span>
                                </span>
                            </th>
                        </tr>
                    </thead>
                    <tbody>
                    </tbody>
                </table>
                <br/>
            </span>
            Therefore, two stacks with 16 cards are formed from the original stack with 32 cards and, for reasons of symmetry, it is sufficient to consider only one stack, i.e. we ask ourselves whether a stack is present in this stack. \({@k@}\) red and red \({@n-k@}\) black cards. In turn,
it is sufficient to consider, for example, only the number of red cards (in this smaller stack). Thus, the classical situation for sampling (from a two-part population, here: red and black cards) is present.
Which distribution is therefore suitable to model the situation?
            <br/>
            [[input:ans0b]] [[validation:ans0b]]
            <br/>
            <br/>
            Please click the "Check" button after entering the solution.
            <br/>
            [[feedback:prt0b]]
        </div>
        <div id="zwischen0c" style="display:none;">
            <span>
                <br/>
                <table style="border-width: 2px; border-style: solid; border-color: rgb(152, 202, 62);">
                    <thead>
                        <tr>
                            <th scope="col">
                                <span style="font-weight: normal;">
                                    <span>
                                        <u>
                                            Original question:
                                        </u>
                                        A well-mixed stack with 16 red and black cards is divided into two equally-sized stacks. How great is the probability that in the smaller stacks \({@k@}\) red and red \({@n-k@}\) black cards?
                                        <br/>
                                        <i>
                                            Round your result to 3 postage points.
                                        </i>
                                    </span>
                                </span>
                            </th>
                        </tr>
                    </thead>
                    <tbody>
                    </tbody>
                </table>
                <br/>
            </span>
            Therefore, two stacks with 16 cards are formed from the original stack with 32 cards and, for reasons of symmetry, it is sufficient to consider only one stack, i.e. we ask ourselves whether a stack is present in this stack. \({@k@}\) red and red \({@n-k@}\) black cards. In turn,
it is sufficient to consider, for example, only the number of red cards (in this smaller stack). Thus, the classical situation for sampling (from a two-part population, here: red and black maps) is present,
i.e. for hypergeometric distribution \(P(Z=k):=H_{n,N,R}({k})=\frac{{R\choose k} {N-R\choose n-k}}{{N\choose n}}, 0\leq k\leq n\). What is the statement regarding the task context?
            <br/>
            [[input:ansZwischen0c]] [[validation:ansZwischen0c]]
            <br/>
            <br/>
            Please click the "Check" button after entering the solution.
            <br/>
            [[feedback:ansZwischen0c]]
        </div>
        <div id="zwischen1" style="display:none;">
            <span>
                <br/>
                <table style="border-width: 2px; border-style: solid; border-color: rgb(152, 202, 62);">
                    <thead>
                        <tr>
                            <th scope="col">
                                <span style="font-weight: normal;">
                                    <span>
                                        <u>
                                            Original question:
                                        </u>
                                        A well-mixed stack with 16 red and black cards is divided into two equally-sized stacks. How great is the probability that in the smaller stacks \({@k@}\) red and red \({@n-k@}\) black cards?
                                        <br/>
                                        <i>
                                            Round your result to 3 postage points.
                                        </i>
                                    </span>
                                </span>
                            </th>
                        </tr>
                    </thead>
                    <tbody>
                    </tbody>
                </table>
                <br/>
            </span>
            Therefore, two stacks with 16 cards are formed from the original stack with 32 cards and, for reasons of symmetry, it is sufficient to consider only one stack, i.e. we ask ourselves whether a stack is present in this stack. \({@k@}\) red and red \({@n-k@}\) black cards. In turn,
it is sufficient to consider, for example, only the number of red cards (in this smaller stack). Thus, the classical situation for sampling (from a two-part population, here: red and black maps) is present,
i.e. for hypergeometric distribution \(P(Z=k):=H_{n,N,R}({k})=\frac{{R\choose k} {N-R\choose n-k}}{{N\choose n}}, 0\leq k\leq n\). Reference \(Z\) the random variable, which the number of successes (e.g. \(k=1\) black card)
when pulling \(n\) Maps from a population of \(N\) maps, of which \(R\) Cards have the desired property (e.g. black card).
            <br/>
            The number of cards in the stack is now \(N\).
            <br/>
            <br/>
            First determine \(n\), i.e. the sample size.
            <br/>
            \(n=\) [[input:ans2]] [[validation:ans2]]
            <br/>
            <br/>
            Please click the "Check" button after entering the solution.
            <br/>
            [[feedback:prt2]]
        </div>
        <div id="zwischen2" style="display:none;">
            <span>
                <br/>
                <table style="border-width: 2px; border-style: solid; border-color: rgb(152, 202, 62);">
                    <thead>
                        <tr>
                            <th scope="col">
                                <span style="font-weight: normal;">
                                    <span>
                                        <u>
                                            Original question:
                                        </u>
                                        A well-mixed stack with 16 red and black cards is divided into two equally-sized stacks. How great is the probability that in the smaller stacks \({@k@}\) red and red \({@n-k@}\) black cards?
                                        <br/>
                                        <i>
                                            Round your result to 3 postage points.
                                        </i>
                                    </span>
                                </span>
                            </th>
                        </tr>
                    </thead>
                    <tbody>
                    </tbody>
                </table>
                <br/>
            </span>
            Define now \(R\)the number of red cards in the original stack.
            <br/>
            <br/>
            \(R=\) [[input:ans3]] [[validation:ans3]]
            <br/>
            <br/>
            Please click the "Check" button after entering the solution.
            <br/>
            [[feedback:prt3]]
        </div>
    </div>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2023010400</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
  Kartenspiel

  wurde entwickelt von
  
    Riko Kelter <riko.kelter(at)uni-siegen.de>
    Alexander Schnurr <schnurr(at)mathematik.uni-siegen.de>
    Susanne Spies <spies(at)mathematik.uni-siegen.de>

  nach einer Idee von
  
    Riko Kelter <riko.kelter(at)uni-siegen.de>
    Alexander Schnurr <schnurr(at)mathematik.uni-siegen.de>
    Susanne Spies <spies(at)mathematik.uni-siegen.de>

  an der Universit&#228;t Siegen.

  Dieses Werk ist lizensiert unter einer Creative Commons
  Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International
  Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie
  http://creativecommons.org/licenses/by-sa/4.0/.

  SPDX-License-Identifier: CC-BY-SA-4.0

  Technische Informationen:

  Diese Aufgabe bindet das Skript &#8216;stackselbstlern.js&#8217; von Michael
  Kallweit f&#252;r die Aufgabennavigation ein.
*/

stackfltfmt:"~f";

/* ---- Zahlen aus der Aufgabenstellung: ----*/
N: 32; /*Anzahl Karten*/
n: 16; /*Stichprobenumfang*/
r: 16; /*Anzahl der roten Karten im Ausgabestapel*/
k: rand_with_step(6,10,1); /*Anzahl der roten Karten im ersten Stapel*/
Liste: [[1, false, "Die Binomialverteilung"], [2, false, "Die Poisson-Verteilung"], [3, false, "Die Gamma-Verteilung"], [4, true, "Die hypergeometrische Verteilung"] ];

test01(x):= is(x>=0) and (x<=1);
/*L&#246;sung der Aufgabe*/
n1: 16
n2: 32
n3: (binomial(r,k)*binomial(N-r,n-k))/binomial(N,n);
TA1: significantfigures(n3, 3);
/*TA1: (binomial(r,k)*binomial(N-r,n-k))/binomial(N,n); */
/*TA1: (binomial(n,n/2)*binomial(n,n/2))/binomial(N,n); Korrekte Antwort f&#252;r Hauptaufgabe*/
zwischen0cListe: [[1, false, "Hierbei bezeichnet \\(Z\\) die Zufallsvariable, die die Anzahl der Erfolge (z.B. Ziehen von \\(r=1\\) schwarzen Karte) beim Ziehen von \\(N\\) Karten aus einer Grundgesamtheit von \\(R\\) Karten z&#228;hlt, von denen \\(k\\) Karten die gew&#252;nschte Eigenschaft
            haben (z.B. schwarze Karte). Die Anzahl aller Karten im Stapel entspricht in der obigen Wahrscheinlichkeitsfunktion \\(N\\)."], [2, true, "Hierbei bezeichnet \\(Z\\) die Zufallsvariable, die die Anzahl der Erfolge (z.B. Ziehen von \\(k=1\\) schwarzen Karte) beim Ziehen von \\(n\\) Karten aus einer Grundgesamtheit von \\(N\\) Karten z&#228;hlt, von denen \\(R\\) Karten die gew&#252;nschte Eigenschaft
            haben (z.B. schwarze Karte). Die Anzahl aller Karten im Stapel entspricht in der obigen Wahrscheinlichkeitsfunktion \\(N\\)."], [3, false, "Hierbei bezeichnet \\(Z\\) die Zufallsvariable, die die Anzahl der Erfolge (z.B. Ziehen von \\(k=1\\) schwarzen Karte) beim Ziehen von \\(r\\) Karten aus einer Grundgesamtheit von \\(N\\) Karten z&#228;hlt, von denen \\(R\\) Karten nicht die gew&#252;nschte Eigenschaft
            haben (z.B. schwarze Karte). Die Anzahl aller Karten im Stapel entspricht in der obigen Wahrscheinlichkeitsfunktion \\(N\\)."], [4, false, "Hierbei bezeichnet \\(R\\) die Zufallsvariable, die die Anzahl der Erfolge (z.B. Ziehen von \\(n=1\\) schwarzen Karte) beim Ziehen von \\(k\\) Karten aus einer Grundgesamtheit von \\(N\\) Karten z&#228;hlt, von denen \\(k\\) Karten die gew&#252;nschte Eigenschaft
            haben (z.B. schwarze Karte). Die Anzahl aller Karten im Stapel entspricht in der obigen Wahrscheinlichkeitsfunktion \\(N\\)."] ];]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>Ein gut durchgemischter Stapel mit je 16 roten und schwarzen Karten wird in zwei gleichgro&#223;e Stapel geteilt. Wie gro&#223; ist die Wahrscheinlichkeit, dass sich in den kleineren Stapeln wiederum je gleichviele rote und schwarze Karten befinden?</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text/>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text/>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text/>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans0b</name>
            <type>radio</type>
            <tans>Liste</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>TA1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>n</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>r</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ansZwischen0c</name>
            <type>radio</type>
            <tans>zwischen0cListe</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <prt>
            <name>ansZwischen0c</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ansZwischen0c</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ansZwischen0c-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <br/>
    <p>
        Klicken Sie zum Fortfahren bitte auf den folgenden Button:
        <button class="buttonweiter" onclick="hide('zwischen0c');show('zwischen1');" type="button">
            Weiter
        </button>
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <br/>
    <p>
        To continue, click on the following button:
        <button class="buttonweiter" onclick="hide('zwischen0c');show('zwischen1');" type="button">
            Next
        </button>
        <br/>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ansZwischen0c-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
&#220;berlegen Sie noch einmal, welche Aussage in Bezug auf den Aufgabenkontext Sinn macht.
    <em>
        Tipp:
    </em>
    Die Anzahl aller Karten im Stapel entspricht in der obigen Wahrscheinlichkeitsfunktion \(N\).
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
Think again what makes sense in terms of the task context.
    <em>
        Tip:
    </em>
    The number of cards in the stack corresponds to the above probability function \(N\).
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt0b</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans0b</sans>
                <tans>4</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt0b-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!




Die hypergeometrische Verteilung eignet sich, da wir eine diskrete Wahrscheinlichkeitsverteilung ben&#246;tigen, welche die Situation modelliert bei der aus einer diskreten Grundgesamtheit eine Anzahl an Elementen ohne Zur&#252;cklegen entnommen
werden und wir nur zwischen zwei F&#228;llen unterscheiden (rot oder schwarz). Die hypergeometrische Verteilung gibt dann Auskunft dar&#252;ber, mit welcher
    <a href="https://de.wikipedia.org/wiki/Wahrscheinlichkeit" title="Wahrscheinlichkeit">
        Wahrscheinlichkeit
    </a>
    in der Stichprobe eine bestimmte Anzahl von Elementen vorkommt,
die die gew&#252;nschte Eigenschaft (rot) haben.
    <br/>
    <p>
        Klicken Sie zum Fortfahren bitte auf den folgenden Button:
        <button class="buttonweiter" onclick="hide('zwischen0b');show('zwischen0c');" type="button">
            Weiter
        </button>
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!




The hypergeometric distribution is suitable since we need a discrete probability distribution that models the situation in which a number of elements are taken from a discrete population without being reset.
and we only distinguish between two cases (red or black). The hypergeometric distribution then provides information about which
    <a href="https://de.wikipedia.org/wiki/Wahrscheinlichkeit" title="Wahrscheinlichkeit">
        Probability
    </a>
    in the sample a certain number of elements occur,
which have the desired property (red).
    <br/>
    <p>
        To continue, click on the following button:
        <button class="buttonweiter" onclick="hide('zwischen0b');show('zwischen0c');" type="button">
            Next
        </button>
        <br/>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt0b-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans0b</sans>
                <tans>1</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt0b-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
&#220;berlegen Sie noch einmal, welche Verteilung hier sinnvoll ist. Es handelt sich bei der Situation offensichtlich
    <i>
        nicht
    </i>
    um eine Bernoullikette. Die Binomialverteilung ist damit ungeeignet.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
Think again what distribution is useful here. The situation is obviously
    <i>
        not
    </i>
    a Bernoulli chain. The binomial distribution is therefore unsuitable.
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt0b-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans0b</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt0b-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
&#220;berlegen Sie noch einmal, welche Verteilung hier sinnvoll ist. Die Poisson-Verteilung ist zwar auf einem diskreten Wahrscheinlichkeitsraum definiert, allerdings m&#252;ssen wir hier die Situation modellieren bei der aus einer diskreten Grundgesamtheit \(k\) mal ohne Zur&#252;cklegen Elemente entnommen werden. Die Poisson-Verteilung eignet sich daf&#252;r offensichtlich nicht.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
Think again what distribution is useful here. The Poisson distribution is defined in a discreet probability space, but here we need to model the situation in the case of a discrete population \(k\) can be removed without resetting elements. The Poisson distribution is obviously not suitable for this.
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt0b-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans0b</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prt0b-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
&#220;berlegen Sie noch einmal, welche Verteilung hier sinnvoll ist. Die Gamma-Verteilung ist nicht auf einem diskreten Wahrscheinlichkeitsraum definiert, und somit hier ungeeignet.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
Think again what distribution is useful here. The gamma distribution is not defined in a discrete probability space and is therefore unsuitable here.
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prt0b-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans0b</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt0b-5-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p dir="ltr" style="text-align: left;">
        Bitte w&#228;hlen Sie eine L&#246;sung aus!
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <p dir="ltr" style="text-align: left;">
        Please choose a solution!
        <br/>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt0b-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>float(ans1)</sans>
                <tans>TA1</tans>
                <testoptions>0.001</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!




Die Wahrscheinlichkeit
    <span>
        , dass sich in den kleineren Stapeln wiederum je gleichviele rote und schwarze Karten befinden
    </span>
    ergibt sich somit zu \({@TA1@}\). Ein M&#246;glichkeit, dies zu berechnen, besteht darin die hypergeometrische Verteilung zu nutzen um die Situation zu modellieren. Es ist  \(P(Z=k):=H_{n,N,R}({k})=\frac{{R\choose k} {N-R\choose n-k}}{{N\choose n}}\), wobei \(N = 32\), \(n=16\), \(r=16\) und \(k={@k@}\).
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!




Probability
    <span>
        that in the smaller stacks there are again equal numbers of red and black cards
    </span>
    thus results \({@TA1@}\). One way to calculate this is to use the hypergeometric distribution to model the situation. It is \(P(Z=k):=H_{n,N,R}({k})=\frac{{R\choose k} {N-R\choose n-k}}{{N\choose n}}\)where \(N = 32\), \(n=16\), \(r=16\) and \(k={@k@}\).
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <div id="prtA-1">
        <span style="font-size: 1.5 em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.




Dass Sie die Aufgabe nicht auf Anhieb l&#246;sen konnten, ist kein Problem! Im Folgenden wird die Aufgabe in kleinere Teilaufgaben unterteilt.
        <p>
            Klicken Sie zum Fortfahren bitte auf den folgenden Button:
            <button class="buttonweiter" onclick="hide('aufgabe');show('zwischen0b');" type="button">
                Weiter
            </button>
            <br/>
        </p>
    </div>
    <div id="prtA-2" style="display:none;">
        <span style="font-size: 1.5 em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
        <p>
            Versuchen Sie nun erneut die Aufgabe zu l&#246;sen. Berechnen Sie daher nun auf Basis der vorherigen Schritte \(P(Z=k):=H_{n,N,R}({k})=\frac{{R\choose k} {N-R\choose n-k}}{{N\choose n}}\) f&#252;r das gesuchte Ereignis \(\{Z={@k@}\}\). Die vorherigen Schritte haben die Ergebnisse \(N={@N@}\), \(n={@n@}\) und \(R={@r@}\) geliefert, sowie
\(k={@k@}\).
        </p>
    </div>
</span>
<span class="multilang" lang="en">
    <div id="prtA-1">
        <span style="font-size: 1.5 em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.




That you couldn't solve the task right away is no problem! In the following, the task is subdivided into smaller subtasks.
        <p>
            To continue, click on the following button:
            <button class="buttonweiter" onclick="hide('aufgabe');show('zwischen0b');" type="button">
                Next
            </button>
            <br/>
        </p>
    </div>
    <div id="prtA-2" style="display:none;">
        <span style="font-size: 1.5 em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
        <p>
            Try to solve the task again. Calculate on the basis of the previous steps \(P(Z=k):=H_{n,N,R}({k})=\frac{{R\choose k} {N-R\choose n-k}}{{N\choose n}}\) for the searched event \(\{Z={@k@}\}\). The previous steps have the results \(N={@N@}\), \(n={@n@}\) and \(R={@r@}\) supplied as well as
\(k={@k@}\).
        </p>
    </div>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>test01(ans1)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>0</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die angegebene Wahrscheinlichkeit liegt nicht zwischen 0 und 1. Bitte korrigieren Sie Ihre L&#246;sung.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The specified probability is not between 0 and 1. Please correct your solution.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>n</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!




Bitte klicken Sie zum fortfahren auf den Button:
    <button class="buttonweiter" onclick="hide('zwischen1');show('zwischen2');" type="button">
        Weiter
    </button>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!




Please click on the button to proceed:
    <button class="buttonweiter" onclick="hide('zwischen1');show('zwischen2');" type="button">
        Next
    </button>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort. 
&#220;berlegen Sie noch einmal, wieviele Karten \(n\) im Stichprobenumfang sind.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
Consider how many cards \(n\) are sampled.
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>r</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Sie k&#246;nnen nun die urspr&#252;ngliche Aufgabe erneut versuchen. Klicken Sie zum Fortfahren bitte auf den folgenden Button:
        <button class="buttonweiter" onclick="hide('zwischen2');show('aufgabe');hide('prtA-1');show('prtA-2');" type="button">
            Weiter
        </button>
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5 em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        You can now try the original task again. To continue, click on the following button:
        <button class="buttonweiter" onclick="hide('zwischen2');show('aufgabe');hide('prtA-1');show('prtA-2');" type="button">
            Next
        </button>
        <br/>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
Lesen Sie noch einmal die Aufgabenstellung.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
Read the task again.
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>