<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lebensdauer einer Fahrradlampe</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <style>
        button.next, button.top-next {
    border: 1px solid black;
background-color: #A9F5BC;
}
button.next:hover, button.top-next:hover {
        background-color: #E0F8E6;
}
button.prev, button.top {
    border: 1px solid black;
    background-color: #DCB25E;
}
button.prev:hover, button.top:hover {
        background-color: #E7D2A7;
}




/* Style f&#252;r Details */
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <details>
        <summary>
            Inhalt und Umfang dieser Aufgabe
        </summary>
        <p>
            In dieser Aufgabe lernen Sie, Wahrscheinlichkeiten f&#252;r unabh&#228;ngige und identisch verteilte Zufallsvariablen zu berechnen. Sie besteht aus den folgenden zwei Teilaufgaben.
        </p>
        <ul>
            <li>
                In
                <b>
                    (a)
                </b>
                Berechnen einer Wahrscheinlichkeit mittels einer Dichtefunktion.
            </li>
            <li>
                In
                <b>
                    (b)
                </b>
                Berechnen einer gemeinsamen Wahrscheinlichkeit f&#252;r unabh&#228;ngige Zufallsvariablen mit identischen Dichtefunktionen.
            </li>
        </ul>
    </details>
    <div class="question-view">
        <p>
            Eine LED, die zum Bau einer Fahrradlampe eingesetzt wird, hat eine
exponentiell verteilte Lebensdauer [gemessen in Jahren] mit Parameter
\(\lambda = {@ 1/mean @}\), d.h. die Dichtefunktion ist
\[
        f(x)=\left\{
                \begin{array}{cl}
                    0 &amp; \mbox{ f&#252;r } x &lt; 0 \\
                    {@ 1/mean @} {@ exp(-x/mean) @} &amp; \mbox{ f&#252;r } x \geq 0.
                \end{array}
        \right.
\]
        </p>
        <div class="question-view">
            <div class="question-view">
                <p>
                    <b>
                        (a)
                    </b>
                    Berechnen Sie die Wahrscheinlichkeit, dass ein zuf&#228;llig
ausgew&#228;hlte LED l&#228;nger als {# minyears #} Jahre funktioniert. Bitte
geben Sie das Ergebnis auf mindestens zwei Nachkommastellen genau
                an.
                </p>
                <!-- Numerical input -->
                <p>
                    Antwort: [[input:ans1]][[validation:ans1]]
                </p>
                [[feedback:ans1]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a.1)
                    </b>
                    Es ist die Wahrscheinlichkeit gesucht, dass eine zuf&#228;llig
ausgew&#228;hlte LED l&#228;nger als {# minyears #} Jahre funktioniert. Welche
der folgenden Ausdr&#252;cke hat den gleichen Wert wie die gesuchte
        Wahrscheinlichkeit?
                </p>
                <!-- Multiple choice input -->
                <p>
                    [[input:ans1a]]
                </p>
                [[validation:ans1a]]
[[feedback:ans1a]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a.2)
                    </b>
                    Bei dichteverteilten Zufallsvariablen k&#246;nnen wir
Wahrscheinlichkeiten mithilfe eines Integrals &#252;ber die Dichte
ausrechnen. Wenn Sie die Wahrscheinlichkeit \(P(X &gt; {@ minyears @})\)
als Integral schreiben, welches Integral ergibt sich?
                </p>
                <!-- Multiple choice input -->
                <p>
                    [[input:ans1b]]
                </p>
                [[validation:ans1b]]
[[feedback:ans1b]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a.3)
                    </b>
                    Im vorherigen Schritt haben Sie das Integral
\[
        \int_{{@ minyears @}}^\infty {@ 1/mean @} {@ exp(-x/mean) @} dx
\]
ausgew&#228;hlt. Was ist eine richtige Stammfunktion f&#252;r den Integranden im
        ausgew&#228;hlten Integral?
                </p>
                <!-- Algebraic input -->
                <p>
                    \(F(x) = \) [[input:ans1c]]
                </p>
                [[validation:ans1c]]
[[feedback:ans1c]]
            </div>
        </div>
        <div class="question-view">
            <div class="question-view">
                <p>
                    <b>
                        (b)
                    </b>
                    Eine Fahrradlampe besteht aus vier LEDs des oben
beschriebenen Typs. Wir bezeichnen ihre Lebensdauern mit
\(X_1,\ldots,X_4\) und nehmen an, dass diese unabh&#228;ngig sind.
                </p>
                <p>
                    Berechnen Sie die Wahrscheinlichkeit, dass nach {# minyears #}
Jahren noch alle vier LEDs funktionieren. Bitte geben Sie das Ergebnis
auf mindestens zwei Nachkommastellen genau an.
                </p>
                <p>
                    Antwort: [[input:ans2]][[validation:ans2]]
                </p>
                [[feedback:ans2]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (b.1)
                    </b>
                    Wir wollen die Wahrscheinlichkeit berechnen, dass nach
{#minyears#} Jahren noch alle vier LEDs funktionieren.  Welches
Ereignis passt zu unserer gesuchten Fragestellung?
                </p>
                <!-- Multiple choice input -->
                <p>
                    [[input:ans2a]]
                </p>
                [[validation:ans2a]]
                [[feedback:ans2a]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (b.2)
                    </b>
                    Jetzt wollen wir die Wahrscheinlichkeit des Ereignisses
\(\{X_1 &gt; {@minyears@}, \dots, X_4 &gt; {@minyears@}\}\) bestimmen.
Welche Gleichungen stimmen?
                </p>
                <!-- Multiple choice input -->
                <p>
                    [[input:ans2b]]
                </p>
                [[validation:ans2b]]
                [[feedback:ans2b]]
            </div>
        </div>
    </div>
    <script>
        /*
    * This program is free software: you can redistribute it and/or
    * modify it under the terms of the GNU General Public License as
    * published by the Free Software Foundation, either version 3 of the
    * License, or (at your option) any later version.
    *
    * This program is distributed in the hope that it will be useful, but
    * WITHOUT ANY WARRANTY; without even the implied warranty of
    * MERCHANTABILITY or FITNESS FOR A PARTICULAR PURPOSE.  See the GNU
    * General Public License for more details.
    *
    * You should have received a copy of the GNU General Public License
    * along with this program.  If not, see <https://www.gnu.org/licenses/>.
    *
    * Copyright (C) 2021 by Daniel Mei&#223;ner <daniel.meissner-i4k@ruhr-uni-bochum.de>
    * SPDX-License-Identifier: GPL-3.0-or-later
    *
    */




/* @license magnet:?xt=urn:btih:1f739d935676111cfff4b4693e3816e664797050&dn=gpl-3.0.txt GPL-v3-or-Later */




(function() {
/* Make sure to create QController instances only once even if the
            script is loaded multiple times.  This is so that the script can
            be included in more than one STACK question on the page and still
            function properly. */
if (window.QController === undefined) {
        window.QController = QController;
        document.addEventListener("DOMContentLoaded", event => {
                const responseform = document.getElementById("responseform");
                if (!responseform) {
                        throw Error("Element with id responseform not found.");
                }

    
    
    
    responseform.querySelectorAll(".stack.que").forEach(stackQue => {
                        const rootElement = stackQue.querySelector(".question-view");

      
      
      
      if (!rootElement) {
                                console.warn("QController: no elements with class question-view inside " + stackQue.id);
                                return;
                        }

      
      
      
      const controller = new QController(rootElement, stackQue.id);
                        controller.installHandlers();
                        controller.loadState();
                        controller.update();
                });
        });
}




/* This class handles the navigation of a single STACK question */
function QController(rootElement, questionId) {
        this.rootElement = rootElement;
        this.id = questionId;
        this.views = rootElement.querySelectorAll(".question-view");
        this.visited = (new Array(this.views.length)).fill(false);
        this.currentView = Array.prototype.findIndex.call(this.views, isLeaf);
        this.visited[this.currentView] = true;
}




QController.prototype.handleVisitedNodes = function(current) {
        var visitedSiblings = [];
        var el = current;
        while (el.nextElementSibling) {
                var index =
                                Array.prototype.indexOf.call(this.views, el.nextElementSibling);
                visitedSiblings.push(this.visited[index])
                el = el.nextElementSibling;
        }

  
  
  
  var withoutHelp = visitedSiblings.every(x => x == false);
        var helpCompleted = visitedSiblings.every(x => x == true);

  
  
  
  current.querySelectorAll(".visited").forEach(function(element) {
                var index = element.getAttribute("data-qcontroller-help");
                if (!index || index === "none")
                        element.hidden = !withoutHelp;
                else if (index === "complete")
                        element.hidden = !helpCompleted;
                else if (index === "partial")
                        element.hidden = withoutHelp || helpCompleted;
                else
                        element.hidden = !visitedSiblings[index - 1];
        });
};




/* Update view according to current controller state */
QController.prototype.update = function() {
        this.views.forEach(hide);

  
  
  
  const current = this.views[this.currentView];
        if (isLeaf(current)) {
                showView(current);
                this.handleVisitedNodes(current);
        } else {
                this.showFirstView(current);
        }
};




/* Install event handlers for all buttons under root element */
QController.prototype.installHandlers = function() {
        const controller = this;
        this.rootElement.querySelectorAll("button").forEach(button => {
                button.addEventListener("click", ev => {
                        const view = button.closest(".question-view");

      
      
      
      if (!view) {
                                console.log("qcontroller: <button> not inside question view");
                                return;
                        }

      
      
      
      const classList = button.classList;
                        let forward = null;

      
      
      
      if (classList.contains("next")) {
                                forward = view.nextElementSibling;
                        } else if (classList.contains("top-next")) {
                                const parent = view.parentNode;
                                if (parent) {
                                        forward = parent.nextElementSibling;
                                } else {
                                        forward = null;
                                }
                        } else if (classList.contains("top")) {
                                forward = containingView(view);
                        } else if (classList.contains("prev")) {
                                forward = view.previousElementSibling;
                        }

      
      
      
      if (forward && forward.classList.contains("question-view")) {
                                controller.visit(forward);
                                controller.saveState();
                                controller.update();
                        }
                });
        });
};




QController.prototype.hasVisited = function(viewIndex) {
        return this.visited[viewIndex];
};




QController.prototype.visit = function(nextView) {
        const index = Array.prototype.findIndex.call(
                this.views,
                x => x === nextView);
        this.currentView = index;
        this.visited[index] = true;
};




/* Save controller state to current session */
QController.prototype.saveState = function() {
        const key = this.id + "-controller-state";
        const state = {
                currentView: this.currentView,
                visited: this.visited
        };
        window.sessionStorage.setItem(key, JSON.stringify(state));
};




/* Load controller state from current session */
QController.prototype.loadState = function() {
        const key = this.id + "-controller-state";
        const state = JSON.parse(window.sessionStorage.getItem(key));

  
  
  
  if (state) {
                this.currentView = state.currentView;
                this.visited = state.visited;
        }
};




/* Show the first leaf question view */
QController.prototype.showFirstView = function(questionView) {
        let view = questionView;

  
  
  
  while (view && !isLeaf(view)) {
                view = view.querySelector(".question-view");
        }

  
  
  
  showView(view);
        this.handleVisitedNodes(view);
}




function showView(view) {
        /* show all containing views */
        while (view) {
                show(view);
                view = containingView(view);
        }
}




/* Return the view containing VIEW */
function containingView(view) {
        let parent = view.parentNode;

  
  
  
  if (parent) {
                parent = parent.closest(".question-view");
        }

  
  
  
  return parent;
}




function show(view) {
        view.style.display = "block";
}




function isLeaf(view) {
        return view.querySelector(".question-view") == null;
}




function isVisible(view) {
        return view.style.display != "none";
}




function hide(view) {
        view.style.display = "none";
}




})();
/* @license-end */
    </script>
</span>
<span class="multilang" lang="en">
    <style>
        button.next, button.top-next {
    border: 1px solid black;
background-color: #A9F5BC;
}
button.next:hover, button.top-next:hover {
        background-color: #E0F8E6;
}
button.prev, button.top {
    border: 1px solid black;
    background-color: #DCB25E;
}
button.prev:hover, button.top:hover {
        background-color: #E7D2A7;
}




/* Style f&#252;r Details */
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <details>
        <summary>
            Content and scope of this task
        </summary>
        <p>
            In this task, you will learn to calculate probabilities for independent and identically distributed random variables. It consists of the following two subtasks.
        </p>
        <ul>
            <li>
                In
                <b>
                    (a)
                </b>
                calculating a probability by means of a density function.
            </li>
            <li>
                In
                <b>
                    (b)
                </b>
                calculating a common probability for independent random variables with identical density functions.
            </li>
        </ul>
    </details>
    <div class="question-view">
        <p>
            An LED used to build a bicycle lamp has a
exponentially distributed life [measured in years] with parameters
\(\lambda = {@ 1/mean @}\), i.e. the density function is
\[
        f(x)=\left\{
                \begin{array}{cl}
                    0 &amp; \mbox{ f&#252;r } x &lt; 0 \\
                    {@ 1/mean @} {@ exp(-x/mean) @} &amp; \mbox{ f&#252;r } x \geq 0.
                \end{array}
        \right.
\]
        </p>
        <div class="question-view">
            <div class="question-view">
                <p>
                    <b>
                        (a)
                    </b>
                    Calculate the probability that a random
selected LED longer than {# minyears #} Years worked. Please
give the result to at least two descendants exactly
on.
                </p>
                <!-- Numerical input -->
                <p>
                    Answer: [[input:ans1]][[validation:ans1]]
                </p>
                [[feedback:ans1]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a)
                    </b>
                    It is the probability that a
selected LED longer than {# minyears #} Years worked. Which
the following expressions have the same value as the searched
Probability?
                </p>
                <!-- Multiple choice input -->
                <p>
                    [[input:ans1a]]
                </p>
                [[validation:ans1a]][[feedback:ans1a]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a.2)
                    </b>
                    With densely distributed random variables we can
Probabilities with an integrator about density
Calculate. If you have the chance \(P(X &gt; {@ minyears @})\)as an integral, which integral results?
                </p>
                <!-- Multiple choice input -->
                <p>
                    [[input:ans1b]]
                </p>
                [[validation:ans1b]][[feedback:ans1b]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a.3)
                    </b>
                    In the previous step you have the integral
\[
        \int_{{@ minyears @}}^\infty {@ 1/mean @} {@ exp(-x/mean) @} dx
\]selected. What is a proper root function for the integrande in
selected integral?
                </p>
                <!-- Algebraic input -->
                <p>
                    \(F(x) = \) [[input:ans1c]]
                </p>
                [[validation:ans1c]][[feedback:ans1c]]
            </div>
        </div>
        <div class="question-view">
            <div class="question-view">
                <p>
                    <b>
                        (b)
                    </b>
                    A bicycle lamp consists of four LEDs of the above
type described. We call their lifespans with
\(X_1,\ldots,X_4\) and assume that they are independent.
                </p>
                <p>
                    Calculate the probability that after {# minyears #}All four LEDs still work for years. Please give the result
to at least two descendants.
                </p>
                <p>
                    Answer: [[input:ans2]][[validation:ans2]]
                </p>
                [[feedback:ans2]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (b.1)
                    </b>
                    We want to calculate the probability that after
{#minyears#} All four LEDs still work for years. What
Event fits our searched question?
                </p>
                <!-- Multiple choice input -->
                <p>
                    [[input:ans2a]]
                </p>
                [[validation:ans2a]][[feedback:ans2a]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (b.2)
                    </b>
                    Now we want the probability of the event
\(\{X_1 &gt; {@minyears@}, \dots, X_4 &gt; {@minyears@}\}\) to determine.
What are the equations?
                </p>
                <!-- Multiple choice input -->
                <p>
                    [[input:ans2b]]
                </p>
                [[validation:ans2b]][[feedback:ans2b]]
            </div>
        </div>
    </div>
    <script>
        /*
    * This program is free software: you can redistribute it and/or
    * modify it under the terms of the GNU General Public License as
    * published by the Free Software Foundation, either version 3 of the
    * License, or (at your option) any later version.
    *
    * This program is distributed in the hope that it will be useful, but
    * WITHOUT ANY WARRANTY; without even the implied warranty of
    * MERCHANTABILITY or FITNESS FOR A PARTICULAR PURPOSE.  See the GNU
    * General Public License for more details.
    *
    * You should have received a copy of the GNU General Public License
    * along with this program.  If not, see <https://www.gnu.org/licenses/>.
*
* Copyright (C) 2021 by Daniel Mei&#223;ner <daniel.meissner-i4k@ruhr-uni-bochum.de>* SPDX-License Identifier: GPL-3.0-or-later
*
*/




/* @license magnet:?xt=urn:btih:1f739d935676111cfff4b4693e3816e664797050&dn=gpl-3.0.txt GPL-v3-orLater */




(function() {
/* Make sure to create QController instances only once even if the
            script is loaded multiple times.  This is so that the script can
            be included in more than one STACK question on the page and still
            function properly. */
if (window.QController === undefined) {
        window.QController = QController;
        document.addEventListener("DOMContentLoaded", event => {
                const responseform = document.getElementById("responseform");
                if (!responseform) {
                        throw Error("Element with id responseform not found.");
                }

    
    
    
    responseform.querySelectorAll(".stack.que").forEach(stackQue => {
                        const rootElement = stackQue.querySelector(".question-view");

      
      
      
      if (!rootElement) {
                                console.warn("QController: no elements with class question-view inside " + stackQue.id);
                                return;
                        }

      
      
      
      const controller = new QController(rootElement, stackQue.id);
                        controller.installHandlers();
                        controller.loadState();
                        controller.update();
                });
        });
}




/* This class handles the navigation of a single STACK question */
function QController(rootElement, questionId) {
        this.rootElement = rootElement;
        this.id = questionId;
        this.views = rootElement.querySelectorAll(".question-view");
        this.visited = (new Array(this.views.length)).fill(false);
        this.currentView = Array.prototype.findIndex.call(this.views, isLeaf);
        this.visited[this.currentView] = true;
}




QController.prototype.handleVisitedNodes = function(current) {
        var visitedSiblings = [];
        var el = current;
        while (el.nextElementSibling) {
                var index =
                                Array.prototype.indexOf.call(this.views, el.nextElementSibling);
                visitedSiblings.push(this.visited[index])
                el = el.nextElementSibling;
        }

  
  
  
  var withoutHelp = visitedSiblings.every(x => x == false);
        var helpCompleted = visitedSiblings.every(x => x == true);

  
  
  
  current.querySelectorAll(".visited").forEach(function(element) {
                var index = element.getAttribute("data-qcontroller-help");
                if (!index || index === "none")
                        element.hidden = !withoutHelp;
                else if (index === "complete")
                        element.hidden = !helpCompleted;
                else if (index === "partial")
                        element.hidden = withoutHelp || helpCompleted;
                else
                        element.hidden = !visitedSiblings[index - 1];
        });
};




/* Update view according to current controller state */
QController.prototype.update = function() {
        this.views.forEach(hide);

  
  
  
  const current = this.views[this.currentView];
        if (isLeaf(current)) {
                showView(current);
                this.handleVisitedNodes(current);
        } else {
                this.showFirstView(current);
        }
};




/* Install event handlers for all buttons under root element */
QController.prototype.installHandlers = function() {
        const controller = this;
        this.rootElement.querySelectorAll("button").forEach(button => {
                button.addEventListener("click", ev => {
                        const view = button.closest(".question-view");

      
      
      
      if (!view) {
                                console.log("qcontroller: <button> not inside question view");
                                return;
                        }

      
      
      
      const classList = button.classList;
                        let forward = null;

      
      
      
      if (classList.contains("next")) {
                                forward = view.nextElementSibling;
                        } else if (classList.contains("top-next")) {
                                const parent = view.parentNode;
                                if (parent) {
                                        forward = parent.nextElementSibling;
                                } else {
                                        forward = null;
                                }
                        } else if (classList.contains("top")) {
                                forward = containingView(view);
                        } else if (classList.contains("prev")) {
                                forward = view.previousElementSibling;
                        }

      
      
      
      if (forward && forward.classList.contains("question-view")) {
                                controller.visit(forward);
                                controller.saveState();
                                controller.update();
                        }
                });
        });
};




QController.prototype.hasVisited = function(viewIndex) {
        return this.visited[viewIndex];
};




QController.prototype.visit = function(nextView) {
        const index = Array.prototype.findIndex.call(
                this.views,
                x => x === nextView);
        this.currentView = index;
        this.visited[index] = true;
};




/* Save controller state to current session */
QController.prototype.saveState = function() {
        const key = this.id + "-controller-state";
        const state = {
                currentView: this.currentView,
                visited: this.visited
        };
        window.sessionStorage.setItem(key, JSON.stringify(state));
};




/* Load controller state from current session */
QController.prototype.loadState = function() {
        const key = this.id + "-controller-state";
        const state = JSON.parse(window.sessionStorage.getItem(key));

  
  
  
  if (state) {
                this.currentView = state.currentView;
                this.visited = state.visited;
        }
};




/* Show the first leaf question view */
QController.prototype.showFirstView = function(questionView) {
        let view = questionView;

  
  
  
  while (view && !isLeaf(view)) {
                view = view.querySelector(".question-view");
        }

  
  
  
  showView(view);
        this.handleVisitedNodes(view);
}




function showView(view) {
        /* show all containing views */
        while (view) {
                show(view);
                view = containingView(view);
        }
}




/* Return the view containing VIEW */
function containingView(view) {
        let parent = view.parentNode;

  
  
  
  if (parent) {
                parent = parent.closest(".question-view");
        }

  
  
  
  return parent;
}




function show(view) {
        view.style.display = "block";
}




function isLeaf(view) {
        return view.querySelector(".question-view") == null;
}




function isVisible(view) {
        return view.style.display != "none";
}




function hide(view) {
        view.style.display = "none";
}




}),
/* @license-end */
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>10.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
  Lebensdauer einer Fahrradlampe

  wurde entwickelt von

    Daniel Mei&#223;ner <daniel.meissner-i4k(at)ruhr-uni-bochum.de>

  nach einer Idee von

    Herold Dehling <herold.dehling(at)ruhr-uni-bochum.de>

  an der Ruhr-Universit&#228;t Bochum.

  Dieses Werk ist lizenziert unter einer Creative Commons
  Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International
  Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie
  http://creativecommons.org/licenses/by-sa/4.0/.

  SPDX-License-Identifier: CC-BY-SA-4.0

  Technische Informationen:

  Diese Aufgabe ist eine mehrteilige Aufgabe, deren einzelne
  Aufgabenteile mithilfe von JavaScript nacheinander sichtbar gemacht
  werden.  Manchen Aufgabenteilen sind zus&#228;tzlich Hilfeschritte
  beigef&#252;gt, die Studierende schrittweise durch den Aufgabenteil
  f&#252;hren.
*/

probabilityp (x) := is(x >= 0 and x <= 1);

mean: rand([10, 12, 8]);
minyears: rand([1/8, 1/4, 1/2])*mean;

pdf: 1/mean * exp(-x/mean);

Tans1: fullratsimp(int(pdf, x, minyears, inf));
Tans2: Tans1^4;

Tans1a: random_permutation([
[P(X > minyears), true],
[P(X >= minyears), true],
[P(X = minyears), false],
[P(X < minyears), false],
[P(X <= minyears), false]
]);

integrals: map (lambda ([x], [x[1], x[2], simplode (x[3])]), [
  [1, true, ["\\(\\displaystyle\\int_{", minyears, "}^\\infty", stack_disp(1/mean, ""), stack_disp(exp(-x/mean),""), "dx\\)"]],
  [2, false, ["\\(\\displaystyle\\int_{-\\infty}^{", minyears, "}", stack_disp(1/mean, ""), stack_disp(exp(-x/mean),""), "dx\\)"]],
  [3, false, ["\\(\\displaystyle\\int_{", minyears, "}^\\infty", stack_disp(1/mean, ""), stack_disp(x,""), stack_disp(exp(-x/mean),""), "dx\\)"]],
  [4, false, ["\\(\\displaystyle\\int_",0,"^{", minyears, "}", stack_disp(1/mean, ""), stack_disp(exp(-x/mean),"") ,"dx\\)"]]
]);

Tans1b: random_permutation(integrals);

Tans2a: random_permutation([
[1, true, sconcat("\\(\\{", stack_disp(X_1 > minyears, ""), ", \\dots, ", stack_disp(X_4 > minyears, ""), "\\}\\)")],
[2, false, sconcat("\\(\\{", stack_disp(X_1 = minyears, ""), ", \\dots, ", stack_disp(X_4 = minyears, ""), "\\}\\)")],
[3, false, sconcat("\\(\\{", stack_disp(X_1 < minyears, ""), ", \\dots, ", stack_disp(X_4 < minyears, ""), "\\}\\)")]
]);

Tans2b: random_permutation([
  [1, true, sconcat(
    "\\(P(", stack_disp(X_1 > minyears, ""), ", \\dots, ", stack_disp(X_4 > minyears, ""),
    ") = P(",
    stack_disp(X_1 > minyears, ""),
    ") \\cdots P(",
    stack_disp(X_4 > minyears, ""),
    ")\\)")],
  [2, false, sconcat(
    "\\(P(", stack_disp(X_1 > minyears, ""), ", \\dots, ", stack_disp(X_4 > minyears, ""),
    ") = P(",
    stack_disp(X_1 > minyears, ""),
    ") + \\cdots + P(",
    stack_disp(X_4 > minyears, ""),
    ")\\)")],
  [3, true, sconcat(
    "\\(P(", stack_disp(X_1 > minyears, ""), ", \\dots, ", stack_disp(X_4 > minyears, ""),
    ") = P(",
    stack_disp(X_1 > minyears, ""),
    ")^4\\)")]
  ]);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{# mean #}, {# minyears #}, {# Tans1 #}, {# Tans2 #}</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed>i</variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>numerical</type>
            <tans>decimalplaces(Tans1,2)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options>allowempty</options>
        </input>
        <input>
            <name>ans1a</name>
            <type>checkbox</type>
            <tans>Tans1a</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ans1b</name>
            <type>radio</type>
            <tans>Tans1b</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ans1c</name>
            <type>algebraic</type>
            <tans>Tans1c</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>Tans2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2a</name>
            <type>radio</type>
            <tans>Tans2a</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>ans2b</name>
            <type>checkbox</type>
            <tans>Tans2b</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <prt>
            <name>ans1</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>probabilityp(ans1)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>ans1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Wahrscheinlichkeiten werden durch Zahlen im Intervall [0,1] beschrieben. Ihre Eingabe liegt nicht in diesem Intervall.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Probabilities are described by numbers in the interval [0,1]. Your input is not in this interval.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans1</sans>
                <tans>Tans1</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Da \(X\) eine stetige Zufallsvariable ist, k&#246;nnen wir die gesuchte Wahrscheinlichkeit mithilfe von Integration der Dichte oder mithilfe der Verteilungsfunktion bestimmen.  Die Verteilungsfunktionist uns hier nicht gegeben, daher w&#228;hlen wir die Integration der Dichtefunktion:
\begin{align*}
{@ P(X &gt; minyears) @} &amp;= \int_{@minyears@}^\infty {@ 1/mean @} {@ exp(-x/mean) @} dx \\
&amp;= \lim_{b\to\infty}\left[ {@-exp(-x/mean)@} \right]_{@minyears@}^b\\
&amp;= {@int(pdf, x, minyears, inf),simp:true@} \approx {@decimalplaces(Tans1,2)@}
\end{align*}
    </p>
    <p>
        Klicken Sie auf
        <button class="top-next" type="button">
            Weiter
        </button>
        , um zum n&#228;chsten Aufgabenteil zu gelangen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        There. \(X\) is a constant random variable, we can determine the desired probability by integrating the density or by using the distribution function. The distribution function is not given to us here, so we choose the integration of the density function:
\begin{align*}
{@ P(X &gt; minyears) @} &amp;= \int_{@minyears@}^\infty {@ 1/mean @} {@ exp(-x/mean) @} dx \\
&amp;= \lim_{b\to\infty}\left[ {@-exp(-x/mean)@} \right]_{@minyears@}^b\\
&amp;= {@int(pdf, x, minyears, inf),simp:true@} \approx {@decimalplaces(Tans1,2)@}
\end{align*}
    </p>
    <p>
        Click on
        <button class="top-next" type="button">
            Next
        </button>
        to get to the next task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="visited" data-qcontroller-help="none">
        Dass Sie die Aufgabe nicht auf Anhieb l&#246;sen konnten, ist kein Problem! Im Folgenden wird die Aufgabe in kleinere Schritte unterteilt. Wenn Sie diese Schritte alle gel&#246;st haben, k&#246;nnen Sie die urspr&#252;ngliche Aufgabe sicherlich l&#246;sen. Klicken Sie dazu auf
        <button class="next" type="button">
            Weiter
        </button>
        .
    </p>
    <p>
        <span class="visited" data-qcontroller-help="1">
            Willkommen zur&#252;ck! W&#228;hrend der Bearbeitung der Hilfe haben Sie erfahren, dass die folgenden beiden Ausdr&#252;cke die gesuchte Wahrscheinlichkeit beschreiben: {@P(X &gt; minyears)@} und {@P(X &gt;= minyears)@}.
        </span>
        <span class="visited" data-qcontroller-help="2">
            Sie haben gelernt, dass die gesuchte Wahrscheinlichkeit mithilfe des Integrals {@int(1/mean * exp(-x/mean), x, minyears, inf),simp:false@} berechnet werden kann.
        </span>
        <span class="visited" data-qcontroller-help="3">
            Eine m&#246;gliche Stammfunktion f&#252;r den Integranden lautet: {@-exp(-x/mean)@}
        </span>
    </p>
    <p class="visited" data-qcontroller-help="partial">
        M&#246;chten Sie die Hilfe fortsetzen? Klicken Sie auf
        <button class="next" type="button">
            Hilfe fortsetzen
        </button>
        .
    </p>
</span>
<span class="multilang" lang="en">
    <p class="visited" data-qcontroller-help="none">
        That you couldn't solve the task right away is no problem! In the following, the object is subdivided into smaller steps. If you have all solved these steps, you can certainly solve the original task. Click on
        <button class="next" type="button">
            Next
        </button>
        .
    </p>
    <p>
        <span class="visited" data-qcontroller-help="1">
            Welcome back! During the processing of the help you have learned that the following two expressions describe the desired probability: {@P(X &gt; minyears)@} and {@P(X &gt;= minyears)@}.
        </span>
        <span class="visited" data-qcontroller-help="2">
            You have learned that the searched probability by using the integral {@int(1/mean * exp(-x/mean), x, minyears, inf),simp:false@} can be calculated.
        </span>
        <span class="visited" data-qcontroller-help="3">
            A possible root function for the integrande is: {@-exp(-x/mean)@}
        </span>
    </p>
    <p class="visited" data-qcontroller-help="partial">
        Would you like to continue the help? Click on
        <button class="next" type="button">
            Continue
        </button>
        .
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>ans1a</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text>selection: setify(ans1a);
correct_solutions: setify(mcq_correct(Tans1a));
none_correct: emptyp(intersection(selection, correct_solutions));</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1a</sans>
                <tans>mcq_correct(Tans1a)</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1a-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Die exakte &#220;bertragung des Ausdrucks &#8222;l&#228;nger als {#minyears #} Jahre&#8220; f&#252;hrt auf das Ereignis {@ X &gt; minyears @}. Da in der Aufgabe allerdings nach Ausdr&#252;cken gefragt ist, die den gleichen Wert wie die gesuchte Wahrscheinlichkeit haben, sind die beiden richtigen Antworten {@P(X&gt;=minyears)@} und {@P(X &gt; minyears)@}. Stetige Zufallsvariablen haben n&#228;mlich die Eigenschaft, dass Ereignissen wie {@X = minyears@} die Wahrscheinlichkeit 0 zugeordnet wird.
    </p>
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        , um mit der Hilfe fortzufahren, oder auf
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        , um zur urspr&#252;nglichen Aufgabenstellung zur&#252;ckzukehren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        The exact transmission of the expression &#8220;longer than {#minyears #} Years&#8221; leads to the event {@ X &gt; minyears @}. However, since the task is asked for expressions which have the same value as the desired probability, the two correct answers are: {@P(X&gt;=minyears)@} and {@P(X &gt; minyears)@}. Constant random variables have the property that events such as {@X = minyears@} the probability 0 is assigned.
    </p>
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        to continue with the help, or to
        <button class="top" type="button">
            Back
        </button>
        to return to the original task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>ans1a-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1a</sans>
                <tans><![CDATA[[P(X > minyears)]]]></tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1a-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Ihre Antwort ist teilweise korrekt.
    </p>
    <p>
        Die ausgew&#228;hlte Antwort ist zwar korrekt. Sie haben allerdings nicht 
alle korrekten Antworten ausgew&#228;hlt. Die exakte &#220;bertragung des Ausdrucks &#8222;l&#228;nger als {#minyears #} Jahre&#8220; f&#252;hrt auf das Ereignis 
{@ X &gt; minyears @}. Allerdings ist hier nach Ausdr&#252;cken gefragt, die den 
gleichen Wert haben. Beachten Sie hierbei, dass \(X\) eine Zufallsvariable mit 
stetiger Verteilung ist, um auf die fehlende Antwort zu kommen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Your answer is partly correct.
    </p>
    <p>
        The selected answer is correct. you do not have
all correct answers selected. The exact transmission of the expression &#8220;longer than {#minyears #} Years&#8221; leads to the event
{@ X &gt; minyears @}. However, there is a demand for expressions which
have the same value. Note that \(X\) a random variable with
constant distribution is to come to the missing answer.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>ans1a-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1a</sans>
                <tans><![CDATA[[P(X >= minyears)]]]></tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1a-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em;
color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Ihre Antwort ist
teilweise korrekt.
    </p>
    <p>
        Die von Ihnen ausgew&#228;hlte Wahrscheinlichkeit ist nicht falsch, da \(X\)
eine Zufallsvariable mit stetiger Verteilung ist. Sie haben allerdings
nicht alle korrekten Antworten ausgew&#228;hlt. Achten Sie genau auf die 
Formulierung &#8222;l&#228;nger als {#minyears#} Jahre&#8220;, um auf 
die fehlende Antwort zu kommen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em;
color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Your answer is
partially correct.
    </p>
    <p>
        The probability you selected is not wrong, as \(X\)is a random variable with constant distribution. You have
not selected all correct answers. Pay attention to
wording &#8216;longer than {#minyears#} Years&#8221;
to get the missing answer.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>ans1a-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>none_correct</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1a-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Alle ausgew&#228;hlten Antworten sind falsch.  Bitte versuchen Sie es noch 
einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        All selected answers are wrong. Please try
once.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans1a-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Ihre Antwort ist teilweise korrekt.
    </p>
    <p>
        Sie haben richtige Antworten ausgew&#228;hlt, allerdings haben Sie auch 
falsche Antworten ausgew&#228;hlt.  Bitte versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Your answer is partly correct.
    </p>
    <p>
        You have selected correct answers, but you also have
wrong answers selected. Please try again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>ans1b</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1b</sans>
                <tans>first(mcq_correct(Tans1b))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1b-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        , um mit der Hilfe fortzufahren, oder auf
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        , um zur urspr&#252;nglichen Aufgabenstellung zur&#252;ckzukehren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        to continue with the help, or to
        <button class="top" type="button">
            Back
        </button>
        to return to the original task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>ans1b-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>member(ans1b, [2, 4])</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1b-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Das ausgew&#228;hlte Integral hat die falschen Grenzen. Wir wollen die Wahrscheinlichkeit {@P(X &gt; minyears)@} berechnen. &#220;berlegen Sie noch einmal, wie die Grenzen lauten m&#252;ssen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The selected integral has the wrong limits. We want the probability {@P(X &gt; minyears)@} calculate. Think again how the boundaries must be.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>ans1b-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1b</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1b-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Das ausgew&#228;hlte Integral hat zwar die richtigen Grenzen, allerdings ist der Integrand falsch. Sie m&#252;ssen &#252;ber die Dichtefunktion integrieren. Bitte versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The selected integral has the right limits, but the integrand is wrong. You need to integrate with the density function. Please try again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans1b-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>ans1c</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text>DF: diff(ans1c, x);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>DF</sans>
                <tans>pdf</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans1c-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Sie haben eine korrekte Stammfunktion eingegeben: Die Ableitung der von Ihnen eingegebenen Funktion lautet {@DF@} und stimmt mit der Dichtefunktion im relevanten Intervall &#252;berein.
    </p>
    <p>
        Klicken Sie
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        , um die urspr&#252;ngliche Frage noch einmal zu probieren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        You entered a correct root function: The derivation of the function you have entered is {@DF@} and agrees with the density function in the relevant interval.
    </p>
    <p>
        Click
        <button class="top" type="button">
            Back
        </button>
        to try the original question again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans1c-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Die Ableitung der von Ihnen angegebenen Funktion lautet {@ DF @}.  Dies stimmt nicht mit dem angegebenen Integranden &#252;berein.
Schauen Sie noch einmal in Ihren Notizen oder im Internet zum Thema
        <em>
            Stammfunktionen von Exponentialfunktionen
        </em>
        .  Falls Sie absolut nicht weiterkommen, k&#246;nnen Sie auch hier Hilfe bekommen:
    </p>
    <details>
        <summary>
            <em>
                Hinweis
            </em>
        </summary>
        Eine Stammfunktion f&#252;r \(x \mapsto e^{ax}\) ist gegeben durch \(x \mapsto \frac{e^{ax}}{a}\).
    </details>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The derivation of the function you provide is {@ DF @}. This is not consistent with the specified integrande.
Take a look at your notes or on the Internet
        <em>
            Trial functions of exponential functions
        </em>
        . If you do not get any further, you can also get help here:
    </p>
    <details>
        <summary>
            <em>
                Note
            </em>
        </summary>
        A master function \(x \mapsto e^{ax}\) is given by \(x \mapsto \frac{e^{ax}}{a}\).
    </details>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>ans2</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>probabilityp(ans2)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>ans2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Wahrscheinlichkeiten werden durch Zahlen im Intervall [0,1] beschrieben. Ihre Eingabe liegt nicht in diesem Intervall.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Probabilities are described by numbers in the interval [0,1]. Your input is not in this interval.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans2</sans>
                <tans>Tans2</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Wegen der Unabh&#228;ngigkeit der Zufallsvariablen \(X_1, \dots, X_4\) kann man das gesuchte Ereignis wie folgt beschreiben:
\[
\{X_1 &gt; {@minyears@}, \dots, X_4 &gt; {@minyears@}\}
\]
Dann kann man die Wahrscheinlichkeit dieses Ereignisses so bestimmen:
\begin{align*}
P(X_1 &gt; {@minyears@}, \dots, X_4 &gt; {@minyears@})
&amp;= P(X_1 &gt; {@minyears@}) \cdots P(X_4 &gt; {@minyears@})\\
&amp;= {@Tans1@} \cdots {@Tans1@} = \left({@ Tans1;@}\right)^4 = {@Tans1^4,simp:true@}
\approx {@ decimalplaces(Tans1^4,2) @}.
\end{align*}
    </p>
    <p>
        Herzlichen Gl&#252;ckwunsch! Sie haben das Ende dieser Aufgabe erreicht.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Due to the independence of random variables \(X_1, \dots, X_4\) can be described as follows:
\[
\{X_1 &gt; {@minyears@}, \dots, X_4 &gt; {@minyears@}\}
\]Then you can determine the likelihood of this event:
\begin{align*}
P(X_1 &gt; {@minyears@}, \dots, X_4 &gt; {@minyears@})
&amp;= P(X_1 &gt; {@minyears@}) \cdots P(X_4 &gt; {@minyears@})\\
&amp;= {@Tans1@} \cdots {@Tans1@} = \left({@ Tans1;@}\right)^4 = {@Tans1^4,simp:true@}
\approx {@ decimalplaces(Tans1^4,2) @}.
\end{align*}
    </p>
    <p>
        Congratulations! They have reached the end of this task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="visited" data-qcontroller-help="none">
        Dass Sie die Aufgabe nicht auf Anhieb l&#246;sen konnten, ist kein Problem! Im Folgenden wird die Aufgabe in kleinere Schritte unterteilt. Wenn Sie diese Schritte alle gel&#246;st haben, k&#246;nnen Sie die urspr&#252;ngliche Aufgabe sicherlich l&#246;sen. Klicken Sie dazu auf
        <button class="next" type="button">
            Weiter
        </button>
        .
    </p>
    <p>
        <span class="visited" data-qcontroller-help="1">
            Willkommen zur&#252;ck! W&#228;hrend der Bearbeitung der Hilfe haben Sie erfahren, dass in der Aufgabe nach der Wahrscheinlichkeit des Ereignisses \(\{{@X[1] &gt; minyears@}, \dots, {@X[4]&gt;minyears@}\}\) gefragt ist.
        </span>
        <span class="visited" data-qcontroller-help="2">
            Aufgrund der Unabh&#228;ngigkeit der Zufallsvariablen \(X_1,\dots,X_4\) l&#228;sst sich die gesuchte Wahrscheinlichkeit also als \(({@P(X[1]&gt;minyears)@})^4\) ausdr&#252;cken.
        </span>
    </p>
    <p class="visited" data-qcontroller-help="partial">
        M&#246;chten Sie die Hilfe fortsetzen? Klicken Sie auf
        <button class="next" type="button">
            Hilfe fortsetzen
        </button>
        .
    </p>
</span>
<span class="multilang" lang="en">
    <p class="visited" data-qcontroller-help="none">
        That you couldn't solve the task right away is no problem! In the following, the object is subdivided into smaller steps. If you have all solved these steps, you can certainly solve the original task. Click on
        <button class="next" type="button">
            Next
        </button>
        .
    </p>
    <p>
        <span class="visited" data-qcontroller-help="1">
            Welcome back! During the processing of the help you have learned that in the task after the probability of the event \(\{{@X[1] &gt; minyears@}, \dots, {@X[4]&gt;minyears@}\}\) is in demand.
        </span>
        <span class="visited" data-qcontroller-help="2">
            Due to the independence of random variables \(X_1,\dots,X_4\) so the desired probability can be \(({@P(X[1]&gt;minyears)@})^4\) express.
        </span>
    </p>
    <p class="visited" data-qcontroller-help="partial">
        Would you like to continue the help? Click on
        <button class="next" type="button">
            Continue
        </button>
        .
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>ans2a</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2a</sans>
                <tans>first(mcq_correct(Tans2a))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans2a-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        , um mit der Hilfe fortzufahren, oder auf
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        , um zur urspr&#252;nglichen Aufgabenstellung zur&#252;ckzukehren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        to continue with the help, or to
        <button class="top" type="button">
            Back
        </button>
        to return to the original task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>ans2a-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2a</sans>
                <tans>2</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans2a-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Die Zufallsvariablen \(X_1, \dots X_4\) bezeichnen die Lebensdauern der jeweiligen LED. Das ausgew&#228;hlte Ereignis besagt, dass alle LEDs eine Lebensdauer von {@minyears@} Jahren haben. Bitte versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The random variables \(X_1, \dots X_4\) designates the lifespans of the respective LEDs. The selected event states that all LEDs have a lifetime of {@minyears@} years. Please try again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>ans2a-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2a</sans>
                <tans>3</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans2a-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Die Zufallsvariablen \(X_1, \dots X_4\) bezeichnen die Lebensdauern der jeweiligen LED. Das ausgew&#228;hlte Ereignis besagt, dass alle LEDs eine
        <em>
            k&#252;rzere
        </em>
        Lebensdauer als {@minyears@} Jahre haben. Bitte versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The random variables \(X_1, \dots X_4\) designates the lifespans of the respective LEDs. The selected event states that all LEDs have
        <em>
            shorter
        </em>
        Lifetime as {@minyears@} years. Please try again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans2a-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>ans2b</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text>selection: setify(ans2b);
correct_solutions: setify(mcq_correct(Tans2b));
partially_correct: subsetp(selection, correct_solutions);
none_correct: emptyp(intersection(selection, correct_solutions));</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2b</sans>
                <tans>mcq_correct(Tans2b)</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans2b-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Wegen der Unabh&#228;ngigkeit der Zufallsvariablen \(X_1, \dots X_4\) k&#246;nnen wir die Wahrscheinlichkeit \(P(X_1 &gt; {@minyears@}, \dots X_4 &gt; {@minyears@})\) als ein Produkt der einzelnen Wahrscheinlichkeiten \(P(X_i &gt; {@minyears@})\) ausdr&#252;cken. Da alle Zufallsvariablen \(X_i\) die gleiche Verteilung haben, k&#246;nnen wir die Rechnung noch zu \(P(X_1 &gt; {@minyears@})^4\) vereinfachen.
    </p>
    <p>
        Klicken Sie auf
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        , um zur urspr&#252;nglichen Aufgabenstellung zur&#252;ckzukehren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        Due to the independence of random variables \(X_1, \dots X_4\) we have the chance \(P(X_1 &gt; {@minyears@}, \dots X_4 &gt; {@minyears@})\) as a product of individual probabilities \(P(X_i &gt; {@minyears@})\) express. Since all random variables \(X_i\) the same distribution, we can still have the invoice \(P(X_1 &gt; {@minyears@})^4\) simplify.
    </p>
    <p>
        Click on
        <button class="top" type="button">
            Back
        </button>
        to return to the original task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>ans2b-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>partially_correct</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans2b-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Ihre Antwort ist teilweise korrekt.
    </p>
    <p>
        Die ausgew&#228;hlte Antwort ist zwar korrekt, Sie haben allerdings
nicht alle korrekten Antworten ausgew&#228;hlt. Beachten Sie, dass auch mehrere 
Antworten korrekt sein k&#246;nnen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Your answer is partly correct.
    </p>
    <p>
        The selected answer is correct, but you have
not selected all correct answers. Note that several
Answers can be correct.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>ans2b-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>none_correct</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans2b-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Die ausgew&#228;hlte Gleichung stimmt nicht. M&#246;glicherweise haben Sie sie aufgrund folgender Rechenregel ausgew&#228;hlt:
\[
P(A \cup B) = P(A) + P(B)
\]
Diese Rechenregel gilt allerdings
        <ol>
            <li>
                nur f&#252;r
                <em>
                    disjunkte
                </em>
                Ereignisse, also Ereignisse, die nicht gleichzeitig eintreten k&#246;nnen, und
            </li>
            <li>
                bezieht sich nur auf eine ODER-Verkn&#252;pfung der beiden Ereignisse.
            </li>
        </ol>
        In unserem Fall sind die Ereignisse \(\{X_1 &gt; {@minyears@}\}, \dots \{X_4&gt;{@minyears@}\}\) weder disjunkt, noch ist das uns interessierende Ereignis eine ODER-Verkn&#252;pfung dieser Ereignisse.
    </p>
    <p>
        Bitte versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The selected equation is not true. You may have selected them on the basis of the following calculation rule:
\[
P(A \cup B) = P(A) + P(B)
\]However, this calculation rule applies
        <ol>
            <li>
                only for
                <em>
                    disjunct
                </em>
                events, that is, events that cannot occur at the same time, and
            </li>
            <li>
                refers only to an OR connection of the two events.
            </li>
        </ol>
        In our case, the events \(\{X_1 &gt; {@minyears@}\}, \dots \{X_4&gt;{@minyears@}\}\) neither disjunct nor is the event of interest to us an OR link of these events.
    </p>
    <p>
        Please try again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans2b-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Ihre Antwort ist teilweise korrekt.
    </p>
    <p>
        Sie haben richtige Antworten ausgew&#228;hlt, allerdings haben Sie auch falsche Antworten ausgew&#228;hlt.  Bitte versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:orange;">
            <i class="fa fa-adjust">
            </i>
        </span>
        Your answer is partly correct.
    </p>
    <p>
        You have selected correct answers, but you have also selected incorrect answers. Please try again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1557205003</deployedseed>
        <deployedseed>1303286914</deployedseed>
        <deployedseed>395837351</deployedseed>
        <deployedseed>1901258845</deployedseed>
        <deployedseed>1018059996</deployedseed>
        <deployedseed>1985556290</deployedseed>
        <deployedseed>1139492242</deployedseed>
        <deployedseed>729921243</deployedseed>
        <deployedseed>953537624</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>decimalplaces(Tans1,2)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>Tans2</value>
            </testinput>
            <expected>
                <name>ans1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>ans2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans2-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>decimalplaces(Tans1,0)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>0</value>
            </testinput>
            <expected>
                <name>ans1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>ans2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans2-2-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>decimalplaces(Tans1,0)</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>decimalplaces(Tans1,0)^4</value>
            </testinput>
            <expected>
                <name>ans1</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans1-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>ans2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans2-2-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>