<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Lebensdauer eines Fahrradschlauchs (tutoriell)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <style>
        /*Style f&#252;r Buttons bei tutoriellen Aufgaben */
button.next, button.top-next {
    border: 1px solid black;
background-color: #A9F5BC;
}
button.next:hover, button.top-next:hover {
        background-color: #E0F8E6;
}
button.prev, button.top {
    border: 1px solid black;
    background-color: #DCB25E;
}
button.prev:hover, button.top:hover {
        background-color: #E7D2A7;
}




/* Style f&#252;r Details */
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <details>
        <summary>
            Inhalt und Umfang dieser Aufgabe
        </summary>
        <p>
            In dieser Aufgabe lernen Sie, Eigenschaften von Zufallsvariablen mit stetiger Verteilung zu bestimmen. Sie besteht aus den folgenden zwei Teilaufgaben.
        </p>
        <ul>
            <li>
                In
                <b>
                    (a)
                </b>
                berechnen Sie eine Wahrscheinlichkeit mit Hilfe einer Dichtefunktion.
            </li>
            <li>
                In
                <b>
                    (b)
                </b>
                berechnen Sie einen Erwartungswert mit Hilfe einer Dichtefunktion.
            </li>
        </ul>
    </details>
    <div class="question-view">
        <p>
            In dieser Aufgabe fassen wir die Lebensdauer eines Fahrradschlauchs
        in Jahren als eine Zufallsvariable \(X\) auf.  Als eine einfache
        Modellannahme gehen wir davon aus, dass die Dichte von \(X\) gegeben
        ist durch

  
  
  
  \[
        f(x) = \begin{cases}
        {@ pdf(x) @}&amp; \text{ f&#252;r } 0 \leq x \leq {@ bb @}\\
        0 &amp; \text{ sonst. }
        \end{cases}.
        \]
        </p>
        <div class="question-view">
            <div class="question-view">
                <p>
                    <b>
                        (a)
                    </b>
                    Berechnen Sie die Wahrscheinlichkeit \(P( {@ lo @} \leq X \leq {@ hi @})\).
                </p>
                Antwort: [[input:SAns1]]
        [[validation:SAns1]]
        [[feedback:prt]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a.1)
                    </b>
                    Bei Zufallsvariablen mit stetiger Verteilung k&#246;nnen wir Wahrscheinlichkeiten mithilfe eines Integrals &#252;ber die Dichte ausrechnen. Wenn Sie die Wahrscheinlichkeit \(P( {@ lo @} \leq X \leq {@ hi @})\) als Integral schreiben, welches Integral ergibt sich?
                </p>
                [[input:integral]]
        [[validation:integral]]
        [[feedback:integral]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a.2)
                    </b>
                    Im vorherigen Schritt haben Sie das Integral \[ \int_{{@ lo @}}^{{@hi@}} {@pdf(x)@} dx \] ausgew&#228;hlt. Geben Sie eine Stammfunktion f&#252;r den Integranden an.
                </p>
                \(F(x) = \) [[input:antiderivative]]
        [[validation:antiderivative]]
        [[feedback:antiderivative]]
            </div>
        </div>
        <div class="question-view">
            <div class="question-view">
                <p>
                    <b>
                        (b)
                    </b>
                    Berechnen Sie die erwartete Lebensdauer eines Schlauchs.
                </p>
                Antwort: [[input:SAns2]] [[validation:SAns2]]
        [[feedback:prt_ans2]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (b.1)
                    </b>
                    Sie k&#246;nnen die erwartete Lebensdauer eines neu gekauften Schlauchs als Integral schreiben.  W&#228;hlen Sie das richtige Integral aus:
                </p>
                [[input:expected_integral]]
        [[validation:expected_integral]]
        [[feedback:expected_integral]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (b.2)
                    </b>
                    Im vorherigen Schritt haben Sie das Integral \[ \int_0^{{@bb@}} {@x*pdf(x)@} dx \] ausgew&#228;hlt. Geben Sie eine Stammfunktion f&#252;r den Integranden an.
                </p>
                \(F(x) = \) [[input:lifespanAntideriv]]
        [[validation:lifespanAntideriv]]
        [[feedback:lifespanAntideriv]]
            </div>
        </div>
        <script>
            /*
    * This program is free software: you can redistribute it and/or
    * modify it under the terms of the GNU General Public License as
    * published by the Free Software Foundation, either version 3 of the
    * License, or (at your option) any later version.
    *
    * This program is distributed in the hope that it will be useful, but
    * WITHOUT ANY WARRANTY; without even the implied warranty of
    * MERCHANTABILITY or FITNESS FOR A PARTICULAR PURPOSE.  See the GNU
    * General Public License for more details.
    *
    * You should have received a copy of the GNU General Public License
    * along with this program.  If not, see <https://www.gnu.org/licenses/>.
    *
    * Copyright (C) 2021 by Daniel Mei&#223;ner <daniel.meissner-i4k@ruhr-uni-bochum.de>
    * SPDX-License-Identifier: GPL-3.0-or-later
    *
    */




/* @license magnet:?xt=urn:btih:1f739d935676111cfff4b4693e3816e664797050&dn=gpl-3.0.txt GPL-v3-or-Later */




(function() {
/* Make sure to create QController instances only once even if the
            script is loaded multiple times.  This is so that the script can
            be included in more than one STACK question on the page and still
            function properly. */
if (window.QController === undefined) {
        window.QController = QController;
        document.addEventListener("DOMContentLoaded", event => {
                const responseform = document.getElementById("responseform");
                if (!responseform) {
                        throw Error("Element with id responseform not found.");
                }

    
    
    
    responseform.querySelectorAll(".stack.que").forEach(stackQue => {
                        const rootElement = stackQue.querySelector(".question-view");

      
      
      
      if (!rootElement) {
                                console.warn("QController: no elements with class question-view inside " + stackQue.id);
                                return;
                        }

      
      
      
      const controller = new QController(rootElement, stackQue.id);
                        controller.installHandlers();
                        controller.loadState();
                        controller.update();
                });
        });
}




/* This class handles the navigation of a single STACK question */
function QController(rootElement, questionId) {
        this.rootElement = rootElement;
        this.id = questionId;
        this.views = rootElement.querySelectorAll(".question-view");
        this.visited = (new Array(this.views.length)).fill(false);
        this.currentView = Array.prototype.findIndex.call(this.views, isLeaf);
        this.visited[this.currentView] = true;
}




QController.prototype.handleVisitedNodes = function(current) {
        var visitedSiblings = [];
        var el = current;
        while (el.nextElementSibling) {
                var index =
                                Array.prototype.indexOf.call(this.views, el.nextElementSibling);
                visitedSiblings.push(this.visited[index])
                el = el.nextElementSibling;
        }

  
  
  
  var withoutHelp = visitedSiblings.every(x => x == false);
        var helpCompleted = visitedSiblings.every(x => x == true);

  
  
  
  current.querySelectorAll(".visited").forEach(function(element) {
                var index = element.getAttribute("data-qcontroller-help");
                if (!index || index === "none")
                        element.hidden = !withoutHelp;
                else if (index === "complete")
                        element.hidden = !helpCompleted;
                else if (index === "partial")
                        element.hidden = withoutHelp || helpCompleted;
                else
                        element.hidden = !visitedSiblings[index - 1];
        });
};




/* Update view according to current controller state */
QController.prototype.update = function() {
        this.views.forEach(hide);

  
  
  
  const current = this.views[this.currentView];
        if (isLeaf(current)) {
                showView(current);
                this.handleVisitedNodes(current);
        } else {
                this.showFirstView(current);
        }
};




/* Install event handlers for all buttons under root element */
QController.prototype.installHandlers = function() {
        const controller = this;
        this.rootElement.querySelectorAll("button").forEach(button => {
                button.addEventListener("click", ev => {
                        const view = button.closest(".question-view");

      
      
      
      if (!view) {
                                console.log("qcontroller: <button> not inside question view");
                                return;
                        }

      
      
      
      const classList = button.classList;
                        let forward = null;

      
      
      
      if (classList.contains("next")) {
                                forward = view.nextElementSibling;
                        } else if (classList.contains("top-next")) {
                                const parent = view.parentNode;
                                if (parent) {
                                        forward = parent.nextElementSibling;
                                } else {
                                        forward = null;
                                }
                        } else if (classList.contains("top")) {
                                forward = containingView(view);
                        } else if (classList.contains("prev")) {
                                forward = view.previousElementSibling;
                        }

      
      
      
      if (forward && forward.classList.contains("question-view")) {
                                controller.visit(forward);
                                controller.saveState();
                                controller.update();
                        }
                });
        });
};




QController.prototype.hasVisited = function(viewIndex) {
        return this.visited[viewIndex];
};




QController.prototype.visit = function(nextView) {
        const index = Array.prototype.findIndex.call(
                this.views,
                x => x === nextView);
        this.currentView = index;
        this.visited[index] = true;
};




/* Save controller state to current session */
QController.prototype.saveState = function() {
        const key = this.id + "-controller-state";
        const state = {
                currentView: this.currentView,
                visited: this.visited
        };
        window.sessionStorage.setItem(key, JSON.stringify(state));
};




/* Load controller state from current session */
QController.prototype.loadState = function() {
        const key = this.id + "-controller-state";
        const state = JSON.parse(window.sessionStorage.getItem(key));

  
  
  
  if (state) {
                this.currentView = state.currentView;
                this.visited = state.visited;
        }
};




/* Show the first leaf question view */
QController.prototype.showFirstView = function(questionView) {
        let view = questionView;

  
  
  
  while (view && !isLeaf(view)) {
                view = view.querySelector(".question-view");
        }

  
  
  
  showView(view);
        this.handleVisitedNodes(view);
}




function showView(view) {
        /* show all containing views */
        while (view) {
                show(view);
                view = containingView(view);
        }
}




/* Return the view containing VIEW */
function containingView(view) {
        let parent = view.parentNode;

  
  
  
  if (parent) {
                parent = parent.closest(".question-view");
        }

  
  
  
  return parent;
}




function show(view) {
        view.style.display = "block";
}




function isLeaf(view) {
        return view.querySelector(".question-view") == null;
}




function isVisible(view) {
        return view.style.display != "none";
}




function hide(view) {
        view.style.display = "none";
}




})();
/* @license-end */
        </script>
    </div>
</span>
<span class="multilang" lang="en">
    <style>
        /*Style f&#252;r Buttons bei tutoriellen Aufgaben */
button.next, button.top-next {
    border: 1px solid black;
background-color: #A9F5BC;
}
button.next:hover, button.top-next:hover {
        background-color: #E0F8E6;
}
button.prev, button.top {
    border: 1px solid black;
    background-color: #DCB25E;
}
button.prev:hover, button.top:hover {
        background-color: #E7D2A7;
}




/* Style f&#252;r Details */
.que.stack .formulation details {
                border: 1px solid #aaa;
                border-radius: 4px;
                padding: .5em .5em 0;
}
.que.stack .formulation summary {
                font-weight: bold;
                margin: -.5em -.5em 0;
                padding: .5em;
}
.que.stack .formulation details[open] {
                padding: .5em;
}
.que.stack .formulation details[open] summary {
                border-bottom: 1px solid #aaa;
                margin-bottom: .5em;
}
    </style>
    <details>
        <summary>
            Content and scope of this task
        </summary>
        <p>
            In this task you will learn to determine properties of random variables with constant distribution. It consists of the following two subtasks.
        </p>
        <ul>
            <li>
                In
                <b>
                    (a)
                </b>
                calculate a probability using a density function.
            </li>
            <li>
                In
                <b>
                    (b)
                </b>
                calculate an expected value using a density function.
            </li>
        </ul>
    </details>
    <div class="question-view">
        <p>
            In this task we grasp the life of a bicycle hose
in years as a random variable \(X\) up. As a simple one
We assume that the density of \(X\) given
by

 
 
 
 \[
        f(x) = \begin{cases}
        {@ pdf(x) @}&amp; \text{ f&#252;r } 0 \leq x \leq {@ bb @}\\
        0 &amp; \text{ sonst. }
        \end{cases}.
        \]
        </p>
        <div class="question-view">
            <div class="question-view">
                <p>
                    <b>
                        (a)
                    </b>
                    Calculate the probability \(P( {@ lo @} \leq X \leq {@ hi @})\).
                </p>
                Answer: [[input:SAns1]][[validation:SAns1]][[feedback:prt]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a)
                    </b>
                    In case of random variables with constant distribution, we can calculate probabilities using an integral over the density. If you have the chance \(P( {@ lo @} \leq X \leq {@ hi @})\) as an integral, which integral results?
                </p>
                [[input:integral]][[validation:integral]][[feedback:integral]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (a.2)
                    </b>
                    In the previous step you have the integral \[ \int_{{@ lo @}}^{{@hi@}} {@pdf(x)@} dx \] selected. Specify a root function for the integrande.
                </p>
                \(F(x) = \) [[input:antiderivative]][[validation:antiderivative]][[feedback:antiderivative]]
            </div>
        </div>
        <div class="question-view">
            <div class="question-view">
                <p>
                    <b>
                        (b)
                    </b>
                    Calculate the expected service life of a hose.
                </p>
                Answer: [[input:SAns2]] [[validation:SAns2]][[feedback:prt_ans2]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (b.1)
                    </b>
                    You can write the expected service life of a newly purchased hose as an integral. Choose the right integral:
                </p>
                [[input:expected_integral]][[validation:expected_integral]][[feedback:expected_integral]]
            </div>
            <div class="question-view">
                <p>
                    <b>
                        (b.2)
                    </b>
                    In the previous step you have the integral \[ \int_0^{{@bb@}} {@x*pdf(x)@} dx \] selected. Specify a root function for the integrande.
                </p>
                \(F(x) = \) [[input:lifespanAntideriv]][[validation:lifespanAntideriv]][[feedback:lifespanAntideriv]]
            </div>
        </div>
        <script>
            /*
    * This program is free software: you can redistribute it and/or
    * modify it under the terms of the GNU General Public License as
    * published by the Free Software Foundation, either version 3 of the
    * License, or (at your option) any later version.
    *
    * This program is distributed in the hope that it will be useful, but
    * WITHOUT ANY WARRANTY; without even the implied warranty of
    * MERCHANTABILITY or FITNESS FOR A PARTICULAR PURPOSE.  See the GNU
    * General Public License for more details.
    *
    * You should have received a copy of the GNU General Public License
    * along with this program.  If not, see <https://www.gnu.org/licenses/>.
*
* Copyright (C) 2021 by Daniel Mei&#223;ner <daniel.meissner-i4k@ruhr-uni-bochum.de>* SPDX-License Identifier: GPL-3.0-or-later
*
*/




/* @license magnet:?xt=urn:btih:1f739d935676111cfff4b4693e3816e664797050&dn=gpl-3.0.txt GPL-v3-orLater */




(function() {
/* Make sure to create QController instances only once even if the
            script is loaded multiple times.  This is so that the script can
            be included in more than one STACK question on the page and still
            function properly. */
if (window.QController === undefined) {
        window.QController = QController;
        document.addEventListener("DOMContentLoaded", event => {
                const responseform = document.getElementById("responseform");
                if (!responseform) {
                        throw Error("Element with id responseform not found.");
                }

    
    
    
    responseform.querySelectorAll(".stack.que").forEach(stackQue => {
                        const rootElement = stackQue.querySelector(".question-view");

      
      
      
      if (!rootElement) {
                                console.warn("QController: no elements with class question-view inside " + stackQue.id);
                                return;
                        }

      
      
      
      const controller = new QController(rootElement, stackQue.id);
                        controller.installHandlers();
                        controller.loadState();
                        controller.update();
                });
        });
}




/* This class handles the navigation of a single STACK question */
function QController(rootElement, questionId) {
        this.rootElement = rootElement;
        this.id = questionId;
        this.views = rootElement.querySelectorAll(".question-view");
        this.visited = (new Array(this.views.length)).fill(false);
        this.currentView = Array.prototype.findIndex.call(this.views, isLeaf);
        this.visited[this.currentView] = true;
}




QController.prototype.handleVisitedNodes = function(current) {
        var visitedSiblings = [];
        var el = current;
        while (el.nextElementSibling) {
                var index =
                                Array.prototype.indexOf.call(this.views, el.nextElementSibling);
                visitedSiblings.push(this.visited[index])
                el = el.nextElementSibling;
        }

  
  
  
  var withoutHelp = visitedSiblings.every(x => x == false);
        var helpCompleted = visitedSiblings.every(x => x == true);

  
  
  
  current.querySelectorAll(".visited").forEach(function(element) {
                var index = element.getAttribute("data-qcontroller-help");
                if (!index || index === "none")
                        element.hidden = !withoutHelp;
                else if (index === "complete")
                        element.hidden = !helpCompleted;
                else if (index === "partial")
                        element.hidden = withoutHelp || helpCompleted;
                else
                        element.hidden = !visitedSiblings[index - 1];
        });
};




/* Update view according to current controller state */
QController.prototype.update = function() {
        this.views.forEach(hide);

  
  
  
  const current = this.views[this.currentView];
        if (isLeaf(current)) {
                showView(current);
                this.handleVisitedNodes(current);
        } else {
                this.showFirstView(current);
        }
};




/* Install event handlers for all buttons under root element */
QController.prototype.installHandlers = function() {
        const controller = this;
        this.rootElement.querySelectorAll("button").forEach(button => {
                button.addEventListener("click", ev => {
                        const view = button.closest(".question-view");

      
      
      
      if (!view) {
                                console.log("qcontroller: <button> not inside question view");
                                return;
                        }

      
      
      
      const classList = button.classList;
                        let forward = null;

      
      
      
      if (classList.contains("next")) {
                                forward = view.nextElementSibling;
                        } else if (classList.contains("top-next")) {
                                const parent = view.parentNode;
                                if (parent) {
                                        forward = parent.nextElementSibling;
                                } else {
                                        forward = null;
                                }
                        } else if (classList.contains("top")) {
                                forward = containingView(view);
                        } else if (classList.contains("prev")) {
                                forward = view.previousElementSibling;
                        }

      
      
      
      if (forward && forward.classList.contains("question-view")) {
                                controller.visit(forward);
                                controller.saveState();
                                controller.update();
                        }
                });
        });
};




QController.prototype.hasVisited = function(viewIndex) {
        return this.visited[viewIndex];
};




QController.prototype.visit = function(nextView) {
        const index = Array.prototype.findIndex.call(
                this.views,
                x => x === nextView);
        this.currentView = index;
        this.visited[index] = true;
};




/* Save controller state to current session */
QController.prototype.saveState = function() {
        const key = this.id + "-controller-state";
        const state = {
                currentView: this.currentView,
                visited: this.visited
        };
        window.sessionStorage.setItem(key, JSON.stringify(state));
};




/* Load controller state from current session */
QController.prototype.loadState = function() {
        const key = this.id + "-controller-state";
        const state = JSON.parse(window.sessionStorage.getItem(key));

  
  
  
  if (state) {
                this.currentView = state.currentView;
                this.visited = state.visited;
        }
};




/* Show the first leaf question view */
QController.prototype.showFirstView = function(questionView) {
        let view = questionView;

  
  
  
  while (view && !isLeaf(view)) {
                view = view.querySelector(".question-view");
        }

  
  
  
  showView(view);
        this.handleVisitedNodes(view);
}




function showView(view) {
        /* show all containing views */
        while (view) {
                show(view);
                view = containingView(view);
        }
}




/* Return the view containing VIEW */
function containingView(view) {
        let parent = view.parentNode;

  
  
  
  if (parent) {
                parent = parent.closest(".question-view");
        }

  
  
  
  return parent;
}




function show(view) {
        view.style.display = "block";
}




function isLeaf(view) {
        return view.querySelector(".question-view") == null;
}




function isVisible(view) {
        return view.style.display != "none";
}




function hide(view) {
        view.style.display = "none";
}




}),
/* @license-end */
        </script>
    </div>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/* This work is licensed under the Creative Commons Attribution-ShareAlike
   4.0 International License. To view a copy of this license, visit
   http://creativecommons.org/licenses/by-sa/4.0/.

   Copyright 2021, 2022 Ruhr-Universit&#228;t Bochum
   SPDX-License-Identifier: CC-BY-SA-4.0

   Authors: Daniel Mei&#223;ner
   Initial idea: Herold Dehling  */

/* utilities */

probabilityp(p) := is(numberp(p) and p >= 0 and p <= 1);

/* actual question variables */

bb: rand([3, 4, 5, 6]);
mm: -2/bb^2;
pdf(x) := -mm*(bb - x);
density_supportp(x) := is(numberp(x) and x >= 0 and x <= bb);

lo: bb - rand([1, 2]);
hi: lo + 1;

/* first char of value string indicates correct/false limits,
 * second char of value string indicates correct/false integrand
 *
 * E. g.: "ft" means that the limits are correct, whereas
 * the density is wrong
 */
integrals: map (lambda ([x], [x[1], x[2], simplode (x[3])]), [
  ["ft", false, ["\\(\\int_", 0, "^", bb, stack_disp(pdf(x), ""), "dx\\)"]],
  ["tt", true, ["\\(\\int_",lo,"^", hi, stack_disp(pdf(x), ""), "dx\\)"]],
  ["tf", false, ["\\(\\int_",lo,"^", hi, stack_disp(x*pdf(x), ""), "dx\\)"]],
  ["ff", false, ["\\(\\int_",0,"^", bb, stack_disp(x*pdf(x), "") ,"dx\\)"]]
]);

expected_integrals: map (lambda ([x], [x[1], x[2], simplode (x[3])]), [
  ["ft", false, ["\\(\\int_{-\\infty}^{\\infty}", stack_disp(x*pdf(x), ""), "dx\\)"]],
  ["tt", true, ["\\(\\int_",0,"^", bb, stack_disp(x*pdf(x), ""), "dx\\)"]],
  ["tf", false, ["\\(\\int_",0,"^", bb, stack_disp(pdf(x), ""), "dx\\)"]],
  ["ff", false, ["\\(\\int_{-\\infty}^{\\infty}", stack_disp(pdf(x), ""), "dx\\)"]]
]);

TAns1: int(pdf(x), x, lo, hi);
TAns2: int(x*pdf(x), x, 0, bb);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>b: {# bb #}, lo: {# lo #}
Wahrscheinlichkeit: {# TAns1 #}
Erwartungswert {# TAns2 #}</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist nicht komplett falsch.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is not completely wrong.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed>i</variantsselectionseed>
        <input>
            <name>antiderivative</name>
            <type>algebraic</type>
            <tans>integrate(pdf(u),u,0,x)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>expected_integral</name>
            <type>radio</type>
            <tans>expected_integrals</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>integral</name>
            <type>radio</type>
            <tans>integrals</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options>nonotanswered</options>
        </input>
        <input>
            <name>lifespanAntideriv</name>
            <type>algebraic</type>
            <tans>integrate(u*pdf(u),u,0,x)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>SAns1</name>
            <type>numerical</type>
            <tans>TAns1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>SAns2</name>
            <type>numerical</type>
            <tans>TAns2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>antiderivative</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text>DF: diff(antiderivative, x);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>DF</sans>
                <tans>pdf(x)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>antiderivative-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Sie haben eine korrekte Stammfunktion eingegeben: Die Ableitung der von Ihnen eingegebenen Funktion lautet {@DF@} und stimmt mit der Dichtefunktion im relevanten Intervall &#252;berein.
    </p>
    <p>
        Klicken Sie
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        , um die urspr&#252;ngliche Frage noch einmal zu probieren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        You entered a correct root function: The derivation of the function you have entered is {@DF@} and agrees with the density function in the relevant interval.
    </p>
    <p>
        Click
        <button class="top" type="button">
            Back
        </button>
        to try the original question again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>antiderivative-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Die Ableitung der von Ihnen angegebenen Funktion lautet {@ DF @}.  Dies stimmt nicht mit dem angegebenen Integranden &#252;berein.
Schauen Sie noch einmal in Ihren Notizen oder im Internet zum Thema
        <em>
            Stammfunktionen von Potenzfunktionen
        </em>
        .  Falls Sie absolut nicht weiterkommen, k&#246;nnen Sie auch hier Hilfe bekommen:
    </p>
    <details>
        <summary>
            <em>
                Hinweis
            </em>
        </summary>
        Eine Stammfunktion f&#252;r \(x \mapsto x^n\) ist gegeben durch \(x \mapsto \frac{x^{n+1}}{n+1}\).
    </details>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The derivation of the function you provide is {@ DF @}. This is not consistent with the specified integrande.
Take a look at your notes or on the Internet
        <em>
            Common functions of potency functions
        </em>
        . If you do not get any further, you can also get help here:
    </p>
    <details>
        <summary>
            <em>
                Note
            </em>
        </summary>
        A master function \(x \mapsto x^n\) is given by \(x \mapsto \frac{x^{n+1}}{n+1}\).
    </details>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>expected_integral</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>expected_integral</sans>
                <tans>first(mcq_correct(expected_integrals))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>expected_integral-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Um den Erwartungswert auszurechnen, m&#252;ssen Sie {@ x*pdf(x) @} in den Grenzen von {#0#} bis {#bb#} integrieren.
    </p>
    Klicken Sie auf
    <button class="next" type="button">
        Weiter
    </button>
    , damit es weitergeht.
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        To calculate the expected value, you must {@ x*pdf(x) @} within the limits of {#0#} to {#bb#} integration.
    </p>
    Click on
    <button class="next" type="button">
        Next
    </button>
    to keep going.
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>expected_integral-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>CasEqual</answertest>
                <sans>expected_integral</sans>
                <tans><![CDATA["ft"]]></tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>expected_integral-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Das ausgew&#228;hlte Integral existiert nicht. Um den Erwartungswert auszurechnen, m&#252;ssen Sie das Integral \[\int_{-\infty}^\infty x f(x) dx\] berechnen. Um dieses Integral zu berechnen, reicht es aus, &#252;ber die Bereiche zu integrieren, auf denen f nicht Null
ist. Das von Ihnen ausgew&#228;hlte Integral hat die falschen Grenzen.
    </p>
    <p>
        Versuchen Sie es noch einmal und klicken Sie auf  &#8222;Pr&#252;fen&#8220;.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The selected integral does not exist. To calculate the expected value, you need the integral \[\int_{-\infty}^\infty x f(x) dx\] calculate. In order to calculate this integral, it is sufficient to integrate over the areas where f is not zero
is. The integral you selected has the wrong limits.
    </p>
    <p>
        Try it again and click &#8220;Check&#8221;.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>expected_integral-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>CasEqual</answertest>
                <sans>expected_integral</sans>
                <tans><![CDATA["tf"]]></tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>expected_integral-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Das ausgew&#228;hlte Integral hat stets den Wert 1, da \(f\) eine Wahrscheinlichkeitsdichte ist. Um den Erwartungswert auszurechnen, m&#252;ssen Sie das Integral \[\int_{-\infty}^\infty x f(x) dx\] berechnen. Um dieses Integral zu berechnen, reicht es aus, &#252;ber die Bereiche zu integrieren, auf denen f nicht Null
ist. Das von Ihnen ausgew&#228;hlte Integral hat den falschen Integranden.
    </p>
    <p>
        Versuchen Sie es noch einmal und klicken Sie auf  &#8222;Pr&#252;fen&#8220;.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The selected integral always has the value 1, since \(f\) is a probability density. To calculate the expected value, you need the integral \[\int_{-\infty}^\infty x f(x) dx\] calculate. In order to calculate this integral, it is sufficient to integrate over the areas where f is not zero
is. The integral you selected has the wrong integral.
    </p>
    <p>
        Try it again and click &#8220;Check&#8221;.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>expected_integral-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>CasEqual</answertest>
                <sans>expected_integral</sans>
                <tans><![CDATA["ff"]]></tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>expected_integral-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Das ausgew&#228;hlte Integral existiert nicht. Um den Erwartungswert auszurechnen, m&#252;ssen Sie das Integral \[\int_{-\infty}^\infty x f(x) dx\] berechnen. Um dieses Integral zu berechnen, reicht es aus, &#252;ber die Bereiche zu integrieren, auf denen f nicht Null
ist. Das von Ihnen ausgew&#228;hlte Integral hat den falschen Integranden und die falschen Grenzen.
    </p>
    <p>
        Versuchen Sie es noch einmal und klicken Sie auf  &#8222;Pr&#252;fen&#8220;.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The selected integral does not exist. To calculate the expected value, you need the integral \[\int_{-\infty}^\infty x f(x) dx\] calculate. In order to calculate this integral, it is sufficient to integrate over the areas where f is not zero
is. The integral selected by you has the wrong integral and the wrong limits.
    </p>
    <p>
        Try it again and click &#8220;Check&#8221;.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>expected_integral-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>integral</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>integral</sans>
                <tans>first(mcq_correct(integrals))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>integral-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Um die Wahrscheinlichkeit \(P( {@lo@} \le X \le {@hi@})\) auszurechnen, m&#252;ssen Sie die Dichte in den Grenzen von {#lo#} bis {#hi#} integrieren. Da die Dichte auf dem Intervall \([{@lo@}, {@hi@}]\) den Wert
{@pdf(x)@} annimmt, ergibt sich der angegebene Integrand.
    </p>
    Klicken Sie auf
    <button class="next" type="button">
        Weiter
    </button>
    , damit es weitergeht.
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        To the likelihood \(P( {@lo@} \le X \le {@hi@})\) to calculate the density within the limits of {#lo#} to {#hi#} integration. Since the density on the interval \([{@lo@}, {@hi@}]\) the value
{@pdf(x)@} the indicated integrand.
    </p>
    Click on
    <button class="next" type="button">
        Next
    </button>
    to keep going.
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>integral-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>CasEqual</answertest>
                <sans>integral</sans>
                <tans><![CDATA["ft"]]></tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>integral-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Um die Wahrscheinlichkeit \(P( {@lo@} \le X \le {@hi@})\) auszurechnen, m&#252;ssen Sie die Dichte in den Grenzen von {#lo#} bis {#hi#} integrieren. Das von Ihnen ausgew&#228;hlte Integral hat die falschen Grenzen.
    </p>
    <p>
        Versuchen Sie es noch einmal und klicken Sie auf  &#8222;Pr&#252;fen&#8220;.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        To the likelihood \(P( {@lo@} \le X \le {@hi@})\) to calculate the density within the limits of {#lo#} to {#hi#} integration. The integral you selected has the wrong limits.
    </p>
    <p>
        Try it again and click &#8220;Check&#8221;.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>integral-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>CasEqual</answertest>
                <sans>integral</sans>
                <tans><![CDATA["tf"]]></tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>integral-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Um die Wahrscheinlichkeit \(P( {@lo@} \le X \le {@hi@})\) auszurechnen, m&#252;ssen Sie die Dichte in den Grenzen von {#lo#} bis {#hi#} integrieren. Das von Ihnen ausgew&#228;hlte Integral hat den falschen Integranden.
    </p>
    <p>
        Versuchen Sie es noch einmal und klicken Sie auf  &#8222;Pr&#252;fen&#8220;.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        To the likelihood \(P( {@lo@} \le X \le {@hi@})\) to calculate the density within the limits of {#lo#} to {#hi#} integration. The integral you selected has the wrong integral.
    </p>
    <p>
        Try it again and click &#8220;Check&#8221;.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>integral-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>CasEqual</answertest>
                <sans>integral</sans>
                <tans><![CDATA["ff"]]></tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>integral-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Um die Wahrscheinlichkeit \(P( {@lo@} \le X \le {@hi@})\) auszurechnen, m&#252;ssen Sie die Dichte in den Grenzen von {#lo#} bis {#hi#} integrieren. Das von Ihnen ausgew&#228;hlte Integral hat den falschen Integranden und die falschen Grenzen.
    </p>
    <p>
        Versuchen Sie es noch einmal und klicken Sie auf  &#8222;Pr&#252;fen&#8220;.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        To the likelihood \(P( {@lo@} \le X \le {@hi@})\) to calculate the density within the limits of {#lo#} to {#hi#} integration. The integral selected by you has the wrong integral and the wrong limits.
    </p>
    <p>
        Try it again and click &#8220;Check&#8221;.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>integral-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>lifespanAntideriv</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text>DF: diff(lifespanAntideriv,x);</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>diff(lifespanAntideriv,x)</sans>
                <tans>x*pdf(x)</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>lifespanAntideriv-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Sie haben eine korrekte Stammfunktion eingegeben: Die Ableitung der von Ihnen eingegebenen Funktion lautet {@DF@} und stimmt mit dem Integranden &#252;berein. Um eine Stammfunktion f&#252;r den Integranden zu erhalten, k&#246;nnen Sie die Formel \(\int x^n \;\mathrm{d}x = \frac{x^{n+1}}{n+1} + c\) verwenden. Hierbei ist \(c \in \mathbb{R}\) eine beliebige Konstante.
    </p>
    <p>
        Klicken Sie
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        , um die urspr&#252;ngliche Frage noch einmal zu probieren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        You entered a correct root function: The derivation of the function you have entered is {@DF@} and agrees with the integrande. To get a root function for the integrande, you can formula \(\int x^n \;\mathrm{d}x = \frac{x^{n+1}}{n+1} + c\) use. Here, \(c \in \mathbb{R}\) any constant.
    </p>
    <p>
        Click
        <button class="top" type="button">
            Back
        </button>
        to try the original question again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>lifespanAntideriv-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Die Ableitung der von Ihnen angegebenen Funktion lautet {@ DF @}.  Dies stimmt nicht mit dem angegebenen Integranden &#252;berein.
Schauen Sie noch einmal in Ihren Notizen oder im Internet zum Thema
        <em>
            Stammfunktionen von Potenzfunktionen
        </em>
        .  Falls Sie absolut nicht weiterkommen k&#246;nnen Sie auch hier Hilfe bekommen:
    </p>
    <details>
        <summary>
            <em>
                Hinweis
            </em>
        </summary>
        Eine Stammfunktion f&#252;r \(x \mapsto x^n\) ist gegeben durch \(x \mapsto \frac{x^{n+1}}{n+1}\).
    </details>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The derivation of the function you provide is {@ DF @}. This is not consistent with the specified integrande.
Take a look at your notes or on the Internet
        <em>
            Common functions of potency functions
        </em>
        . If you do not get any further, you can also get help here:
    </p>
    <details>
        <summary>
            <em>
                Note
            </em>
        </summary>
        A master function \(x \mapsto x^n\) is given by \(x \mapsto \frac{x^{n+1}}{n+1}\).
    </details>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>SAns1</sans>
                <tans>TAns1</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Die angegebene Wahrscheinlichkeit ist gegeben als das Integral &#252;ber die Dichte der Zufallsvariable in den Grenzen von {#lo#} bis {#hi#}.
    </p>
    <p>
        Klicken Sie auf
        <button class="top-next" type="button">
            Weiter
        </button>
        , um zum n&#228;chsten Aufgabenteil zu gelangen
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The indicated probability is given as the integral over the density of random variable within the limits of {#lo#} to {#hi#}.
    </p>
    <p>
        Click on
        <button class="top-next" type="button">
            Next
        </button>
        to get to the next task
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="visited" data-qcontroller-help="none">
        Dass Sie die Aufgabe nicht auf Anhieb l&#246;sen konnten, ist kein
        Problem! Im Folgenden wird die Aufgabe in kleinere Schritte
        unterteilt. Wenn Sie diese Schritte alle gel&#246;st haben, k&#246;nnen Sie
        die urspr&#252;ngliche Aufgabe sicherlich l&#246;sen. Klicken Sie dazu auf
        <button class="next" type="button">
            Weiter
        </button>
        .
    </p>
    <p class="visited" data-qcontroller-help="2">
        Willkommen zur&#252;ck!  In der Hilfe haben Sie gelernt, dass Sie die
gesuchte Wahrscheinlichkeit als den Wert des Integrals {@ int(pdf(x),
x, lo, hi) @} bestimmen k&#246;nnen.  Au&#223;erdem haben Sie eine Stammfunktion
f&#252;r den Integranden angegeben, zum Beispiel {@ ev(ratsimp(integrate(pdf(x),x)), simp)@}.
    </p>
</span>
<span class="multilang" lang="en">
    <p class="visited" data-qcontroller-help="none">
        That you couldn't solve the task right away is not a
Problem! In the following, the task is carried out in smaller steps
divided. If you have all solved these steps, you can
the original task will certainly solve. Click on
        <button class="next" type="button">
            Next
        </button>
        .
    </p>
    <p class="visited" data-qcontroller-help="2">
        Welcome back! In the help you have learned that you have
sought probability as the value of the integral {@ int(pdf(x),
x, lo, hi) @} can determine. In addition, you have a master function
for the integrande, for example {@ ev(ratsimp(integrate(pdf(x),x)), simp)@}.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>probabilityp(SAns1)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>0</truenextnode>
                <trueanswernote>prt-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    Wahrscheinlichkeiten sollten im Intervall [0, 1] liegen, das trifft f&#252;r Ihre Eingabe nicht zu.
</span>
<span class="multilang" lang="en">
    Probabilities should be in the interval [0, 1], this does not apply to your input.
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt_ans2</name>
            <value>1.0000000</value>
            <autosimplify>0</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>SAns2</sans>
                <tans>TAns2</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt_ans2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Der Erwartungswert einer stetigen Zufallsvariablen ist gegeben als \(\int_{-\infty}^\infty x f(x) dx\). Da die Dichte \(f\) in diesen Fall positive Werte nur im Intervall [0, {#bb#}] annimmt, ergibt sich der Erwartungswert durch das Berechnen des Integrals
\[
\int_0^{@bb@} {@x * pdf(x)@} dx.
\]
    </p>
    <p>
        Herzlichen Gl&#252;ckwunsch! Sie haben das Ende dieser Aufgabe erreicht.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        The expected value of a steady random variable is given as \(\int_{-\infty}^\infty x f(x) dx\). Since the density \(f\) in this case positive values only in the interval [0, {#bb#}], the expected value is obtained by calculating the integral
\[
\int_0^{@bb@} {@x * pdf(x)@} dx.
\]
    </p>
    <p>
        Congratulations! They have reached the end of this task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt_ans2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="visited" data-qcontroller-help="none">
        Dass Sie die Aufgabe nicht auf Anhieb l&#246;sen konnten, ist kein
        Problem! Im Folgenden wird die Aufgabe in kleinere Schritte
        unterteilt. Wenn Sie diese Schritte alle gel&#246;st haben, k&#246;nnen Sie
        die urspr&#252;ngliche Aufgabe sicherlich l&#246;sen. Klicken Sie dazu auf
        <button class="next" type="button">
            Weiter
        </button>
        .
    </p>
    <p class="visited" data-qcontroller-help="2">
        Willkommen zur&#252;ck!  In der Hilfe haben Sie gelernt, dass Sie den
gesuchte Erwartungswert als den Wert des Integrals {@ int(x*pdf(x),
x, 0, bb) @} bestimmen k&#246;nnen.  Au&#223;erdem haben Sie eine Stammfunktion
f&#252;r den Integranden angegeben, zum Beispiel {@ ev(ratsimp(integrate(x*pdf(x),x)), simp) @}.
    </p>
</span>
<span class="multilang" lang="en">
    <p class="visited" data-qcontroller-help="none">
        That you couldn't solve the task right away is not a
Problem! In the following, the task is carried out in smaller steps
divided. If you have all solved these steps, you can
the original task will certainly solve. Click on
        <button class="next" type="button">
            Next
        </button>
        .
    </p>
    <p class="visited" data-qcontroller-help="2">
        Welcome back! In the help you have learned that
expected value as the value of the integral {@ int(x*pdf(x),
x, 0, bb) @} can determine. In addition, you have a master function
for the integrande, for example {@ ev(ratsimp(integrate(x*pdf(x),x)), simp) @}.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>density_supportp(SAns2)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>0</truenextnode>
                <trueanswernote>prt_ans2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt_ans2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    Der von Ihnen eingegebene Wert liegt nicht im Intervall [0, {#bb#}]. Die angegebene Wahrscheinlichkeitsdichte nimmt allerdings nur in diesem Intervall positive Werte an.  Daher kann der Erwartungswert nicht au&#223;erhalb dieses Intervalls liegen.  Versuchen Sie es noch einmal.
</span>
<span class="multilang" lang="en">
    The value you enter is not in the interval [0, {#bb#}]. However, the indicated probability density only assumes positive values at this interval. Therefore, the expected value cannot lie outside this interval. Try it again.
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>1395178303</deployedseed>
        <deployedseed>1668060886</deployedseed>
        <deployedseed>91458270</deployedseed>
        <deployedseed>656659085</deployedseed>
        <deployedseed>904343076</deployedseed>
        <deployedseed>1831377901</deployedseed>
        <deployedseed>1103511096</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>antiderivative</name>
                <value/>
            </testinput>
            <testinput>
                <name>expected_integral</name>
                <value/>
            </testinput>
            <testinput>
                <name>integral</name>
                <value/>
            </testinput>
            <testinput>
                <name>lifespanAntideriv</name>
                <value/>
            </testinput>
            <testinput>
                <name>SAns1</name>
                <value>TAns1</value>
            </testinput>
            <testinput>
                <name>SAns2</name>
                <value>TAns2</value>
            </testinput>
            <expected>
                <name>antiderivative</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>expected_integral</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>integral</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>lifespanAntideriv</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>prt</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt_ans2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt_ans2-1-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>