<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Modellierung eines Zufallsexperiments (Ablaufplan)</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <style>
        .buttonweiter {
    border: 1px solid black;
background-color: #A9F5BC; 
}
.buttonweiter:hover {
        background-color: #E0F8E6;
}
.buttonzurueck {
    border: 1px solid black;
    background-color: #DCB25E; 
}
.buttonzurueck:hover {
        background-color: #E7D2A7;
}
.buttontipp {
    border: 1px solid black;
    background-color: #c9d3f6; 
}
.buttontipp:hover {
        background-color: #dbe1f9;
}
    </style>
    <p>
        In einer Urne befinden sich \(n=4\) Kugeln, die mit den Zahlen 1 bis 4 beschriftet sind.
    </p>
    <p>
        Wir betrachten das Zufallsexperiment, dass \(k=3\) Kugeln ohne Zur&#252;cklegen gezogen werden. Dabei wird die Reihenfolge, in der die Kugeln gezogen werden, ebenfalls beobachtet.
    </p>
    <div style="margin-top: 1rem;margin-bottom: 1rem">
        <ul class="nav nav-pills" id="oer-aufgabe-10-ablauf">
            <li class="nav-item bg-light">
                <span class="nav-link active">
                    a) Ergebnisraum angeben
                </span>
            </li>
            <li class="nav-item bg-light">
                <span class="nav-link">
                    b) Ereignis angeben
                </span>
            </li>
            <li class="nav-item bg-light">
                <span class="nav-link">
                    c) Wahrscheinlichkeit angeben
                </span>
            </li>
        </ul>
    </div>
    <!-- Erinnerung an alle Dinge, die aus vorherigen Aufgabenteilen bekannt sind -->
    <div class="alert alert-primary rounded border border-primary" id="oer-aufgabe-10_div-bekanntes" style="display:none;padding: 1rem">
    </div>
    <div id="oer-aufgabe-10_aufgabe1" style="display:box">
        <p>
            <b>
                a)
            </b>
            Geben Sie den Ergebnisraum des Zufallsexperiments an.
        </p>
        <p>
            Antwort: \(\Omega=\) [[input:ans1]] [[validation:ans1]]
        </p>
        [[feedback:prt1]]
    </div>
    <div id="oer-aufgabe-10_aufgabe2" style="display:none">
        <p>
            <b>
                b)
            </b>
            Wir betrachten nun das Ereignis \(A\), dass {@summe_produkt_string@} der Zahlen, die auf den gezogenen Kugeln stehen, {@n1@} ist. Bitte geben Sie das Ereignis in der Mengenform an.
        </p>
        <p>
            Antwort: \(A=\) [[input:ans2]] [[validation:ans2]]
        </p>
        [[feedback:prt2]]
    </div>
    <div id="oer-aufgabe-10_aufgabe3" style="display:none">
        <p>
            <b>
                c)
            </b>
            Geben Sie die Wahrscheinlichkeit von \(A\) an.
        </p>
        <p>
            Antwort: \(P(A)=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </p>
        <div id="oer-aufgabe-10_Hinweis_Annahme">
            <i>
                Hinweis: Wir gehen davon aus, dass alle Ergebnisse mit gleicher Wahrscheinlichkeit auftreten.
            </i>
        </div>
    </div>
    <p id="oer-aufgabe-10_Hinweis_Eingabe" style="margin-top:0.5rem;margin-bottom:0.5rem">
        <i>
            Hinweis: Eine Menge k&#246;nnen Sie in der Form
            <code>
                {1,2,3}
            </code>
            und ein Tupel in der Form
            <code>
                [1,2,3]
            </code>
            angeben.
        </i>
    </p>
    <p id="oer-aufgabe-10_Hinweis_Pruefen" style="margin-top:0.5rem;margin-bottom:0.5rem">
        Bitte klicken Sie auf "Pr&#252;fen", wenn Sie die L&#246;sung eingegeben haben.
    </p>
    <script>
        (function() {
                "use strict";
                const id = document.currentScript.closest(".que.stack").id;
                const key = id + "-menu-state";

    
    
    
    function update(state) {
                                const listGroup = document.getElementById("oer-aufgabe-10-ablauf");
                                listGroup.querySelectorAll("li > span").forEach(function (item, ind) {
                                                if (ind == state) {
                                                                item.classList.add("active");
                                                } else {
                                                                item.classList.remove("active");
                                                }
                                });
                }

    
    
    
    function load() {
                                const state = window.sessionStorage.getItem(key);
                                if (state) {
                                                return state;
                                } else {
                                                return 0;
                                }
                }

    
    
    
    function storeAndUpdate(state) {
                                window.sessionStorage.setItem(key, state);
                                update(state);
                }

    
    
    
    document.addEventListener("DOMContentLoaded", function(event) {
                                update(load());
                });

    
    
    
    let QMenu = {};
                QMenu.update = update;
                QMenu.storeAndUpdate = storeAndUpdate;
                window.QMenu = QMenu;
})();
    </script>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
<span class="multilang" lang="en">
    <style>
        .buttonweiter {
    border: 1px solid black;
background-color: #A9F5BC; 
}
.buttonweiter:hover {
        background-color: #E0F8E6;
}
.buttonzurueck {
    border: 1px solid black;
    background-color: #DCB25E; 
}
.buttonzurueck:hover {
        background-color: #E7D2A7;
}
.buttontipp {
    border: 1px solid black;
    background-color: #c9d3f6; 
}
.buttontipp:hover {
        background-color: #dbe1f9;
}
    </style>
    <p>
        In a urn \(n=4\) Balls marked with numbers 1 to 4.
    </p>
    <p>
        We consider the random experiment that \(k=3\) without being pulled back. The sequence in which the balls are drawn is also observed.
    </p>
    <div style="margin-top: 1rem;margin-bottom: 1rem">
        <ul class="nav nav-pills" id="oer-aufgabe-10-ablauf">
            <li class="nav-item bg-light">
                <span class="nav-link active">
                    a) Specify the results
                </span>
            </li>
            <li class="nav-item bg-light">
                <span class="nav-link">
                    b) Specify event
                </span>
            </li>
            <li class="nav-item bg-light">
                <span class="nav-link">
                    c) Specify probability
                </span>
            </li>
        </ul>
    </div>
    <!-- Erinnerung an alle Dinge, die aus vorherigen Aufgabenteilen bekannt sind -->
    <div class="alert alert-primary rounded border border-primary" id="oer-aufgabe-10_div-bekanntes" style="display:none;padding: 1rem">
    </div>
    <div id="oer-aufgabe-10_aufgabe1" style="display:box">
        <p>
            <b>
                (a)
            </b>
            Specify the result space of the random experiment.
        </p>
        <p>
            Answer: \(\Omega=\) [[input:ans1]] [[validation:ans1]]
        </p>
        [[feedback:prt1]]
    </div>
    <div id="oer-aufgabe-10_aufgabe2" style="display:none">
        <p>
            <b>
                (b)
            </b>
            We now look at the event \(A\)that {@summe_produkt_string@} the numbers standing on the drawn balls, {@n1@} is. Please indicate the event in the quantity form.
        </p>
        <p>
            Answer: \(A=\) [[input:ans2]] [[validation:ans2]]
        </p>
        [[feedback:prt2]]
    </div>
    <div id="oer-aufgabe-10_aufgabe3" style="display:none">
        <p>
            <b>
                (c)
            </b>
            Give the probability of \(A\) on.
        </p>
        <p>
            Answer: \(P(A)=\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
        </p>
        <div id="oer-aufgabe-10_Hinweis_Annahme">
            <i>
                Note: We assume that all results occur with the same probability.
            </i>
        </div>
    </div>
    <p id="oer-aufgabe-10_Hinweis_Eingabe" style="margin-top:0.5rem;margin-bottom:0.5rem">
        <i>
            Note: A lot you can in the form
            <code>
                {1,2,3}
            </code>
            and a tuple in shape
            <code>
                [1,2,3]
            </code>
            indicate.
        </i>
    </p>
    <p id="oer-aufgabe-10_Hinweis_Pruefen" style="margin-top:0.5rem;margin-bottom:0.5rem">
        Please click "Check" if you have entered the solution.
    </p>
    <script>
        (function() {
                "use strict";
                const id = document.currentScript.closest(".que.stack").id;
                const key = id + "-menu-state";

    
    
    
    function update(state) {
                                const listGroup = document.getElementById("oer-aufgabe-10-ablauf");
                                listGroup.querySelectorAll("li > span").forEach(function (item, ind) {
                                                if (ind == state) {
                                                                item.classList.add("active");
                                                } else {
                                                                item.classList.remove("active");
                                                }
                                });
                }

    
    
    
    function load() {
                                const state = window.sessionStorage.getItem(key);
                                if (state) {
                                                return state;
                                } else {
                                                return 0;
                                }
                }

    
    
    
    function storeAndUpdate(state) {
                                window.sessionStorage.setItem(key, state);
                                update(state);
                }

    
    
    
    document.addEventListener("DOMContentLoaded", function(event) {
                                update(load());
                });

    
    
    
    let QMenu = {};
                QMenu.update = update;
                QMenu.storeAndUpdate = storeAndUpdate;
                window.QMenu = QMenu;
})();
    </script>
    <script src="https://www.rub.de/ak-mathe-digital/stackselbstlern.js">
    </script>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*Autor der Aufgabe: Jonas Lache, Ruhr-Universit&#228;t Bochum*/
/*ID der Aufgabe: oer-aufgabe-10*/
/* evtl. zu aktualisierende Fremdmaterialien/Links/Bilder/Daten: 
- Im Fragetext wird auf eine Javascript-Datei der Ruhr-Universit&#228;t Bochum verwiesen*/
/*Sonstige Hinweise:
- Aufgrund des langen mathematischen Ausdrucks wird je nach Browser nicht der gesamte Ausdruck in der Validierung angezeigt.
*/

Omega: set([1,2,3],[1,3,2],[1,2,4],[1,4,2],[1,3,4],[1,4,3],
  [2,1,3],[2,3,1],[2,1,4],[2,4,1],[2,3,4],[2,4,3],
  [3,1,2],[3,2,1],[3,1,4],[3,4,1],[3,2,4],[3,4,2],
  [4,1,2],[4,2,1],[4,2,3],[4,3,2],[4,1,3],[4,3,1]
);
Omega_part1: listify(subset(Omega,lambda([li],is(li[1]=1))));
Omega_part2: listify(subset(Omega,lambda([li],is(li[1]=2))));
Omega_part3: listify(subset(Omega,lambda([li],is(li[1]=3))));
Omega_part4: listify(subset(Omega,lambda([li],is(li[1]=4))));
tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
summe_produkt_string: rand([tr("die Summe","the sum"),tr("das Produkt","the product")]);
n1: if summe_produkt_string=tr("die Summe","the sum") then rand_with_step(6,9,1) else rand([6,8,12,24]);
rechenzeichen: if summe_produkt_string=tr("die Summe","the sum") then "+" else "*";
rechenzeichen_symbol: if summe_produkt_string=tr("die Summe","the sum") then "+" else "\\cdot";
A1: setify(sublist(listify(Omega),lambda([Liste],apply(rechenzeichen,Liste)=n1)));
p1: length(A1)/length(Omega);

/*Buggy-Rules:*/
 Omega_f1: map(setify,Omega); /*Ergebnisraum ohne Beachtung der Reihenfolge*/]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@sum@}</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>(</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>Omega</tans>
            <boxsize>193</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>A1</tans>
            <boxsize>50</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>p1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/*
Knoten 1: Ist die Antwort korrekt?
Knoten 2: Ist die Antwort eine nichtleere Menge?
Knoten 3: Ist jedes der Elemente ein Tupel/Liste? 
Knoten 4: Ist die Antwort eine Teilmenge der gesuchten Menge? (d.h. fehlen nur Ergebnisse?) 
Knoten 5: Gibt es mindestens einen Eintrag in einem Tupel, der nicht im Definitionsbereich {1,2,3,4} liegt?
Knoten 6: Haben alle Elemente genau drei Eintr&#228;ge?
Knoten 7: Gibt es mindestens ein Element mit doppelten Eintr&#228;gen?
Knoten 8: Ist jedes der Elemente eine Menge? 
Knoten 9: Ist jedes der Elemente eine Zahl? 
*/

ans1_L: listify(ans1); /*ans1 als Liste (um besser an die Elemente zu kommen)*/

/*Liste mit allen Tupeleintr&#228;gen, die nicht im Definitionsbereich liegen:*/
/* falscheEintraege: sublist(unique(apply(append,ans1_L)),lambda([x],freeof(x,set(1,2,3,4))));*/
/*Wird nirgendwo verwendet*/

/*Liste mit allen Tupeln, die Eintr&#228;ge haben, die nicht im Definitionsbereich liegen:*/
/*TupelFalscheEintraege: sublist(ans1_L,lambda([x],intersection(setify(x),set(1,2,3,4))#setify(x)));*/ 

 /*Liste mit allen Ergebnissen, die keine Tupel sind:*/
/*keineTupel: sublist(ans1_L,lambda([x],listp(x)=false));*/

 /*Liste mit allen Ergebnissen (Tupeln), die nicht drei EIntr&#228;ge haben:*/
/*nichtDreiEintraege: sublist(ans1_L,lambda([x],length(x)#3));*/

/*doppelteEintraege: sublist(ans1_L,lambda([x],length(x)#length(unique(x))));*/

 /*Ein Beispiel f&#252;r ein korrektes Element, was nicht in ans1 enthalten ist:*/
/*fehlendesElement: listify(setdifference(Omega,ans1))[1];*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>Omega</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 1w -->
    <div id="oer-aufgabe-10_prt1_button" style="display:box">
        <p>
            Bei diesem Zufallsexperiment hat der Ergebnisraum genau 24 Elemente. Bei Experimenten, in denen mit Beachtung der Reihenfolge gezogen wird, werden die einzelnen Ergebnisse durch Tupel beschrieben. Da drei Kugeln gezogen werden, sind es Tupel der L&#228;nge 3. Und weil die Kugeln nach dem Ziehen nicht zur&#252;ckgelegt werden, sind in den Tupeln keine doppelten Elemente enthalten. Dies f&#252;hrt zu dem Ergebnisraum
            <br/>
            \(
\begin{equation}
				\begin{split}
				\Omega = &amp;\{ {@tex1(sequenceify(Omega_part1))@}, \\\
								&amp;{@tex1(sequenceify(Omega_part2))@}, \\\
								&amp;{@tex1(sequenceify(Omega_part3))@},\\\
								&amp;{@tex1(sequenceify(Omega_part4))@} \}
				\end{split}
\end{equation}
\)
        </p>
        <p>
            Bitte klicken Sie auf den folgenden Button:
            <button class="buttonweiter" id="oer-aufgabe-10-btn1" onclick="
            hide('oer-aufgabe-10_aufgabe1');
            show('oer-aufgabe-10_aufgabe2');
            show('oer-aufgabe-10_div-bekanntes');" type="button">
                Zum n&#228;chsten Aufgabenteil
            </button>
        </p>
    </div>
    <script>
        (function() {
// Ablaufplan:
const btn = document.getElementById("oer-aufgabe-10-btn1");
btn.addEventListener("click", function(event) {
        QMenu.storeAndUpdate(1);
});




// Einf&#252;gen der bekannten Infos:
const ansDiv = document.querySelector("#oer-aufgabe-10_div-bekanntes");
				//ansDiv.style.display = "block";
				const paragraph = document.createElement("p");
				paragraph.append("Bitte bearbeiten Sie nun den n&#228;chsten Aufgabenteil. Zur Erinnerung:");
				ansDiv.append(paragraph);
				const paragraph2 = document.createElement("p");
				paragraph2.append("\\"+"("+"\\"+"begin{equation}"+ "\\"+"begin{split}"+"\\Omega"+"=&"+"\\"+"{"+{#tex1(sequenceify(Omega_part1))#}+", "+"\\"+"\\"+"\\ "+"&"+{#tex1(sequenceify(Omega_part2))#}+", "+"\\"+"\\"+"\\ "+"&"+{#tex1(sequenceify(Omega_part3))#}+","+"\\"+"\\"+"\\ "+"&"+{#tex1(sequenceify(Omega_part4))#}+"\\"+"}"+"\\"+"end{split}"+"\\"+"end{equation}"+"\\"+")");
				ansDiv.append(paragraph2);
})();
    </script>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 1w -->
    <div id="oer-aufgabe-10_prt1_button" style="display:box">
        <p>
            In this random experiment, the result space has exactly 24 elements. In experiments in which the sequence is observed, the individual results are described by Tupel. Since three balls are drawn, they are tuples of length 3. And because the balls are not covered after drawing, no double elements are included in the tuples. This leads to the result area
            <br/>
            \(
\begin{equation}
				\begin{split}
				\Omega = &amp;\{ {@tex1(sequenceify(Omega_part1))@}, \\\
								&amp;{@tex1(sequenceify(Omega_part2))@}, \\\
								&amp;{@tex1(sequenceify(Omega_part3))@},\\\
								&amp;{@tex1(sequenceify(Omega_part4))@} \}
				\end{split}
\end{equation}
\)
        </p>
        <p>
            Please click the following button:
            <button class="buttonweiter" id="oer-aufgabe-10-btn1" onclick="
            hide('oer-aufgabe-10_aufgabe1');
            show('oer-aufgabe-10_aufgabe2');
            show('oer-aufgabe-10_div-bekanntes');" type="button">
                The next task
            </button>
        </p>
    </div>
    <script>
        (function() {
// Ablaufplan:
const btn = document.getElementById("oer-aufgabe-10-btn1");
btn.addEventListener("click", function(event) {
        QMenu.storeAndUpdate(1);
});




// Einf&#252;gen der bekannten Infos:
const ansDiv = document.querySelector("#oer-aufgabe-10_div-bekanntes");
				//ansDiv.style.display = "block";
				const paragraph = document.createElement("p");
				paragraph.append("Bitte bearbeiten Sie nun den n&#228;chsten Aufgabenteil. Zur Erinnerung:");
				ansDiv.append(paragraph);
				const paragraph2 = document.createElement("p");
				paragraph2.append("\\"+"("+"\\"+"begin{equation}"+ "\\"+"begin{split}"+"\\Omega"+"=&"+"\\"+"{"+{#tex1(sequenceify(Omega_part1))#}+", "+"\\"+"\\"+"\\ "+"&"+{#tex1(sequenceify(Omega_part2))#}+", "+"\\"+"\\"+"\\ "+"&"+{#tex1(sequenceify(Omega_part3))#}+","+"\\"+"\\"+"\\ "+"&"+{#tex1(sequenceify(Omega_part4))#}+"\\"+"}"+"\\"+"end{split}"+"\\"+"end{equation}"+"\\"+")");
				ansDiv.append(paragraph2);
})();
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 1f -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 1f -->
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>setp(ans1) and length(ans1)#0</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 2w -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 2w -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 2f -->
    <p>
        Ein Ergebnisraum ist laut Definition immer eine nichtleere Menge, Ihre Antwort ist es aber nicht.
        <br/>
        Bitte geben Sie eine nichtleere Menge ein. Klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 2f -->
    <p>
        According to definition, a result room is always a non-empty amount, but your answer is not.
        <br/>
        Please enter a non-empty amount. Then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>unique(map(listp,ans1_L))</sans>
                <tans>[true]</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>3</truenextnode>
                <trueanswernote>prt1-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 3w -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 3w -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>7</falsenextnode>
                <falseanswernote>prt1-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 3f - d.h. es gibt mind. ein Element, das kein Tupel/Liste ist. &#220;berpr&#252;fe in Knoten 8, ob es alle Elemente Mengen sind -->
</span>
<span class="multilang" lang="en">
    <!-- Knoten 3f - d.h. es gibt mind. ein Element, das kein Tupel/Liste ist. &#220;berpr&#252;fe in Knoten 8, ob es alle Elemente Mengen sind -->
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>subsetp(ans1,Omega)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 4w -->
    <p>
        Die von Ihnen angegebene Menge ist eine Teilmenge des korrekten Ergebnisraums. Das bedeutet, dass alle Elemente Ihrer Menge richtig sind, aber welche fehlen.
    </p>
    <p>
        Gehen Sie strukturiert an die Sache heran und sortieren Sie die Ergebnisse (z.B. danach, welche Kugel zuerst gezogen wird). So vergessen Sie kein Ergebnis.
        <br/>
        Tipp: Insgesamt sind es 24 Ergebnisse. Eines der fehlenden Elemente lautet {@listify(setdifference(Omega,ans1))[1]@}.
    </p>
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 4w -->
    <p>
        The quantity you provide is a subset of the correct result area. This means that all elements of your amount are correct, but which are missing.
    </p>
    <p>
        Go to the item in a structured manner and sort the results (e.g. after which the ball is pulled first). So do not forget a result.
        <br/>
        Tip: Overall, there are 24 results. One of the missing elements is {@listify(setdifference(Omega,ans1))[1]@}.
    </p>
    <p>
        Please correct your input and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prt1-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 4f - d.h., dass eine Menge von Tupeln angegeben wurde, von denen welche falsch sind. &#220;berpr&#252;fe in Knoten 5, ob alle Tupel nur Eintr&#228;ge aus {1,2,3,4} haben -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 4f - d.h., dass eine Menge von Tupeln angegeben wurde, von denen welche falsch sind. &#220;berpr&#252;fe in Knoten 5, ob alle Tupel nur Eintr&#228;ge aus {1,2,3,4} haben -->
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>unique(map(lambda([x],member(x,set(1,2,3,4))),apply(append,ans1_L)))</sans>
                <tans>[true]</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>5</truenextnode>
                <trueanswernote>prt1-5-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Konten 5w - d.h. es gibt mindestens einen Eintrag in einem Tupel, der nicht im Definitionsbereich {1,2,3,4} liegt -->
</span>
<span class="multilang" lang="en">
    <!-- Konten 5w - d.h. es gibt mindestens einen Eintrag in einem Tupel, der nicht im Definitionsbereich {1,2,3,4} liegt -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 5f - d.h. es gibt keinen einen Eintrag in einem Tupel, der nicht im Definitionsbereich {1,2,3,4} liegt -->
    <p>
        In Ihrer Menge befindet sich mindestens ein Tupel, in dem sich ein Eintrag au&#223;erhalb des Definitionsbereichs befindet. Bei dem betrachteten Zufallsexperiment sind die Kugeln mit den Zahlen von 1 bis 4 beschriftet. Daraus ergibt sich, dass der Definitionsbereich f&#252;r Eintr&#228;ge der Tupel (d.h. der Ergebnisse) \( \{1,2,3,4\} \) ist.
    </p>
    <p>
        Konkret betrifft das die folgenden Tupel: {@sequenceify(sublist(ans1_L,lambda([x],intersection(setify(x),set(1,2,3,4))#setify(x))))@}
    </p>
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!-- Knoten 5f - d.h. es gibt keinen einen Eintrag in einem Tupel, der nicht im Definitionsbereich {1,2,3,4} liegt -->
    <p>
        In your quantity there is at least one tuple in which an entry is located outside the definition range. In the random experiment considered, the balls are labeled with the numbers from 1 to 4. It follows that the definition range for entries of the tuple (i.e. the results) \( \{1,2,3,4\} \) is.
    </p>
    <p>
        Specifically, this concerns the following tuples: {@sequenceify(sublist(ans1_L,lambda([x],intersection(setify(x),set(1,2,3,4))#setify(x))))@}
    </p>
    <p>
        Please correct your input and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>5</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>unique(map(length,ans1_L))</sans>
                <tans>[3]</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>6</truenextnode>
                <trueanswernote>prt1-6-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 6w - d.h., dass eine Liste mit Tupeln angegeben wurde, von denen alle L&#228;nge 3 haben, aber mindestens eines ist falsch (doppelter oder nicht definierter Eintrag) -->
</span>
<span class="multilang" lang="en">
    <!-- Knoten 6w - d.h., dass eine Liste mit Tupeln angegeben wurde, von denen alle L&#228;nge 3 haben, aber mindestens eines ist falsch (doppelter oder nicht definierter Eintrag) -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-6-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 6f -->
    <p>
        In Ihrer Menge befindet sich mindestens ein Tupel, das nicht drei Eintr&#228;ge hat, n&#228;mlich: {@sequenceify(sublist(ans1_L,lambda([x],length(x)#3)))@}
        <br/>
        Da aber ein Zufallsexperiment modelliert werden soll, bei dem drei Kugeln gezogen werden, werden alle Ergebnisse durch Tupel mit drei Eintr&#228;gen (z.B.
        <code>
            [1,2,3]
        </code>
        ) modelliert.
    </p>
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 6f -->
    <p>
        In your quantity there is at least one tuple that does not have three entries, namely: {@sequenceify(sublist(ans1_L,lambda([x],length(x)#3)))@}
        <br/>
        However, since a random experiment is to be modelled, in which three balls are drawn, all results are obtained by tuple with three entries (e.g.
        <code>
            [1,2,3]
        </code>
        ) modelled.
    </p>
    <p>
        Please correct your input and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>6</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>length(sublist(ans1_L,lambda([L],length(L)#length(unique(L)))))</sans>
                <tans>0</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-7-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 6w - d.h. es gibt kein Element mit doppelten Eintr&#228;gen -->
</span>
<span class="multilang" lang="en">
    <!-- Knoten 6w - d.h. es gibt kein Element mit doppelten Eintr&#228;gen -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-7-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 6f - d.h. es gibt mindestens ein Element mit doppelten Eintr&#228;gen -->
    <p>
        In Ihrer Menge befindet sich mindestens ein Tupel, in dem sich Eintr&#228;ge doppeln, n&#228;mlich: {@sequenceify(sublist(ans1_L,lambda([x],length(x)#length(unique(x)))))@}
        <br/>
        Da aber ein Zufallsexperiment modelliert werden soll, bei dem die Kugeln nach dem Ziehen
        <i>
            nicht
        </i>
        zur&#252;ckgelegt werden, werden alle Ergebnisse durch Tupel ohne doppelte Eintr&#228;ge (z.B.
        <code>
            [1,2,3]
        </code>
        ) modelliert.
    </p>
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!-- Knoten 6f - d.h. es gibt mindestens ein Element mit doppelten Eintr&#228;gen -->
    <p>
        In your quantity there is at least one tuple in which entries double, namely: {@sequenceify(sublist(ans1_L,lambda([x],length(x)#length(unique(x)))))@}
        <br/>
        However, since a random experiment is to be modeled, in which the balls are to be modeled after drawing
        <i>
            not
        </i>
        all results will be completed by Tupel without duplicate entries (e.g.
        <code>
            [1,2,3]
        </code>
        ) modelled.
    </p>
    <p>
        Please correct your input and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>7</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>unique(map(setp,ans1_L))</sans>
                <tans>[true]</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-8-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 8w -->
    <p>
        Ihre Menge enth&#228;lt als Elemente Mengen. Da aber in Mengen die Reihenfolge der Elemente nicht beachtet wird, modelliert man bei Zufallsexperimenten, in denen die Reihenfolge beachtet werden soll, die Ergebnisse nicht durch Mengen, sondern durch Tupel.
    </p>
    <p>
        Bitte geben Sie eine Menge an, deren Elemente Tupel (z.B.
        <code>
            [1,2,3]
        </code>
        ) sind. Klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 8w -->
    <p>
        Their quantity contains quantities as elements. However, since the order of the elements is not observed in quantities, the results are modeled by random experiments in which the order is to be observed, not by quantities, but by tuple.
    </p>
    <p>
        Please specify a quantity whose elements Tupel (e.g.
        <code>
            [1,2,3]
        </code>
        ) are. Then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>8</falsenextnode>
                <falseanswernote>prt1-8-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 8f - d.h. es gibt mind. ein Element, das kein Tupel/Liste ist. Es sind aber nicht ausschlie&#223;lich Mengen genannt worden. &#220;berpr&#252;fe in Knoten 9, ob alle Elemente Zahlen sind -->
</span>
<span class="multilang" lang="en">
    <!-- Knoten 8f - d.h. es gibt mind. ein Element, das kein Tupel/Liste ist. Es sind aber nicht ausschlie&#223;lich Mengen genannt worden. &#220;berpr&#252;fe in Knoten 9, ob alle Elemente Zahlen sind -->
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>8</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>unique(map(numberp,ans1_L))</sans>
                <tans>[true]</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-9-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 9w -->
    <p>
        Ihre Menge enth&#228;lt als Elemente Zahlen. Ein Ergebnisraum enth&#228;lt laut Definition alle m&#246;glichen Ergebnisse eines Zufallsexperiments. Wir modellieren hier ein Zufallsexperiment, bei dem wir
        <i>
            ohne Zur&#252;cklegen
        </i>
        Kugeln ziehen und die
        <i>
            Reihenfolge beachten
        </i>
        .
Dies modelliert man durch Tupel und nicht durch Zahlen.
    </p>
    <p>
        Bitte geben Sie eine Menge an, deren Elemente alle Ergebnisse des in der Aufgabenstellung genannten Zufallsexperiments sind und modellieren Sie die Ergebnisse durch Tupel. Klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 9w -->
    <p>
        Their quantity contains numbers as elements. According to definition, a result room contains all possible results of a random experiment. We are modeling a random experiment in which we
        <i>
            without resetting
        </i>
        Pulling balls and the
        <i>
            Note order
        </i>
        .
This is modeled by tuple and not by numbers.
    </p>
    <p>
        Please specify a quantity whose elements are all the results of the random experiment mentioned in the task and model the results by Tupel. Then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-9-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 9f -->
    <p>
        Ein Ergebnisraum enth&#228;lt laut Definition alle m&#246;glichen Ergebnisse eines Zufallsexperiments. Ihre Menge enth&#228;lt aber Elemente, die - zumindest f&#252;r dieses konkreten Beispiel - die Ergebnisse nicht korrekt modellieren. Konkret betrifft das folgende Elemente: {@sequenceify(sublist(ans1_L,lambda([x],listp(x)=false)))@}
    </p>
    <p>
        Bitte geben Sie eine Menge an, deren Elemente alle Ergebnisse des in der Aufgabenstellung genannten Zufallsexperiments sind und verwenden Sie eine geeignete Modellierung. Klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 9f -->
    <p>
        According to definition, a result room contains all possible results of a random experiment. Their quantity contains elements which, at least for this specific example, do not correctly model the results. Specifically, the following elements apply: {@sequenceify(sublist(ans1_L,lambda([x],listp(x)=false)))@}
    </p>
    <p>
        Please specify a quantity whose elements are all the results of the random experiment mentioned in the task and use a suitable modeling. Then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/*
Knoten 1: Ist die Antwort korrekt?
Knoten 2: Ist die Antwort eine Menge?
Knoten 3: Ist die Menge leer?
Knoten 4: Ist ans2 Teilmenge von Omega?
Knoten 5: Ist ans2 nur unvollst&#228;ndig oder enth&#228;lt es falsche Elemente
*/

ans2_L: listify(ans2);

/*Liste mit allen Elementen, die nicht in Omega liegen:*/
/*falscheElemente2: listify(setdifference(ans2,Omega));*/

 /*Liste mit allen Tupeln, die nicht die geforderte Summen-/Produkt-Eigenschaft erf&#252;llen*/
/*falscheSumme2: sublist(listify(ans2_L),lambda([Liste],apply(rechenzeichen,Liste)#n1));*/

 /*Ein Beispiel f&#252;r ein korrektes Element, was nicht in ans2 enthalten ist*/
/*fehlendesElement2: listify(setdifference(A1,ans2))[1];*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans2</sans>
                <tans>A1</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 1w-->
    <div id="oer-aufgabe-10_prt2_button" style="display:box">
        <p>
            F&#252;r Elemente \((\omega_1,\omega_2,\omega_3)\in\Omega\) gibt es nur eine begrenzte Anzahl an M&#246;glichkeiten, sodass \(\omega_1\,{@rechenzeichen_symbol@}\,\omega_2\,{@rechenzeichen_symbol@}\,\omega_3 = {@n1@}\) gilt. Es gibt sogar nur eine m&#246;gliche Kombination von Kugeln, auf die diese Eigenschaft zutrifft, n&#228;mlich {@sequenceify(ans2_L[1])@}. Da bei diesem Zufallsexperiment aber die Reihenfolge, in der die Kugeln gezogen werden, beachtet wird, ist in der Ereignismenge \(A\) je ein Tupel f&#252;r jede m&#246;gliche Permutation dieser drei Zahlen enthalten. Daraus ergibt sich die L&#246;sung \(A={@A1@}\)
        </p>
        <p>
            Bitte klicken Sie auf den folgenden Button:
            <button class="buttonweiter" id="oer-aufgabe-10-btn2" onclick="
            hide('oer-aufgabe-10_aufgabe2');
            show('oer-aufgabe-10_aufgabe3');
            hide('oer-aufgabe-10_Hinweis_Eingabe');" type="button">
                Zum n&#228;chsten Aufgabenteil
            </button>
        </p>
    </div>
    <script>
        // Ablaufplan:
(function() {
const btn = document.getElementById("oer-aufgabe-10-btn2");
btn.addEventListener("click", function(event) {
        QMenu.storeAndUpdate(2);
});




// Einf&#252;gen der bekannten Infos:
const ansDiv = document.querySelector("#oer-aufgabe-10_div-bekanntes");
				//ansDiv.style.display = "block";
				const paragraph = document.createElement("p");
				paragraph.append("und "+"\\"+"("+"\\"+"small{A="+{#tex1(ans2)#}+"}"+"\\"+")");
				ansDiv.append(paragraph);
})();
    </script>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 1w-->
    <div id="oer-aufgabe-10_prt2_button" style="display:box">
        <p>
            For elements \((\omega_1,\omega_2,\omega_3)\in\Omega\) there are only a limited number of possibilities, so \(\omega_1\,{@rechenzeichen_symbol@}\,\omega_2\,{@rechenzeichen_symbol@}\,\omega_3 = {@n1@}\) Article 2 There is even only one possible combination of balls to which this property applies, namely {@sequenceify(ans2_L[1])@}. Since in this random experiment, however, the order in which the balls are drawn is observed, the event quantity is \(A\) each contain a tuple for each possible permutation of these three numbers. This results in the solution \(A={@A1@}\)
        </p>
        <p>
            Please click the following button:
            <button class="buttonweiter" id="oer-aufgabe-10-btn2" onclick="
            hide('oer-aufgabe-10_aufgabe2');
            show('oer-aufgabe-10_aufgabe3');
            hide('oer-aufgabe-10_Hinweis_Eingabe');" type="button">
                The next task
            </button>
        </p>
    </div>
    <script>
        // Ablaufplan:
(function() {
const btn = document.getElementById("oer-aufgabe-10-btn2");
btn.addEventListener("click", function(event) {
        QMenu.storeAndUpdate(2);
});




// Einf&#252;gen der bekannten Infos:
const ansDiv = document.querySelector("#oer-aufgabe-10_div-bekanntes");
				//ansDiv.style.display = "block";
				const paragraph = document.createElement("p");
				paragraph.append("und "+"\\"+"("+"\\"+"small{A="+{#tex1(ans2)#}+"}"+"\\"+")");
				ansDiv.append(paragraph);
})();
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 1f -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 1f -->
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>setp(ans2)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>2</truenextnode>
                <trueanswernote>prt2-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 2w - Antwort ist eine Menge. -->
</span>
<span class="multilang" lang="en">
    <!-- Knoten 2w - Antwort ist eine Menge. -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 2f - Antwort ist keine Menge. -->
    <p>
        Sie sollten das Ereignis als Menge angeben, Ihre Antwort ist aber keine.
        <br/>
        Bitte geben Sie eine Menge ein. Klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!-- Knoten 2f - Antwort ist keine Menge. -->
    <p>
        You should specify the event as a quantity, but your answer is not.
        <br/>
        Please enter a lot. Then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>emptyp(ans2)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 3w - Es wurde eine leere Menge angegeben-->
    <p>
        Sie haben die leere Menge angegeben. Es gibt aber Ergebnisse, f&#252;r die gilt, dass {@summe_produkt_string@} der Zahlen, die auf den gezogenen Kugeln stehen, {@n1@} ist.  Ein Beispiel daf&#252;r ist {@listify(A1)[1]@}.
    </p>
    <p>
        Bitte geben Sie eine nichtleere Menge ein und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!-- Knoten 3w - Es wurde eine leere Menge angegeben-->
    <p>
        You have given the empty quantity. However, there are results for which: {@summe_produkt_string@} the numbers standing on the drawn balls, {@n1@} is. An example is {@listify(A1)[1]@}.
    </p>
    <p>
        Please enter a non-empty quantity and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt2-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 3f - Es wurde keine leere Menge angegeben-->
</span>
<span class="multilang" lang="en">
    <!-- Knoten 3f - Es wurde keine leere Menge angegeben-->
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>subsetp(ans2,Omega)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prt2-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 4w - ans2 ist Teilmenge von Omega. &#220;berpr&#252;fe in Knoten 5, ob falsche Elemente dabei sind oder ob die Antwort nur unvollst&#228;ndig ist. -->
</span>
<span class="multilang" lang="en">
    <!-- Knoten 4w - ans2 ist Teilmenge von Omega. &#220;berpr&#252;fe in Knoten 5, ob falsche Elemente dabei sind oder ob die Antwort nur unvollst&#228;ndig ist. -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!-- Knoten 4f - ans2 ist KEINE Teilmenge von ans1 -->
    <p>
        Ihre Antwort ist keine Teilmenge Ihres Ergebnisraums \(\Omega\), das hei&#223;t in Ihrer Menge sind Elemente zu finden, die keine m&#246;glichen Ergebnisse des Zufallsexperiments sind.
    </p>
    <p>
        Konkret betrifft das die folgenden Elemente: {@sequenceify( listify(setdifference(ans2,Omega)) )@}
    </p>
    <p>
        Bitte geben Sie eine Teilmenge von \(\Omega\) an und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!-- Knoten 4f - ans2 ist KEINE Teilmenge von ans1 -->
    <p>
        Your answer is not a subset of your earnings room \(\Omega\), i.e. in your quantity you can find elements that are not possible results of chance experiment.
    </p>
    <p>
        Specifically, this concerns the following elements: {@sequenceify( listify(setdifference(ans2,Omega)) )@}
    </p>
    <p>
        Please give a subset of \(\Omega\) and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans><![CDATA[is(length(setdifference(ans2,A1))>0)]]></sans>
                <tans>false</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-5-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 5w -->
    <p>
        Ihre Antwort ist eine Teilmenge des Ereignisses \(A\), aber unvollst&#228;ndig. Eines der fehlenden Elemente lautet {@listify(setdifference(A1,ans2))[1]@}.
    </p>
    <p>
        Bitte erg&#228;nzen Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 5w -->
    <p>
        Your answer is a part of the event \(A\)but incomplete. One of the missing elements is {@listify(setdifference(A1,ans2))[1]@}.
    </p>
    <p>
        Please add your input and then click "Check".
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 5f -->
    <p>
        Ihre Antwort ist eine Teilmenge des Ergebnisraums \(\Omega\). Allerdings enth&#228;lt sie Ergebnisse, die nicht die geforderte Eigenschaft erf&#252;llen, und somit keine Elemente von \(A\) sind. Dies betrifft die folgenden Elemente: {@sequenceify(sublist(listify(ans2_L),lambda([Liste],apply(rechenzeichen,Liste)#n1)))@}
    </p>
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 5f -->
    <p>
        Your answer is a subset of the result area \(\Omega\). However, it contains results that do not meet the required property and thus do not contain any elements of \(A\) are. This relates to the following elements: {@sequenceify(sublist(listify(ans2_L),lambda([Liste],apply(rechenzeichen,Liste)#n1)))@}
    </p>
    <p>
        Please correct your input and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>/*
Knoten 1: Ist die Antwort eine Zahl?
Knoten 2: Ist die Antwort korrekt?
Knoten 3: Ist die Antwort die Gegenwahrscheinlichkeit der korrekten Antwort?
Knoten 4: Ist die Antwort 0&#8804;p&#8804;1?
Knoten 5: Wurde #Omega/#A gerechnet?
*/</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>numberp(ans3)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 1w -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 1w -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 1f -->
    <p>
        Ihr Wert ist keine Zahl. Wahrscheinlichkeiten sind definitionsgem&#228;&#223; aber immer eine Zahl.
    </p>
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 1f -->
    <p>
        Their value is not a number. Probabilities are by definition always a number.
    </p>
    <p>
        Please correct your input and then click "Check" again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>p1</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 2w -->
    <p>
        Sind ein Ergebnisraum \(\Omega\) und ein Ereignis \(A\subseteq\Omega\) bekannt, dann kann die Wahrscheinlichkeit des Ereignisses mit der Laplace-Formel \(P(A)=\frac{\#A}{\#\Omega}\) berechnet werden (wobei "\(\#\)" die Anzahl der Elemente bezeichnet). In unserem Fall gilt \(\#A = {@length(A1)@}\) sowie \(\#\Omega = {@length(Omega)@}\), also gilt \(P(A)=\frac{{@length(A1)@}}{{@length(Omega)@}}={@ans3@}\).
    </p>
    <p>
        &#220;brigens: Um die Laplace-Formel anwenden zu k&#246;nnen, sind gar keine genauen Kenntnisse &#252;ber die Elemente von \(\Omega\) oder \(A\) notwendig. Wichtig ist nur, dass jeweils die Anzahl der Elemente bekannt ist. Aus diesem Grund gibt es eine Reihe von Formeln, mit denen man diese bestimmen kann, ohne jedes einzelne Element aufschreiben zu m&#252;ssen.
    </p>
    <script>
        window.onload= function(){
            document.getElementById('oer-aufgabe-10_Hinweis_Eingabe').style.display="none";
            document.getElementById('oer-aufgabe-10_Hinweis_Annahme').style.display="none";
            document.getElementById('oer-aufgabe-10_Hinweis_Pruefen').style.display="none";
            document.getElementById('oer-aufgabe-10_div-bekanntes').style.display="none";
};
    </script>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 2w -->
    <p>
        Are the result room \(\Omega\) and an event \(A\subseteq\Omega\) known, then the probability of the event with the Laplace formula \(P(A)=\frac{\#A}{\#\Omega}\) are calculated (where "\(\#\)" denotes the number of elements). In our case \(\#A = {@length(A1)@}\) and \(\#\Omega = {@length(Omega)@}\), so apply \(P(A)=\frac{{@length(A1)@}}{{@length(Omega)@}}={@ans3@}\).
    </p>
    <p>
        By the way: In order to be able to apply the Laplace formula, no precise knowledge of the elements of \(\Omega\) or \(A\) necessary. It is only important that the number of elements is known in each case. For this reason, there are a number of formulas with which you can determine them without having to write down each individual element.
    </p>
    <script>
        window.onload= function(){
            document.getElementById('oer-aufgabe-10_Hinweis_Eingabe').style.display="none";
            document.getElementById('oer-aufgabe-10_Hinweis_Annahme').style.display="none";
            document.getElementById('oer-aufgabe-10_Hinweis_Pruefen').style.display="none";
            document.getElementById('oer-aufgabe-10_div-bekanntes').style.display="none";
};
    </script>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>2</falsenextnode>
                <falseanswernote>prt3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 2f -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 2f -->
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>1-p1</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-3-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 3w -->
    <p>
        Sie haben die Gegenwahrscheinlichkeit der gesuchten Wahrscheinlichkeit angegeben.
    </p>
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 3w -->
    <p>
        They have indicated the counter probability of the desired probability.
    </p>
    <p>
        Please correct your input and then click "Check" again.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>3</falsenextnode>
                <falseanswernote>prt3-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 3f -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 3f -->
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>3</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans><![CDATA[is(ans3>=0) and is(ans3<=1)]]></sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>4</truenextnode>
                <trueanswernote>prt3-4-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 4w -->
</span>
<span class="multilang" lang="en">
    <!--  Knoten 4w -->
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>4</falsenextnode>
                <falseanswernote>prt3-4-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 4f -->
    <p>
        Ihre Antwort liegt nicht im Intervall \([0,1]\). Wahrscheinlichkeiten nehmen definitionsgem&#228;&#223; aber immer einen Wert zwischen 0 und 1 an.
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 4f -->
    <p>
        Your answer is not in interval \([0,1]\). Probabilities, however, always assume a value between 0 and 1.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>4</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans3</sans>
                <tans>length(Omega)/length(A1)</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-5-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 5w -->
    <p>
        Haben Sie vielleicht die Anzahl aller Ergebnisse durch die Anzahl der Ergebnisse des Ereignisses geteilt (anstatt umgekehrt)?
    </p>
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 5w -->
    <p>
        Have you perhaps shared the number of results by the number of results of the event (instead of vice versa)?
    </p>
    <p>
        Please correct your input and then click "Check" again.
        <br/>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-5-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <!--  Knoten 4f -->
    <p>
        Bitte korrigieren Sie Ihre Eingabe und klicken Sie dann erneut auf "Pr&#252;fen".
    </p>
    <p>
        Tipp: Verwenden Sie die Laplace-Verteilung
    </p>
</span>
<span class="multilang" lang="en">
    <!--  Knoten 4f -->
    <p>
        Please correct your input and then click "Check" again.
    </p>
    <p>
        Tip: Use the Laplace distribution
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>Omega</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>A1</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>p1</value>
            </testinput>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>