<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Passagiere auf einem Transatlantikflug</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <style>
        button.top-next,
button.next {
                border: 1px solid black;
                background-color: #A9F5BC;
}




button.top-next:hover,
button.next:hover {
                background-color: #E0F8E6;
}




button.top,
button.prev {
                border: 1px solid black;
                background-color: #DCB25E;
}
button.top:hover,
button.prev:hover {
                background-color: #E7D2A7;
}
    </style>
    <div class="question-view">
        <p>
            Die Lufthansa wei&#223; aus Erfahrung, dass Passagiere mit einer
Wahrscheinlichkeit von {# pp * 100 #}% zum Flug erscheinen. F&#252;r einen
Transatlantikflug hat sie {# nn #} Tickets verkauft. Wir bezeichnen
mit \(X\) die Anzahl der Passagiere, die zum Flug erscheinen.
        </p>
        <div class="question-view">
            <div class="question-view">
                <p>
                    Bestimmen Sie n&#228;herungsweise die Wahrscheinlichkeit, dass nicht
mehr als {# max_missing #} Passagiere zum Abflug erscheinen. Verwenden
Sie dazu die Normalapproximation und geben Sie das Ergebnis auf
mindestens zwei Nachkommastellen genau an.
                </p>
                Antwort: [[input:ans4]][[validation:ans4]]
[[feedback:ans4]]
            </div>
            <div class="question-view">
                <p>
                    Benennen Sie die Verteilung der Zufallsvariablen \(X\).
                </p>
                [[input:ans1]]
[[validation:ans1]]
[[feedback:prt1]]
            </div>
            <div class="question-view">
                <p>
                    Im vorigen Schritt haben Sie gelernt, dass \(X \sim {@Bin(nn, pp)@}\) gilt.
                </p>
                <p>
                    Geben Sie den Erwartungswert und die Varianz von \(X\) an.
                </p>
                <div>
                    Erwartungswert: [[input:ans2]] [[validation:ans2]]
                </div>
                [[feedback:prt2]]
                <div>
                    Varianz: [[input:ans3]] [[validation:ans3]]
                </div>
                [[feedback:prt3]]
                [[feedback:teil2_vollstaendig]]
            </div>
            <div class="question-view">
                <p>
                    Im vorigen Schritt haben Sie {@E(X) = mean@} und {@Var(X) = variance@} bestimmt.
                </p>
                <p>
                    Definieren Sie eine neue Zufallsvariable \(Z\) durch eine
geeigneten Transformation von \(X\), sodass \(Z\) approximativ
standard-normalverteilt ist.
                </p>
                \(Z = \) [[input:hilfe3]]
[[validation:hilfe3]]
[[feedback:hilfe3]]
            </div>
            <div class="question-view">
                <p>
                    Im vorigen Schritt haben Sie die neue Zufallsvariable {@Z = (X - mean)/sqrt(variance)@} definiert.
                </p>
                <p>
                    Dr&#252;cken Sie das Ereignis {@ X &lt;= max_missing @} mit Hilfe der
        Zufallsvariablen \(Z\) aus.
                </p>
                <p>
                    \({@ {X &lt;= max_missing} @} = \) [[input:hilfe4]]
                </p>
                [[validation:hilfe4]]
        [[feedback:hilfe4]]
            </div>
        </div>
        <script>
            /*
    * This program is free software: you can redistribute it and/or
    * modify it under the terms of the GNU General Public License as
    * published by the Free Software Foundation, either version 3 of the
    * License, or (at your option) any later version.
    *
    * This program is distributed in the hope that it will be useful, but
    * WITHOUT ANY WARRANTY; without even the implied warranty of
    * MERCHANTABILITY or FITNESS FOR A PARTICULAR PURPOSE.  See the GNU
    * General Public License for more details.
    *
    * You should have received a copy of the GNU General Public License
    * along with this program.  If not, see <https://www.gnu.org/licenses/>.
    *
    * Copyright (C) 2021 by Daniel Mei&#223;ner <daniel.meissner-i4k@ruhr-uni-bochum.de>
    * SPDX-License-Identifier: GPL-3.0-or-later
    *
    */




/* @license magnet:?xt=urn:btih:1f739d935676111cfff4b4693e3816e664797050&dn=gpl-3.0.txt GPL-v3-or-Later */




(function() {
/* Make sure to create QController instances only once even if the
            script is loaded multiple times.  This is so that the script can
            be included in more than one STACK question on the page and still
            function properly. */
if (window.QController === undefined) {
        window.QController = QController;
        document.addEventListener("DOMContentLoaded", event => {
                const responseform = document.getElementById("responseform");
                if (!responseform) {
                        throw Error("Element with id responseform not found.");
                }

    
    
    
    responseform.querySelectorAll(".stack.que").forEach(stackQue => {
                        const rootElement = stackQue.querySelector(".question-view");

      
      
      
      if (!rootElement) {
                                console.warn("QController: no elements with class question-view inside " + stackQue.id);
                                return;
                        }

      
      
      
      const controller = new QController(rootElement, stackQue.id);
                        controller.installHandlers();
                        controller.loadState();
                        controller.update();
                });
        });
}




/* This class handles the navigation of a single STACK question */
function QController(rootElement, questionId) {
        this.rootElement = rootElement;
        this.id = questionId;
        this.views = rootElement.querySelectorAll(".question-view");
        this.visited = (new Array(this.views.length)).fill(false);
        this.currentView = Array.prototype.findIndex.call(this.views, isLeaf);
        this.visited[this.currentView] = true;
}




QController.prototype.handleVisitedNodes = function(current) {
        var visitedSiblings = [];
        var el = current;
        while (el.nextElementSibling) {
                var index =
                                Array.prototype.indexOf.call(this.views, el.nextElementSibling);
                visitedSiblings.push(this.visited[index])
                el = el.nextElementSibling;
        }

  
  
  
  var withoutHelp = visitedSiblings.every(x => x == false);
        var helpCompleted = visitedSiblings.every(x => x == true);

  
  
  
  current.querySelectorAll(".visited").forEach(function(element) {
                var index = element.getAttribute("data-qcontroller-help");
                if (!index || index === "none")
                        element.hidden = !withoutHelp;
                else if (index === "complete")
                        element.hidden = !helpCompleted;
                else if (index === "partial")
                        element.hidden = withoutHelp || helpCompleted;
                else
                        element.hidden = !visitedSiblings[index - 1];
        });
};




/* Update view according to current controller state */
QController.prototype.update = function() {
        this.views.forEach(hide);

  
  
  
  const current = this.views[this.currentView];
        if (isLeaf(current)) {
                showView(current);
                this.handleVisitedNodes(current);
        } else {
                this.showFirstView(current);
        }
};




/* Install event handlers for all buttons under root element */
QController.prototype.installHandlers = function() {
        const controller = this;
        this.rootElement.querySelectorAll("button").forEach(button => {
                button.addEventListener("click", ev => {
                        const view = button.closest(".question-view");

      
      
      
      if (!view) {
                                console.log("qcontroller: <button> not inside question view");
                                return;
                        }

      
      
      
      const classList = button.classList;
                        let forward = null;

      
      
      
      if (classList.contains("next")) {
                                forward = view.nextElementSibling;
                        } else if (classList.contains("top-next")) {
                                const parent = view.parentNode;
                                if (parent) {
                                        forward = parent.nextElementSibling;
                                } else {
                                        forward = null;
                                }
                        } else if (classList.contains("top")) {
                                forward = containingView(view);
                        } else if (classList.contains("prev")) {
                                forward = view.previousElementSibling;
                        }

      
      
      
      if (forward && forward.classList.contains("question-view")) {
                                controller.visit(forward);
                                controller.saveState();
                                controller.update();
                        }
                });
        });
};




QController.prototype.hasVisited = function(viewIndex) {
        return this.visited[viewIndex];
};




QController.prototype.visit = function(nextView) {
        const index = Array.prototype.findIndex.call(
                this.views,
                x => x === nextView);
        this.currentView = index;
        this.visited[index] = true;
};




/* Save controller state to current session */
QController.prototype.saveState = function() {
        const key = this.id + "-controller-state";
        const state = {
                currentView: this.currentView,
                visited: this.visited
        };
        window.sessionStorage.setItem(key, JSON.stringify(state));
};




/* Load controller state from current session */
QController.prototype.loadState = function() {
        const key = this.id + "-controller-state";
        const state = JSON.parse(window.sessionStorage.getItem(key));

  
  
  
  if (state) {
                this.currentView = state.currentView;
                this.visited = state.visited;
        }
};




/* Show the first leaf question view */
QController.prototype.showFirstView = function(questionView) {
        let view = questionView;

  
  
  
  while (view && !isLeaf(view)) {
                view = view.querySelector(".question-view");
        }

  
  
  
  showView(view);
        this.handleVisitedNodes(view);
}




function showView(view) {
        /* show all containing views */
        while (view) {
                show(view);
                view = containingView(view);
        }
}




/* Return the view containing VIEW */
function containingView(view) {
        let parent = view.parentNode;

  
  
  
  if (parent) {
                parent = parent.closest(".question-view");
        }

  
  
  
  return parent;
}




function show(view) {
        view.style.display = "block";
}




function isLeaf(view) {
        return view.querySelector(".question-view") == null;
}




function isVisible(view) {
        return view.style.display != "none";
}




function hide(view) {
        view.style.display = "none";
}




})();
/* @license-end */
        </script>
    </div>
</span>
<span class="multilang" lang="en">
    <style>
        button.top-next,
button.next {
                border: 1px solid black;
                background-color: #A9F5BC;
}




button.top-next:hover,
button.next:hover {
                background-color: #E0F8E6;
}




button.top,
button.prev {
                border: 1px solid black;
                background-color: #DCB25E;
}
button.top:hover,
button.prev:hover {
                background-color: #E7D2A7;
}
    </style>
    <div class="question-view">
        <p>
            Lufthansa knows from experience that passengers with one
Probability {# pp * 100 #}% to the flight. For one
Transatlantic flight has it {# nn #} Sold tickets. We call
with \(X\) the number of passengers coming to the flight.
        </p>
        <div class="question-view">
            <div class="question-view">
                <p>
                    approximately determine the probability that not
more than {# max_missing #} Passengers appearing for departure. Use
To do this, give the normal approximation and give the result
at least two descendants exactly.
                </p>
                Answer: [[input:ans4]][[validation:ans4]][[feedback:ans4]]
            </div>
            <div class="question-view">
                <p>
                    Name the distribution of random variables \(X\).
                </p>
                [[input:ans1]][[validation:ans1]][[feedback:prt1]]
            </div>
            <div class="question-view">
                <p>
                    In the previous step you have learned that \(X \sim {@Bin(nn, pp)@}\) Article 2
                </p>
                <p>
                    Specify the expected value and the variance of \(X\) on.
                </p>
                <div>
                    Expected value: [[input:ans2]] [[validation:ans2]]
                </div>
                [[feedback:prt2]]
                <div>
                    Variance: [[input:ans3]] [[validation:ans3]]
                </div>
                [[feedback:prt3]][[feedback:teil2_vollstaendig]]
            </div>
            <div class="question-view">
                <p>
                    In the previous step you have {@E(X) = mean@} and {@Var(X) = variance@} destined.
                </p>
                <p>
                    Define a new random variable \(Z\) by
suitable transformation of \(X\)so that \(Z\) approximative
standard-normally distributed.
                </p>
                \(Z = \) [[input:hilfe3]][[validation:hilfe3]][[feedback:hilfe3]]
            </div>
            <div class="question-view">
                <p>
                    In the previous step you have the new random variable {@Z = (X - mean)/sqrt(variance)@} defined.
                </p>
                <p>
                    Press the event {@ X &lt;= max_missing @} with the help of
Random variable \(Z\) out.
                </p>
                <p>
                    \({@ {X &lt;= max_missing} @} = \) [[input:hilfe4]]
                </p>
                [[validation:hilfe4]][[feedback:hilfe4]]
            </div>
        </div>
        <script>
            /*
    * This program is free software: you can redistribute it and/or
    * modify it under the terms of the GNU General Public License as
    * published by the Free Software Foundation, either version 3 of the
    * License, or (at your option) any later version.
    *
    * This program is distributed in the hope that it will be useful, but
    * WITHOUT ANY WARRANTY; without even the implied warranty of
    * MERCHANTABILITY or FITNESS FOR A PARTICULAR PURPOSE.  See the GNU
    * General Public License for more details.
    *
    * You should have received a copy of the GNU General Public License
    * along with this program.  If not, see <https://www.gnu.org/licenses/>.
*
* Copyright (C) 2021 by Daniel Mei&#223;ner <daniel.meissner-i4k@ruhr-uni-bochum.de>* SPDX-License Identifier: GPL-3.0-or-later
*
*/




/* @license magnet:?xt=urn:btih:1f739d935676111cfff4b4693e3816e664797050&dn=gpl-3.0.txt GPL-v3-orLater */




(function() {
/* Make sure to create QController instances only once even if the
            script is loaded multiple times.  This is so that the script can
            be included in more than one STACK question on the page and still
            function properly. */
if (window.QController === undefined) {
        window.QController = QController;
        document.addEventListener("DOMContentLoaded", event => {
                const responseform = document.getElementById("responseform");
                if (!responseform) {
                        throw Error("Element with id responseform not found.");
                }

    
    
    
    responseform.querySelectorAll(".stack.que").forEach(stackQue => {
                        const rootElement = stackQue.querySelector(".question-view");

      
      
      
      if (!rootElement) {
                                console.warn("QController: no elements with class question-view inside " + stackQue.id);
                                return;
                        }

      
      
      
      const controller = new QController(rootElement, stackQue.id);
                        controller.installHandlers();
                        controller.loadState();
                        controller.update();
                });
        });
}




/* This class handles the navigation of a single STACK question */
function QController(rootElement, questionId) {
        this.rootElement = rootElement;
        this.id = questionId;
        this.views = rootElement.querySelectorAll(".question-view");
        this.visited = (new Array(this.views.length)).fill(false);
        this.currentView = Array.prototype.findIndex.call(this.views, isLeaf);
        this.visited[this.currentView] = true;
}




QController.prototype.handleVisitedNodes = function(current) {
        var visitedSiblings = [];
        var el = current;
        while (el.nextElementSibling) {
                var index =
                                Array.prototype.indexOf.call(this.views, el.nextElementSibling);
                visitedSiblings.push(this.visited[index])
                el = el.nextElementSibling;
        }

  
  
  
  var withoutHelp = visitedSiblings.every(x => x == false);
        var helpCompleted = visitedSiblings.every(x => x == true);

  
  
  
  current.querySelectorAll(".visited").forEach(function(element) {
                var index = element.getAttribute("data-qcontroller-help");
                if (!index || index === "none")
                        element.hidden = !withoutHelp;
                else if (index === "complete")
                        element.hidden = !helpCompleted;
                else if (index === "partial")
                        element.hidden = withoutHelp || helpCompleted;
                else
                        element.hidden = !visitedSiblings[index - 1];
        });
};




/* Update view according to current controller state */
QController.prototype.update = function() {
        this.views.forEach(hide);

  
  
  
  const current = this.views[this.currentView];
        if (isLeaf(current)) {
                showView(current);
                this.handleVisitedNodes(current);
        } else {
                this.showFirstView(current);
        }
};




/* Install event handlers for all buttons under root element */
QController.prototype.installHandlers = function() {
        const controller = this;
        this.rootElement.querySelectorAll("button").forEach(button => {
                button.addEventListener("click", ev => {
                        const view = button.closest(".question-view");

      
      
      
      if (!view) {
                                console.log("qcontroller: <button> not inside question view");
                                return;
                        }

      
      
      
      const classList = button.classList;
                        let forward = null;

      
      
      
      if (classList.contains("next")) {
                                forward = view.nextElementSibling;
                        } else if (classList.contains("top-next")) {
                                const parent = view.parentNode;
                                if (parent) {
                                        forward = parent.nextElementSibling;
                                } else {
                                        forward = null;
                                }
                        } else if (classList.contains("top")) {
                                forward = containingView(view);
                        } else if (classList.contains("prev")) {
                                forward = view.previousElementSibling;
                        }

      
      
      
      if (forward && forward.classList.contains("question-view")) {
                                controller.visit(forward);
                                controller.saveState();
                                controller.update();
                        }
                });
        });
};




QController.prototype.hasVisited = function(viewIndex) {
        return this.visited[viewIndex];
};




QController.prototype.visit = function(nextView) {
        const index = Array.prototype.findIndex.call(
                this.views,
                x => x === nextView);
        this.currentView = index;
        this.visited[index] = true;
};




/* Save controller state to current session */
QController.prototype.saveState = function() {
        const key = this.id + "-controller-state";
        const state = {
                currentView: this.currentView,
                visited: this.visited
        };
        window.sessionStorage.setItem(key, JSON.stringify(state));
};




/* Load controller state from current session */
QController.prototype.loadState = function() {
        const key = this.id + "-controller-state";
        const state = JSON.parse(window.sessionStorage.getItem(key));

  
  
  
  if (state) {
                this.currentView = state.currentView;
                this.visited = state.visited;
        }
};




/* Show the first leaf question view */
QController.prototype.showFirstView = function(questionView) {
        let view = questionView;

  
  
  
  while (view && !isLeaf(view)) {
                view = view.querySelector(".question-view");
        }

  
  
  
  showView(view);
        this.handleVisitedNodes(view);
}




function showView(view) {
        /* show all containing views */
        while (view) {
                show(view);
                view = containingView(view);
        }
}




/* Return the view containing VIEW */
function containingView(view) {
        let parent = view.parentNode;

  
  
  
  if (parent) {
                parent = parent.closest(".question-view");
        }

  
  
  
  return parent;
}




function show(view) {
        view.style.display = "block";
}




function isLeaf(view) {
        return view.querySelector(".question-view") == null;
}




function isVisible(view) {
        return view.style.display != "none";
}




function hide(view) {
        view.style.display = "none";
}




}),
/* @license-end */
        </script>
    </div>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>10.0000000</defaultgrade>
        <penalty>0.0000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2020120600</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[/*
  Passagiere eines Transatlantikflugs

  wurde entwickelt von

    Daniel Mei&#223;ner <daniel.meissner-i4k(at)ruhr-uni-bochum.de>

  nach einer Idee von

    Herold Dehling <herold.dehling(at)ruhr-uni-bochum.de>

  an der Ruhr-Universit&#228;t Bochum.

  Dieses Werk ist lizenziert unter einer Creative Commons
  Namensnennung-Weitergabe unter gleichen Bedingungen 4.0 International
  Lizenz. Um eine Kopie der Lizenz zu erhalten, besuchen Sie
  http://creativecommons.org/licenses/by-sa/4.0/.

  SPDX-License-Identifier: CC-BY-SA-4.0

  Technische Informationen:

  Diese Aufgabe enth&#228;lt Hilfeschritte, die mithilfe von JavaScript
  nacheinander sichtbar gemacht werden.
*/

/* utilities */

probabilityp(p) := is(numberp(p) and p >= 0 and p <= 1);

/* actual question variables */

pp: 9/10;
nn: rand([400, 900]);

mean: nn * pp;
variance: nn * pp * (1 - pp);
std: sqrt(variance);

max_missing: mean + rand([1, 2])*std;

Tans1: random_permutation([
  [simplode(["Bin(", nn, ", ", pp, ")-Verteilung"]), true],
  [simplode(["Exp(", nn * pp, ")-Verteilung"]), false],
  [simplode(["Poisson(", nn * pp, ")-Verteilung"]), false],
  [simplode(["N(", nn, ", ", pp, ")-Verteilung"]), false]
]);
Tans1b: binomial(nn,k)*pp^k*(1-pp)^(nn-k);
Tans2: mean;
Tans3: variance;
Tans4: decimalplaces(cdf_normal(max_missing, mean, std), 2);]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text/>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{# nn #} {# max_missing #}</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Ihre Antwort ist nicht komplett falsch.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:orange;">
        <i class="fa fa-adjust">
        </i>
    </span>
    Your answer is not completely wrong.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed>i</variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>dropdown</type>
            <tans>Tans1</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>0</mustverify>
            <showvalidation>0</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans2</name>
            <type>numerical</type>
            <tans>decimalplaces(Tans2,2)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans3</name>
            <type>numerical</type>
            <tans>decimalplaces(Tans3,2)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>ans4</name>
            <type>numerical</type>
            <tans>ev(Tans4, float)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>hilfe3</name>
            <type>algebraic</type>
            <tans>(X - mean)/sqrt(variance)</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <input>
            <name>hilfe4</name>
            <type>algebraic</type>
            <tans><![CDATA[{ Z <= (max_missing - mean)/std }]]></tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>ans4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text>exakte_loesung: cdf_binomial(max_missing, nn, pp);
approx_loesung: cdf_normal(max_missing, mean, sqrt(variance));</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans4</sans>
                <tans>approx_loesung</tans>
                <testoptions>0.005</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty>0</truepenalty>
                <truenextnode>-1</truenextnode>
                <trueanswernote>ans4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Um eine gute N&#228;herung f&#252;r diese Wahrscheinlichkeit zu erhalten, kann die Normalapproximation verwendet werden. Dazu muss die Variable \(X\) standardisiert werden:
\[
Z = {@ (X - nn*pp)/sqrt(nn*pp*(1-pp)) @}.
\]
Dann gilt
\[
P(X \le {@ max_missing @}) = P\left(Z \le \frac{{@ max_missing @} - {@ nn*pp@}}{\sqrt{{@nn * pp * (1-pp)@}}}\right) 
= {@ P(Z &lt;= (max_missing - mean)/std) @}\approx {@ decimalplaces(approx_loesung, 2) @}.
\]
    </p>
    <p>
        Herzlichen Gl&#252;ckwunsch, Sie sind am Ende der Aufgabe angelangt.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In order to obtain a good approximation for this probability, normal approximation can be used. To this end, the variable \(X\) to be standardized:
\[
Z = {@ (X - nn*pp)/sqrt(nn*pp*(1-pp)) @}.
\]Then,
\[
P(X \le {@ max_missing @}) = P\left(Z \le \frac{{@ max_missing @} - {@ nn*pp@}}{\sqrt{{@nn * pp * (1-pp)@}}}\right) 
= {@ P(Z &lt;= (max_missing - mean)/std) @}\approx {@ decimalplaces(approx_loesung, 2) @}.
\]
    </p>
    <p>
        Congratulations, you have reached the end of the task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty>0</falsepenalty>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p class="visited" data-qcontroller-help="none">
        Dass Sie die Aufgabe nicht auf Anhieb l&#246;sen konnten, ist kein Problem! Im Folgenden wird die Aufgabe in kleinere Schritte unterteilt. Wenn Sie diese Schritte alle gel&#246;st haben, k&#246;nnen Sie die urspr&#252;ngliche Aufgabe sicherlich l&#246;sen. Klicken Sie dazu auf
        <button class="next" type="button">
            Weiter
        </button>
        .
    </p>
    <p>
        <span class="visited" data-qcontroller-help="1">
            Willkommen zur&#252;ck!  W&#228;hrend der Bearbeitung der Hilfe haben Sie erfahren, dass \(X\) eine {@first(mcq_correct(Tans1))@} hat.
        </span>
        <span class="visited" data-qcontroller-help="2">
            Sie haben den Erwartungswert {@E(X) = mean@} und die Varianz {@Var(X) = variance@} bestimmt.
        </span>
        <span class="visited" data-qcontroller-help="3">
            Dann haben Sie eine neue Zufallsvariable {@Z = (X-mean)/sqrt(variance)@} eingef&#252;hrt.
        </span>
        <span class="visited" data-qcontroller-help="4">
            Schlie&#223;lich haben Sie das Ereignis {@{X &lt;= max_missing}@} mithilfe der neuen Zufallsvariablen \(Z\) wie folgt geschrieben: {@{Z &lt;= (max_missing-mean)/sqrt(variance)}@}
        </span>
    </p>
    <p class="visited" data-qcontroller-help="partial">
        M&#246;chten Sie die Hilfe fortsetzen? Klicken Sie auf
        <button class="next" type="button">
            Hilfe fortsetzen
        </button>
        .
    </p>
</span>
<span class="multilang" lang="en">
    <p class="visited" data-qcontroller-help="none">
        That you couldn't solve the task right away is no problem! In the following, the object is subdivided into smaller steps. If you have all solved these steps, you can certainly solve the original task. Click on
        <button class="next" type="button">
            Next
        </button>
        .
    </p>
    <p>
        <span class="visited" data-qcontroller-help="1">
            Welcome back! During the processing of the help you have learned that \(X\) a {@first(mcq_correct(Tans1))@} has.
        </span>
        <span class="visited" data-qcontroller-help="2">
            You have the expected value {@E(X) = mean@} and the variance {@Var(X) = variance@} destined.
        </span>
        <span class="visited" data-qcontroller-help="3">
            Then you have a new random variable {@Z = (X-mean)/sqrt(variance)@} introduced.
        </span>
        <span class="visited" data-qcontroller-help="4">
            After all, you have the event {@{X &lt;= max_missing}@} using the new random variables \(Z\) as follows: {@{Z &lt;= (max_missing-mean)/sqrt(variance)}@}
        </span>
    </p>
    <p class="visited" data-qcontroller-help="partial">
        Would you like to continue the help? Click on
        <button class="next" type="button">
            Continue
        </button>
        .
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>2</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>probabilityp(ans4)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>0</truenextnode>
                <trueanswernote>ans4-3-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>ans4-3-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Wahrscheinlichkeiten sollten im Intervall [0, 1] liegen, das trifft f&#252;r Ihre Eingabe nicht zu.
    </p>
    <p>
        Bitte versuchen Sie es noch einmal: Geben Sie eine Wahrscheinlichkeit ein und dr&#252;cken Sie auf &#8222;Pr&#252;fen&#8220;.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Probabilities should be in the interval [0, 1], this does not apply to your input.
    </p>
    <p>
        Please try again: Enter a probability and press &#8220;Check&#8221;.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>hilfe3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text>loesung: (X - mean)/sqrt(variance)
varloesung: (X - mean)/variance</text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>hilfe3</sans>
                <tans>loesung</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>hilfe3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Richtige Antwort, gut gemacht!
    <p>
        Mit dem Zentralen Grenzwertsatz folgt, dass die standardisierte Variable \(Z = {@ (X - nn*pp)/sqrt(nn*pp*(1-pp)) @}\) ann&#228;hernd standard-normalverteilt ist.
    </p>
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        , um mit der Hilfe fortzufahren,
oder auf
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        ,
um zur urspr&#252;nglichen Aufgabenstellung zur&#252;ckzukehren.
    </p>
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:green;">
        <i class="fa fa-check">
        </i>
    </span>
    Right answer, well done!
    <p>
        The central limit value set follows that the standardised variable \(Z = {@ (X - nn*pp)/sqrt(nn*pp*(1-pp)) @}\) is approximately standard-normally distributed.
    </p>
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        to continue with the help,
or
        <button class="top" type="button">
            Back
        </button>
        ,
to return to the original task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>hilfe3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    <span style="font-size: 1.5em; color:red;">
        <i class="fa fa-times">
        </i>
    </span>
    Wrong answer.
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>hilfe3</sans>
                <tans>varloesung</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>hilfe3-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Haben Sie m&#246;glicherweise Varianz und Standardabweichung verwechselt? Standardisieren Sie die Variable \(X\), indem Sie den Erwartungswert subtrahieren und durch die Standardabweichung teilen, um eine approximativ standard-normalverteilte Variable zu erhalten.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Have you possibly confused variance and standard deviation? Standardize the variable \(X\)by subtracting the expected value and by sharing the standard deviation to obtain an approximate standard-normally distributed variable.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>hilfe3-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Standardisieren Sie die Variable \(X\), indem Sie den Erwartungswert subtrahieren und durch die Standardabweichung teilen, um eine approximativ standard-normalverteilte Variable zu erhalten.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Standardize the variable \(X\)by subtracting the expected value and by sharing the standard deviation to obtain an approximate standard-normally distributed variable.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>hilfe4</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[loesung: { Z <= (max_missing - mean)/std }
varloesung: { Z <= (max_missing - mean)/variance }]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>hilfe4</sans>
                <tans>loesung</tans>
                <testoptions/>
                <quiet>1</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>hilfe4-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Die folgenden Umformungsschritte zeigen, dass \({@{X &lt;= max_missing} @} = {@ loesung @}\) gilt:
\[
{@ X &lt;= max_missing@} \iff {@ X - mean &lt;= max_missing -mean @}
\iff {@ (X - mean)/std &lt;= (max_missing-mean)/std @}
\iff {@ Z &lt;=(max_missing-mean)/std@}
\]
    </p>
    <p>
        Klicken Sie auf
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        ,
um zur urspr&#252;nglichen Aufgabenstellung zur&#252;ckzukehren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        The following forming steps show that \({@{X &lt;= max_missing} @} = {@ loesung @}\) the following:
\[
{@ X &lt;= max_missing@} \iff {@ X - mean &lt;= max_missing -mean @}
\iff {@ (X - mean)/std &lt;= (max_missing-mean)/std @}
\iff {@ Z &lt;=(max_missing-mean)/std@}
\]
    </p>
    <p>
        Click on
        <button class="top" type="button">
            Back
        </button>
        ,
to return to the original task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>hilfe4-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>hilfe4</sans>
                <tans>varloesung</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>hilfe4-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Haben Sie m&#246;glicherweise Varianz und Standardabweichung verwechselt? Im letzten Schritt haben Sie den folgende Zusammenhang zwischen \(X\) und der Standardisierung \(Z\) gefunden:
\[
Z = {@ (X - mean)/std @}
\]
Nutzen Sie diesen Zusammenhang, um die Ungleichung {@X &lt;= max_missing @} in eine Ungleichung mit \(Z\) umzuformen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Have you possibly confused variance and standard deviation? In the last step, you have the following link between \(X\) and standardization \(Z\) found:
\[
Z = {@ (X - mean)/std @}
\]Take advantage of this connection to the unequalification {@X &lt;= max_missing @} in an inequality with \(Z\) to reform.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>hilfe4-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Im letzten Schritt haben Sie den folgende Zusammenhang zwischen \(X\) und der Standardisierung \(Z\) gefunden:
\[
Z = {@ (X - mean)/std @}
\]
Nutzen Sie diesen Zusammenhang, um die Ungleichung {@X &lt;= max_missing @} in eine Ungleichung mit \(Z\) umzuformen.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        In the last step, you have the following link between \(X\) and standardization \(Z\) found:
\[
Z = {@ (X - mean)/std @}
\]Take advantage of this connection to the unequalification {@X &lt;= max_missing @} in an inequality with \(Z\) to reform.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>first (mcq_correct (Tans1))</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Ein Passagier kann entweder erscheinen (Erfolg) oder nicht erscheinen (Misserfolg). Die Zufallsvariable \(X\) z&#228;hlt daher die Anzahl der Erfolge in {#nn#} unabh&#228;ngigen Bernoulli-Experimenten und ist somit binomialverteilt mit den Parametern \(n = {@ nn @}\) und \(p = {@ pp @}\).
    </p>
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        , um mit der Hilfe fortzufahren,
oder auf
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        ,
um zur urspr&#252;nglichen Aufgabenstellung zur&#252;ckzukehren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        A passenger can either appear (following) or not appear (success). The random variable \(X\) the number of successes in {#nn#} independent Bernoulli experiments and is thus binomially distributed with the parameters \(n = {@ nn @}\) and \(p = {@ pp @}\).
    </p>
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        to continue with the help,
or
        <button class="top" type="button">
            Back
        </button>
        ,
to return to the original task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Ein Passagier kann entweder erscheinen (Erfolg) oder nicht erscheinen (Misserfolg). Die Zufallsvariable \(X\) z&#228;hlt daher die Anzahl der Erfolge in {#nn#} unabh&#228;ngigen Bernoulli-Experimenten. Bitte versuchen Sie es noch einmal.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        A passenger can either appear (following) or not appear (success). The random variable \(X\) the number of successes in {#nn#} independent Bernoulli experiments. Please try again.
    </p>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt2</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans2</sans>
                <tans>Tans2</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt2-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Der Erwartungswert einer binomialverteilten Zufallsvariablen \(X\) mit Parametern \(n = {@ nn @}\) und \(p = {@ pp @}\) ist gegeben durch die Formel
\[
\mathbb{E}[X] = np = {@ nn * pp @}.
\]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        The expected value of a binary-distributed random variable \(X\) with parameters \(n = {@ nn @}\) and \(p = {@ pp @}\) is given by the formula
\[
\mathbb{E}[X] = np = {@ nn * pp @}.
\]
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt2-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Der Erwartungswert einer binomialverteilten Zufallsvariablen \(X\) mit Parametern \(n = {@ nn @}\) und \(p = {@ pp @}\) ist mit einer kurze Formel berechenbar, in der \(n\) und \(p\) vorkommen.
    </p>
    <details>
        <summary>
            <em>
                Hinweis
            </em>
        </summary>
        \[
\mathbb{E}[X] = np.
\]
    </details>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The expected value of a binary-distributed random variable \(X\) with parameters \(n = {@ nn @}\) and \(p = {@ pp @}\) can be calculated with a short formula in which \(n\) and \(p\) to be found.
    </p>
    <details>
        <summary>
            <em>
                Note
            </em>
        </summary>
        \[
\mathbb{E}[X] = np.
\]
    </details>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>prt3</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans3</sans>
                <tans>Tans3</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt3-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Richtige Antwort, gut gemacht!
    </p>
    <p>
        Die Varianz einer binomialverteilten Zufallsvariablen \(X\) mit Parametern \(n = {@ nn @}\) und \(p = {@ pp @}\) ist gegeben durch die Formel
\[
\mathrm{Var}[X] = np(1-p) = {@ nn * pp * (1 - pp) @}.
\]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:green;">
            <i class="fa fa-check">
            </i>
        </span>
        Right answer, well done!
    </p>
    <p>
        The variance of a binomial distributed random variables \(X\) with parameters \(n = {@ nn @}\) and \(p = {@ pp @}\) is given by the formula
\[
\mathrm{Var}[X] = np(1-p) = {@ nn * pp * (1 - pp) @}.
\]
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt3-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Falsche Antwort.
    </p>
    <p>
        Die Varianz einer binomialverteilten Zufallsvariablen \(X\) mit Parametern \(n = {@ nn @}\) und \(p = {@ pp @}\)  ist durch eine kurze Formel berechenbar, in der \(n\) und \(p\) vorkommen.
    </p>
    <details>
        <summary>
            <em>
                Hinweis
            </em>
        </summary>
        \[
\mathrm{Var}[X] = np(1-p).
\]
    </details>
</span>
<span class="multilang" lang="en">
    <p>
        <span style="font-size: 1.5em; color:red;">
            <i class="fa fa-times">
            </i>
        </span>
        Wrong answer.
    </p>
    <p>
        The variance of a binomial distributed random variables \(X\) with parameters \(n = {@ nn @}\) and \(p = {@ pp @}\) can be calculated by a short formula in which \(n\) and \(p\) to be found.
    </p>
    <details>
        <summary>
            <em>
                Note
            </em>
        </summary>
        \[
\mathrm{Var}[X] = np(1-p).
\]
    </details>
</span>
]]></text>
                </falsefeedback>
            </node>
        </prt>
        <prt>
            <name>teil2_vollstaendig</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>0</feedbackstyle>
            <feedbackvariables>
                <text/>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans2</sans>
                <tans>Tans2</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>1</truenextnode>
                <trueanswernote>teil2_vollstaendig-1-T</trueanswernote>
                <truefeedback format="html">
                    <text/>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>teil2_vollstaendig-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>NumAbsolute</answertest>
                <sans>ans3</sans>
                <tans>Tans3</tans>
                <testoptions>0.01</testoptions>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>teil2_vollstaendig-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Klicken Sie auf
        <button class="next" type="button">
            Weiter
        </button>
        , um mit der Hilfe fortzufahren,
oder auf
        <button class="top" type="button">
            Zur&#252;ck
        </button>
        ,
um zur urspr&#252;nglichen Aufgabenstellung zur&#252;ckzukehren.
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        Click on
        <button class="next" type="button">
            Next
        </button>
        to continue with the help,
or
        <button class="top" type="button">
            Back
        </button>
        ,
to return to the original task.
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>teil2_vollstaendig-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
        <deployedseed>2045828694</deployedseed>
        <deployedseed>2139405385</deployedseed>
        <deployedseed>2133928116</deployedseed>
        <deployedseed>1696743305</deployedseed>
        <deployedseed>1234969437</deployedseed>
        <deployedseed>338725501</deployedseed>
        <deployedseed>1501523990</deployedseed>
        <deployedseed>1695580723</deployedseed>
        <deployedseed>662877477</deployedseed>
        <deployedseed>985397353</deployedseed>
        <deployedseed>1415345076</deployedseed>
        <deployedseed>633060561</deployedseed>
        <deployedseed>497251047</deployedseed>
        <deployedseed>232485328</deployedseed>
        <deployedseed>318852464</deployedseed>
        <qtest>
            <testcase>1</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value>first(mcq_correct(Tans1))</value>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>decimalplaces(Tans2,2)</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>decimalplaces(Tans3,2)</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>decimalplaces(Tans4, 2)</value>
            </testinput>
            <expected>
                <name>ans4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt1</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt1-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-T</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>2</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value/>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value>2</value>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value>4</value>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value>cdf_normal (max_missing, 2, 2)</value>
            </testinput>
            <expected>
                <name>ans4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans4-2-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt1</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt2-1-F</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore>0.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>prt3-1-F</expectedanswernote>
            </expected>
        </qtest>
        <qtest>
            <testcase>3</testcase>
            <description/>
            <testinput>
                <name>ans1</name>
                <value/>
            </testinput>
            <testinput>
                <name>ans2</name>
                <value/>
            </testinput>
            <testinput>
                <name>ans3</name>
                <value/>
            </testinput>
            <testinput>
                <name>ans4</name>
                <value/>
            </testinput>
            <expected>
                <name>ans4</name>
                <expectedscore>1.0000000</expectedscore>
                <expectedpenalty>0.0000000</expectedpenalty>
                <expectedanswernote>ans4-1-T</expectedanswernote>
            </expected>
            <expected>
                <name>prt1</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>prt2</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
            <expected>
                <name>prt3</name>
                <expectedscore/>
                <expectedpenalty/>
                <expectedanswernote>NULL</expectedanswernote>
            </expected>
        </qtest>
    </question>
    

</quiz>