<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
        <name>
            <text>Vier Urnenmodelle</text>
        </name>
        <questiontext format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    <p>
    </p>
    <p>
        Wie viele M&#246;glichkeiten gibt es, {@k1@} Kugeln aus einer Urne mit insgesamt {@n1@} Kugeln zu ziehen ({@ZL[2]@} Zur&#252;cklegen der Kugeln und {@RF[2]@} Beachtung der Reihenfolge)?
        <br/>
    </p>
    <p>
        [[input:ans1]] [[validation:ans1]]
    </p>
</span>
<span class="multilang" lang="en">
    <p>
    </p>
    <p>
        How many possibilities are there, {@k1@} Balls of a urn with a total {@n1@} To pull balls ({@ZL[2]@} laying down the balls and {@RF[2]@} Observing the order)?
        <br/>
    </p>
    <p>
        [[input:ans1]] [[validation:ans1]]
    </p>
</span>
]]></text>
        </questiontext>
        <generalfeedback format="html">
            <text/>
        </generalfeedback>
        <defaultgrade>1.0000000</defaultgrade>
        <penalty>0.1000000</penalty>
        <hidden>0</hidden>
        <idnumber/>
        <stackversion>
            <text>2019090200</text>
        </stackversion>
        <questionvariables>
            <text><![CDATA[

tr(de,en):=castext("[[lang code='de']]{@de@}[[/lang]][[lang code='en']]{@en@}[[/lang]]");
ZL: rand([[1,tr("mit","with")],[2,tr("ohne","without")]]);
RF: rand([[1,tr("mit","with")],[2,tr("ohne","without")]]);
fall: [ZL[1],RF[1]];

n1: rand_with_step(5,10,1);
k1: rand_with_step(2,n1-2,1);

TA11: n1^k1; /*mit Zur&#252;cklegen, mit Reihenfolge*/
TA12: binomial(n1+k1-1,k1); /*mit Zur&#252;cklegen, ohne Reihenfolge*/
TA21: binomial(n1,k1)*k1!; /*ohne Zur&#252;cklegen, mit Reihenfolge*/
TA22: binomial(n1,k1); /*ohne Zur&#252;cklegen, ohne Reihenfolge*/
Mat: matrix([TA11,TA12],[TA21,TA22]);

TA: Mat[ZL[1],RF[1]];]]></text>
        </questionvariables>
        <specificfeedback format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    [[feedback:prt1]]
</span>
<span class="multilang" lang="en">
    [[feedback:prt1]]
</span>
]]></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text>{@TA@}</text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text/>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Richtige Antwort, gut gemacht!
</span>
<span class="multilang" lang="en">
    Right answer, well done!
</span>
]]></text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Ihre Antwort ist teilweise korrekt.
</span>
<span class="multilang" lang="en">
    Your answer is partly correct.
</span>
]]></text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text><![CDATA[
<span class="multilang" lang="de">
    Falsche Antwort.
</span>
<span class="multilang" lang="en">
    Wrong answer.
</span>
]]></text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed/>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>TA</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint/>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords/>
            <allowwords/>
            <forbidfloat>0</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options/>
        </input>
        <prt>
            <name>prt1</name>
            <value>1.0000000</value>
            <autosimplify>1</autosimplify>
            <feedbackstyle>1</feedbackstyle>
            <feedbackvariables>
                <text><![CDATA[textZL: if ans1=Mat[1,1] or ans1=Mat[1,2] then "mit" else "ohne";
textRF: if ans1=Mat[1,1] or ans1=Mat[2,1] then "mit" else "ohne";
TAM: set(TA11,TA12,TA21,TA22);]]></text>
            </feedbackvariables>
            <node>
                <name>0</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>ans1</sans>
                <tans>TA</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>=</truescoremode>
                <truescore>1.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-1-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    <p>
        Sie haben die korrekte Anzahl an M&#246;glichkeiten angegeben.
        <br/>
    </p>
</span>
<span class="multilang" lang="en">
    <p>
        You have specified the correct number of options.
        <br/>
    </p>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>=</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>1</falsenextnode>
                <falseanswernote>prt1-1-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
            <node>
                <name>1</name>
                <description/>
                <answertest>AlgEquiv</answertest>
                <sans>member(ans1,TAM)</sans>
                <tans>true</tans>
                <testoptions/>
                <quiet>0</quiet>
                <truescoremode>+</truescoremode>
                <truescore>0.0000000</truescore>
                <truepenalty/>
                <truenextnode>-1</truenextnode>
                <trueanswernote>prt1-2-T</trueanswernote>
                <truefeedback format="html">
                    <text><![CDATA[
<span class="multilang" lang="de">
    Ihre L&#246;sung w&#228;re korrekt, wenn wir den Fall {@textZL@} Zur&#252;cklegen der Kugeln und {@textRF@} Beachtung der Reihenfolge betrachten w&#252;rden. Wir betrachten aber den Fall {@ZL[2]@} Zur&#252;cklegen der Kugeln und {@RF[2]@} Beachtung der Reihenfolge.
    <br/>
</span>
<span class="multilang" lang="en">
    Your solution would be correct if we had the case {@textZL@} laying down the balls and {@textRF@} I would like to see the order. We consider the case {@ZL[2]@} laying down the balls and {@RF[2]@} Observing the order.
    <br/>
</span>
]]></text>
                </truefeedback>
                <falsescoremode>-</falsescoremode>
                <falsescore>0.0000000</falsescore>
                <falsepenalty/>
                <falsenextnode>-1</falsenextnode>
                <falseanswernote>prt1-2-F</falseanswernote>
                <falsefeedback format="html">
                    <text/>
                </falsefeedback>
            </node>
        </prt>
    </question>
    

</quiz>